```python
# Initialize Otter
import otter
grader = otter.Notebook("test.ipynb")
```

```python
import pandas as pd
import numpy as np
import matplotlib.pylab as plt
%matplotlib inline
from datascience import *
path="https://raw.githubusercontent.com/fundatosudea/material_20242/develop_otter"
```

<!-- BEGIN QUESTION -->

Ingresa tu nombre y número de carnet:

```python
# BEGIN SOLUTION NO PROMPT
Nomre="Ingrese"
NumCarnet="1111"
# END SOLUTION
""" # BEGIN PROMPT
Nombre="..."
NumCarnet=".."
""" # END PROMPT
```

Output:
```
' # BEGIN PROMPT\nNombre="..."\nNumCarnet=".."\n'
```

```python
grader.check("qt")
```

<!-- END QUESTION -->

P1. Esta es una pregunta de pruena


```python
# BEGIN SOLUTION NO PROMPT
characters_q1 = 5
# END SOLUTION
""" # BEGIN PROMPT
characters_q1 = ...
""" # END PROMPT
```

Output:
```
' # BEGIN PROMPT\ncharacters_q1 = ...\n'
```

```python
grader.check("q1")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: 1)
```
>>> assert characters_q1 == 5
```


Tests for qt (points: 0)
```
>>> assert isinstance(Nombre, str)
>>> assert isinstance(NumCarnet, str)
>>> assert len(Nombre) > 3
>>> assert len(NumCarnet) > 3
```
