```python
# Initialize Otter
import otter
grader = otter.Notebook("hw02.ipynb")
```

# Tarea 2: Arreglos y Tablas
Por favor complete este notebook llenando las celdas que se piden. Antes de comenzar, ejecute la celda previa para cargar las pruebas dadas.



```python
# Ejecute esta celda e ignore lo que esta arroja.Esto es solo para verificar que ejecutó la primera celda.
# Si obtuene un NameError, ¡por favor corra la primera celca de este notebook!

grader
```

**Recursos útiles:**
- [Referencia de Python](http://data8.org/fa23/reference/): Hoja de referencia de métodos útiles de matrices y tablas utilizados en Data 8.

**Lecturas recomendadas:**
- [Arreglos](https://inferentialthinking.com/chapters/05/1/Arrays.html)
- [Guía de arreglos de Data 8](https://drive.google.com/file/d/1vTduvivHyEX7cdpMCq9B2jC4iy6A3Tru/view?usp=sharing)
- [¿Qué es la ciencia de datos?](http://www.inferentialthinking.com/chapters/01/what-is-data-science.html)
- [Causalidad y experimentos](http://www.inferentialthinking.com/chapters/02/causality-and-experiments.html)
- [Programando en Python](http://www.inferentialthinking.com/chapters/03/programming-in-python.html)

Para todos los problemas en los que deba escribir explicaciones y oraciones, **deberá** proporcionar la respuesta en el espacio designado. Además, a lo largo de esta tarea y de todas las futuras, ¡asegúrese de no reasignar variables en todo el notebook! Por ejemplo, si usa `temperatura` en su respuesta a una pregunta, no la reasigne más adelante. De lo contrario, ¡reprobará las pruebas que creía que estaba aprobando anteriormente!

**Fecha de entrega:**

Esta tarea deberá ser entregada ** ___ a las ___**. Si quiere puntos extra, entréguela el (un día antes). Trabajo que se entregue por fuera del plazo no será recibido.

**Nota: Esta tarea tiene pruebas ocultas. Eso significa que aunque los exámenes indiquen que se aprobó al 100%, no significa que su calificación final será del 100%. Realizaremos más pruebas ocultas de corrección una vez que todos entreguen la tarea.**


Evite compartir respuestas o hacer copia con sus compañeros. Establezca en cambio discusiones grupales con otros estudiantes y docentes del curso.

Si desea contar con el tiempo necesario para pedir ayuda o asesoría en caso de estar atascado, se recomienda iniciar temprano con la realización de las tareas. Tenga en cuenta el cronograma del curso y de entregas.

## 1. Creando Arreglos

```python
# Ejecute esta celda para configurar su notebook, pero por favor no la cambie.

import numpy as np
from datascience import *
import warnings
warnings.simplefilter('ignore', FutureWarning)
```

**Pregunta 1.** Cree un arreglo llamado `numeros_extraños` que contenga los siguientes números (en el orden dado) **(4 Points)**:

1. -2
2. El piso (floor) de 17.6
3. 3
4. 5 a la potencia del techo (ceil) de 5.3

*Hint:* `floor` y `ceil` son funciones en el módulo `math`. Importar módulos se trabaja en el Lab 2.

*Nota:* Las listas de Python son y se comportan diferente de los arreglos de NumPy. En este curso usaremos arreglos de NumPy, así que cuando haga un **arreglo**, no haga una lista de Python.


```python
# Nuestra solución requiere una línea extra de código antes de crear numeros_extraños

...
numeros_extraños = ...
numeros_extraños
```

Output:
```
Ellipsis
```

```python
grader.check("q1_1")
```

**Pregunta 2.** Cree un arreglo llamado `titulo_libros` que contenga las tres siguientes cadenas de caracteres: "Eats", "Shoots", and "and Leaves". **(4 Puntos)**


```python
titulo_libros = ...
titulo_libros
```

```python
grader.check("q1_2")
```

Las cadenas de caracteres (strings) tienen un método llamado `join`. `join` toma como argumento un arreglo de cadenas de caracteres y regresa una cadena de caracteres. Específicamente, el resultado es una sola cadena de caracteres, que es la [concatenación](https://es.wikipedia.org/wiki/Concatenaci%C3%B3n) ("juntar") todas las cadenas de caracteres en el arreglo.


**Pregunta 3.** Use el arreglo `titulo_libros` y el método `join` para crear dos arreglos **(4 untos)**:

1. "Eats, Shoots, and Leaves" (llámelo `usando_comas`)
2. "Eats Shoots and Leaves" (llámelo `sin_comas`)



```python
usando_comas = ...
sin_comas = ...

# Estas líneas son solo para imprimir sus respuestas.
print('usando_comas:', usando_comas)
print('sin_comas:', sin_comas)
```

```python
grader.check("q1_3")
```

## 2. Indexando Arreglos

Estos ejercicios están pensados en dar práctica al momento de acceder a elementos individuales de los arreglos. En Python (y en muchos lenguajes de programación), cada elemento es accesible mediante su  *índice*; por ejemplo, el primer elmento es el elemento con índice 0. Los índices deben ser números **enteros**.


***Nota:* Si usted tiene experiencia previa programando, la notación de corchetes le debe ser familiar. NO use la notación de corchetes cuando esté indexando (i.e. `arr[0]`), pues esto puede conducir a un tipo de dato diferente del que estamos esperando. Le puede además hacer fallar en la revisión.**

En caso de sentirse atascado, revise la [Referencia de Python](http://data8.org/fa23/reference/)



**Pregunta 1.** La siguiente celda crea un arreglo de algunos números. Asigne `tercer_elemento`  al tercer elemento de `algunos_numeros`. **(4 Puntos)**


```python
algunos_numeros = make_array(-1, -3, -6, -10, -15)

tercer_elemento = ...
tercer_elemento
```

```python
grader.check("q2_1")
```

**Pregunta 2.** La próxima celda crea una tabla que muestra información sobre los elementos de `algunos_numeros` y su orden. Corra la celda para ver la tabla parcialmente completa y llene la información faltante (las celdas que dicen "Ellipsis") mediante la asignación `espacio_a`, `espacio_b`, `espacio_c`, y `espacio_d` para corregir los elementos en la tabla. **(4 Puntos)**

*Hint* Reemplace los `...` con cadenas de caracteres o números. Recuerde que los índices deben ser **enteros**.


```python
espacio_a = ...
espacio_b = ...
espacio_c = ...
espacio_d = ...
elementos_de_algunos_numeros = Table().with_columns(
    "English name for position", make_array("primero", "segundo", espacio_a, espacio_b, "quinto"),
    "Index",                     make_array(espacio_c, 1, 2, espacio_d, 4),
    "Element",                   algunos_numeros)
elementos_de_algunos_numeros
```

```python
grader.check("q2_2")
```

**Pregunta 3.** En ocasiones usted querrá encontrar el **último** elemento de un arreglo. Suponga que el arreglo tiene 142 elementos. ¿Cuál es el índice del último elemento? **(4 Puntos)**

*Nota:* Su respuesta deberá ser un número positivo.


```python
indice_ultimo_elemento = ...
```

```python
grader.check("q2_3")
```

Con frecuencia, usted no conocerá el número de elementos de un arreglo, su *longitud*. (Por ejemplo, puede ser un conjunto de datos que usted encuentra en Internet.) La función `len` toma un solo argumento, un arreglo, y entrega un número entero que representa la longitud o tamaño de dicho arreglo.

**Pregunta 4.** La siguiente celda carga un arreglo llamado `año_nacimiento_presidentes`. Llamando `tabla.column(...)` en una tabla entrega un arreglo con la columna especificada, en este caso la columna `Año de nacimiento` de la tabla `nacimiento_presidentes`. El último elemento en ese arreglo es el más reciente entre los años de nacimiento entre todos los ex-presidentes de los Estados Unidos. Asigne ese año a `año_nacimiento_mas_reciente` **(4 Puntos)**


**Nota:** Evite googlear la respuesta. Usted debería ser capaz de responder esta pregunta únicamente usando métodos de tablas.


```python
año_nacimiento_presidentes = Table.read_table("president_births.csv").column('Birth Year')  #Modificar una vez se ajuste la tabla.

most_recent_birth_year = ...
most_recent_birth_year
```

```python
grader.check("q2_4")
```

**Pregunta 5.** Para finalizar, asigne `min_año_nacimiento` al mínimo del primer, decimosexto y el último año listados en `año_nacimiento_presidentes`. **(4 Puntos)**

**Nota:** Use la función `min` de Python y métodos de tabla para encontrar la respuesta. Evite calcular manualmente el resultado.



```python
min_año_nacimiento = ...
min_año_nacimiento
```

```python
grader.check("q2_5")
```

## 3. Aritmética Básica de Arreglos

**Pregunta 1.** Multiplique los números 42, -4224, 424224242, y 250 por 157. Asigne cada variable en la siguiente celda de tal forma que  `primer_producto` se asigne al resultado de $42 * 157$, `segundo_producto` al resultado de $-4224 * 157$ y así sucesivamente. **(4 Puntos)**

*Nota*: Para esta pregunta, **no** use arreglos.


```python
primer_producto = ...
segundo_producto = ...
tercer_producto = ...
cuarto_producto = ...
print("Primer Producto:", primer_producto)
print("Segundo Producto:", segundo_producto)
print("Tercer Producto:", tercer_producto)
print("Cuarto Producto:", cuarto_producto)
```

```python
grader.check("q3_1")
```

**Pregunta 2.** Ahora, haga los mismos cálculos, pero esta vez usando un arreglo llamado `numeros`y un solo operador de multiplicación (`*`). Almacene los 4 resultados en un arreglo llamado `productos`. **(4 Puntos)**

```python
numeros = ...
productos = ...
productos
```

```python
grader.check("q3_2")
```

**Pregunta 3.** Ups, cometimos un error. En lugar de 157, queremos que multiplique cada número por 1577. Calcule los productos correctos en la celda de abajo usando aritmética de arreglos. Note que su trabajo es realmente sencillo si previamente definió un arreglo que contiene los 4 números. **(4 Puntos)**

```python
productos_correctos = ...
productos_correctos
```

```python
grader.check("q3_3")
```

**Pregunta 4.** Hemos cargado un arreglo de temperaturas en la siguiente celda. Cada número es la temperatura más alta registrada en un día en una estación de medición del cloma, mayoritariamente en los Estados Unidos. Al ser datos de la agencia gubernamental de Estados Unidos [NOAA](https://www.noaa.gov/), todas las temperaturas están en Fahrenheit.

Convierta todas las temperaturas a Celsius, primero sustrayendo 32 de ellas. Luego, multiplique los resultados por by $\frac{5}{9}$, i.e. $C = (F - 32) * \frac{5}{9}$. Luego de convertir las temperaturas a Celsius, asegúrese de **REDONDEAR** el resultado final al entero más cercano usando la función `np.round`. **(4 Puntos)**

```python
max_temps = Table.read_table("temperatures.csv").column("Daily Max Temperature") #Modificar una vez se ajuste la tabla.


celsius_max_temps = ...
celsius_temps_redondeadas = ...
celsius_temps_redondeadas
```

```python
grader.check("q3_4")
```

**Pregunta 5.** La celda de abajo carga todas las más *bajas* temperaturas de cada día (en Fahrenheit). Calcule el intervalo diario de temperaturas de cada día. Esto es, calcule la diferencia, por día, entre las respectivas temperaturas máxima y mínima.**¡Preste atención a las unidades y de su respuesta en Celsius!** Asegúrese de **NO** redondear su respuesta.


*Hint:* Use `min_temps` y/o `max_temps`, y preste atención a la conversión de unidades. Escriba el cálculo matemático a mano si no logra avanzar.



```python
min_temps = Table.read_table("temperatures.csv").column("Daily Min Temperature")

celsius_temps_rangos = ...
celsius_temps_rangos
```

```python
grader.check("q3_5")
```

## 4. Old Faithful

*Old* Faithful es un geyser en Yellowstone que erupciona cada 44-125 minutos (de acuerdo con [Wikipedia](https://es.wikipedia.org/wiki/Old_Faithful#:~:text=El%20Old%20Faithful%2C%20traducido%20al,%2C%20en%20Wyoming%2C%20Estados%20Unidos)). A las personas se les dice con frecuencia que erupciona cada una hora, pero de hecho el tiempo de espera entre erupciones es más variable. Veámoslo.



**Pregunta 1.** La primera línea de la celda de abajo asigna `tiempos_espera` a un arreglo con 272 tiempos de espera consecutivos entre erupciones, tomados de un conjunto de datos de 1938. Asigne los nombres `menor`, `mayor` y `promedio`de forma tal que la asignación `print` esté correcta. **(4 Puntos)**



```python
tiempos_espera = Table.read_table('old_faithful.csv').column('waiting')

menor = ...
mayor = ...
promedio = ...

print("Old Faithful erupciona cada", menor, "a", mayor, "minutos y cada", promedio, "minutos en promedio.")
```

```python
grader.check("q4_1")
```

**Pregunta 2.** Asigne `mayor_diferencia` al mayor cambio en la espera entre dos erupciones consecutivas. Por ejemplo, la tercera erupción ocurrió luego de 74 minutos y la cuarta luego de 62 minutos, así que la diferencia fue 74 - 62 = 12 minutos. **(4 Puntos)**

*Hint*: Queremos el valor absoluto del valor de la mayor diferencia.

*Nota*: `np.diff()` calcula la diferencia entre valores subsecuentes en un arreglo. Por ejemplo, llamando `np.diff()` en el arreglo `make_array(1, 8, 3, 5)`, evalúa `array([8 - 1, 3 - 8, 5 - 3])` o `array([7, -5, 2])`.



```python
# np.diff() calcula la diferencia entre valores subsecuentes
# en un arreglo de NumPy.
diferencias = np.diff(tiempos_espera)
mayor_diferencia = ...
mayor_diferencia
```

```python
grader.check("q4_2")
```

**Pregunta 3.** La tabla `faithful_erupcion_nums` contiene dos columnas: `numero_erupcion`, la cual representa el número de esa erupción, y `espera`, que indica el tiempo que se tuvo que esperar luego de esa erupción. Por ejemplo, tome las primeras dos filas de la tabla:

| Número de erupción | Espera  |
|--------------------|---------|
| 1                  | 79      |
| 2                  | 54      |

Podemos leer esto como sigue: luego de la primera erupción, esperamos 79 minutos hasta que ocurrió la segunda erupción. Así, luego de la segunda erupción, esperamos 54 minutos por la tercera erupción.

Suponga que Oscar y Wendy comenzaron viendo al Old Faithful al principio de la primera erupción. Suponga que ellos observan hasta el final de la décima erupción.

We can read this as follows: after the first eruption, we waited 79 minutes for the second eruption. Then, after the second eruption, we waited 54 minutes for the third eruption. Durante parte de ese tiempo estarán observando erupciones, y por el resto del tiempo estarán esperando que Old Faithful haga erupción. ¿Cuántos minutos pasarán esperando que ocurran erupciones? **(4 Puntos)**

*Hint #1:* Puede iniciar usando el método `take` en la tabla `faithful_erupcion_nums`.

*Hint #2:* `tiempo_espera_primeras_nueve` deberá ser un arreglo.


```python
# Las siguientes dos líneas de código cargan la tabla faithful_erupcion_nums
faithful = Table.read_table('old_faithful.csv').drop("eruptions")
faithful_erupcion_numss = faithful.with_column("eruption number", np.arange(faithful.num_rows) + 1).select(1, 0)

tiempo_espera_primeras_nueve = ...
tiempo_total_espera_decima = ...
tiempo_total_espera_decima
```

```python
grader.check("q4_3")
```

**Pregunta 4.** Imaginemos que su estimación del siguiente tiempo de espera siempre fue la duración del tiempo de espera anterior. Si siempre adivinara el tiempo de espera anterior, ¿qué tan grande sería, en promedio, el error al adivinar los tiempos de espera?

Por ejemplo, dado que los primeros cuatro tiempos de espera son 79, 54, 74 y 62, la diferencia promedio entre su "predicción" y el tiempo real solo para la segunda, tercera y cuarta erupciones sería: $\frac{|79-54|+ |54-74|+ |74-62|}{3} = 19$.


```python
diferencias = np.diff(tiempos_espera)
error_promedio = ...
error_promedio
```

```python
grader.check("q4_4")
```

## 5. Tablas

**Pregunta 1.** Suponga que tiene 4 manzanas y 3 piñas. (Quizás esté usando Python para resolver un problema de álgebra del bachillerato.)Cree una tabla que contenga esta información. Esta deberá contener dos columnas: `nombre de la fruta` y `cantidad`. Asigne la nueva tabla a la variable `frutas`. **(4 Puntos)**



*Note:* Use letras minúsculas y una sola palabra para el nombre de cada fruta, por ejemplo `"manzana"`.


```python
# Nuestra solución usa una sola sentencia divida en 3 líneas.
frutas = ...
        ...
        ...
frutas
```

```python
grader.check("q5_1")
```

**Question 2.** El archivo `inventory.csv` contiene información sobr eel inventario de un puesto de frutas. Cada fila representa el contenido de una caja de fruta. Cargue el archivo como una tabla llamada `inventario` usando la función `Table.read_table()`. `Table.read_table(...)` toma un argumento (el nombre de un archivo en formato cadena de caracteres) y regresa una tabla. **(4 Puntos)**


```python
inventario = ...
inventario
```

```python
grader.check("q5_2")
```

**Pregunta3.** ¿Cada caja en el puesto  si cada de frutas contiene una fruta diferente? Asigne `todas_diferentes` a `True` si cada caja contiene contiene una fruta diferente y `False` si múltiples cajas contienen la misma fruta. **(4 Puntos)**

*Hint:* No tiene que escribir un código que calcule el valor True/False para `todas_diferentes`. Solo mire la tabla y asigne el valor `True` o `False` de acuerdo a lo que ve en la tabla.


```python
todas_diferentes = ...
todas_diferentes
```

```python
grader.check("q5_3")
```

**Pregunta 4.** El archivo `sales.csv` contiene el número de frutas vendidas de cada caja el sábado pasado. Esta tiene una columna extra llamada `price per fruit ($)`, la cual es el precio *por unidad*  por fruta en dicha caja. Las filas están en el mismo orden que en la tabla `inventario`.Cargue estos datos en una tabla llamada `ventas`. **(5 Puntos)**

```python
ventas = ...
ventas
```

```python
grader.check("q5_4")
```

**Pregunta 5.** ¿Cuántas frutas se vendieron en total en la tienda ese día? **(5 Puntos)**


```python
total_frutas_vendidas = ...
total_frutas_vendidas
```

```python
grader.check("q5_5")
```

**Pregunta 6.** ¿Cuál fue la ganancia total de la tienda (el precio total de todas las frutas vendidas) **(5 Puntos)**

*Hint:* Si se atasca, piense primero en cómo podría calcular la ganancia total solo de las ventas de uvas.




```python
ganancia_total = ...
ganancia_total
```

```python
grader.check("q5_6")
```

**Pregunta 7.** Cree una nueva tabla llamada `inventario_restante`. Esta deberá contener las mismas filas y columnas de `inventario`, exceptuando que la cantidad de frutas vendidas de cada caja deberá ser restada de la cuenta de la caja **original**, de forma que la columna `count` sea **actualizada para ser**  la cantidad de fruta que queda luego del sábado. **(5 Puntos)**


```python
inventario_restante = ...
    ...
    ...
    ...
inventario_restante
```

```python
grader.check("q5_7")
```

¡Terminó la tarea 2!

**Información sobre la entrega:** Asegúrese de correr todas las pruebas y verificar que se ejectuen correctamente. Aquí, seleccione **Guardar Notebook** En el menú de **Archivo** y **corra la  última celda** para luego darle click en descargar el archivo zip. Con el archivo, vaya a [Gradescope](https://www.gradescope.com/courses/576157) y cargue el archivo zip correspondiente a esta tarea. El nombre de esta tarea es "HW 01 Autograder". **Es su responsabilidad asegurarse que su trabajo se guarde antes de correr la última celda**.

Una vez la haya entregado, su tarea en Gradescope deberá mostrar todas las pruebas que usted aprobó en el notebook de esta tarea.

## Estudiantes de Fundamentos de Ciencia de Datos

¡Felicitaciones por terminar la tarea 2!

<img src="./kim.jpg" width="30%" alt="Close up picture of dog smiling."/>

Perro de la semana: **Kim**

## Entrega

Asegúrese que ha ejecutado todas las celdas en su notebook en orden antes de ejecutar la celda de abajo. De esta forma todas las imágenes/gráficas apareceran en la salida. La celda de abajo va a generar un archivo con extensión .zip para que usted lo entregue. **¡Por favor guarde antes de exportar!**


```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_1 (points: [0, 0])
```
>>> import numpy as np
>>> # It looks like you didn't make an array.
>>> type(weird_numbers) == np.ndarray
True
```
```
>>> len(weird_numbers) == 4
True
```


Tests for q1_2 (points: [0, 0, 0, 0, 4])
```
>>> import numpy as np
>>> # It looks like you didn't make an array.
>>> type(book_title_words) == np.ndarray
True
```
```
>>> # It looks like you included commas in the text.
>>> # The three pieces of text in the array should be:
>>> #   "Eats"
>>> #   "Shoots"
>>> #   "and Leaves"
>>> not any([',' in text for text in book_title_words])
True
```
```
>>> # It looks like you didn't include both words in the
>>> # last piece of text.  It should be the actual string:
>>> #   "and Leaves"
>>> 'and ' in book_title_words.item(2)
True
```
```
>>> len(book_title_words)
3
```
```
>>> book_title_words.item(0) == 'Eats' and book_title_words.item(1) == 'Shoots' and book_title_words.item(2) == 'and Leaves'
True
```


Tests for q1_3 (points: [2, 2])
```
>>> with_commas == 'Eats, Shoots, and Leaves'
True
```
```
>>> without_commas == 'Eats Shoots and Leaves'
True
```


Tests for q2_1 (points: [0, 4])
```
>>> # It looks like you wrote:
>>> #   some_numbers.item(3)
>>> # But the third element has index 2, not 3.
>>> third_element != -10
True
```
```
>>> third_element == -6
True
```


Tests for q2_2 (points: [1, 1, 1, 1])
```
>>> elements_of_some_numbers.column(0).item(2) == "third"
True
```
```
>>> elements_of_some_numbers.column(0).item(3) == "fourth"
True
```
```
>>> elements_of_some_numbers.column(1).item(0) == 0
True
```
```
>>> elements_of_some_numbers.column(1).item(3) == 3
True
```


Tests for q2_3 (points: [0])
```
>>> # Make sure that you assign index_of_last_element to an index (integer), 
>>> # not the value at that index
>>> type(index_of_last_element) == int
True
```


Tests for q2_4 (points: [4])
```
>>> most_recent_birth_year == 1917
True
```


Tests for q2_5 (points: [0])
```
>>> min_of_birth_years > 0
True
```


Tests for q3_1 (points: [0])
```
>>> first_product > 0, second_product < 0, third_product > 0, fourth_product > 0
(True, True, True, True)
```


Tests for q3_2 (points: [0])
```
>>> # products should be an array!
>>> len(products) == 4
True
```


Tests for q3_3 (points: [0])
```
>>> len(correct_products) == 4
True
```


Tests for q3_4 (points: [1, 1, 1, 1])
```
>>> # It looks like you multiplied and subtracted in the wrong
>>> # order.
>>> sum(celsius_temps_rounded) != 356705.0
True
```
```
>>> # Don't forget to round your final results
>>> celsius_temps_rounded.item(1) == 31.0
True
```
```
>>> sum(celsius_temps_rounded)
1280677.0
```
```
>>> len(celsius_temps_rounded)
65000
```


Tests for q3_5 (points: [0, 0, 4, 0])
```
>>> # celsius_temperature_ranges should be an array
>>> type(celsius_temperature_ranges) is np.ndarray
True
```
```
>>> # Don't round your answers
>>> np.isclose(celsius_temperature_ranges.item(0), 7.222)
False
```
```
>>> np.isclose(round(sum(celsius_temperature_ranges)), 768487)
True
```
```
>>> len(celsius_temperature_ranges) == len(max_temperatures)
True
```


Tests for q4_1 (points: [0, 0])
```
>>> # Hint: shortest is a number between 40 and 50.
>>> 40 <= shortest <= 50
True
```
```
>>> # Hint: the average is between the shortest and the longest
>>> shortest <= average <= longest
True
```


Tests for q4_2 (points: [0, 0])
```
>>> # Hint: If you are getting 47 as your answer, you might be computing the biggest change 
>>> # rather than the biggest decrease!
>>> biggest_decrease == 47
False
```
```
>>> # Hint: biggest decrease is above 30, but not 47.
>>> 30 <= biggest_decrease < 47
True
```


Tests for q4_3 (points: [0, 4])
```
>>> len(first_nine_waiting_times)
9
```
```
>>> total_waiting_time_until_tenth
633
```


Tests for q4_4 (points: [0])
```
>>> 15 <= average_error <= 25
True
```


Tests for q5_1 (points: [4])
```
>>> fruits.sort(0)
fruit name | amount
apple      | 4
orange     | 3
pineapple  | 3
```


Tests for q5_2 (points: [4])
```
>>> inventory.sort(0).column(0).item(0)
25274
```


Tests for q5_3 (points: [0])
```
>>> all_different in {True, False}
True
```


Tests for q5_4 (points: [5])
```
>>> sales.sort(0)
box ID | fruit name | count sold | price per fruit ($)
25274  | apple      | 0          | 0.8
26187  | strawberry | 25         | 0.15
43566  | peach      | 17         | 0.8
48800  | orange     | 35         | 0.6
52357  | strawberry | 102        | 0.25
53686  | kiwi       | 3          | 0.5
57181  | strawberry | 101        | 0.2
57930  | grape      | 355        | 0.06
```


Tests for q5_5 (points: [0])
```
>>> # We're asking for the number of *pieces* of fruit, not the
>>> # number of kinds of fruit or the number of boxes from which
>>> # there were sales.
>>> total_fruits_sold > 10
True
```


Tests for q5_6 (points: [0])
```
>>> # If you're stuck, here's a hint: You want to multiply the count
>>> # sold in each box by the per-item price of fruits in that box.
>>> # You can use elementwise multiplication for that.
>>> # Then you want the sum of those products.  Use sum().
>>> 50 <= total_revenue <= 150
True
```


Tests for q5_7 (points: [0, 2, 3])
```
>>> # It looks like your table doesn't have all 3 columns that are
>>> # in the inventory table.
>>> remaining_inventory.num_columns
3
```
```
>>> #It looks like you forgot to subtract off the sales.
>>> remaining_inventory.column("count").item(0) != 45
True
```
```
>>> remaining_inventory.where(1, 'grape')
box ID | fruit name | count
57930  | grape      | 162
```
