```python
# Initialize Otter
import otter
grader = otter.Notebook("hw04.ipynb")
```

# Tarea 4: Funciones, Tablas y Grupos

Por favor complete este notebook llenando las celdas que se especifican. Antes de comenzar, ejecute la celda anterior para cargar las pruebas.



**Recursos útiles:**
- [Referencia de Python](http://data8.org/fa23/reference/): Hoja de trucos y métodos útiles de métodos de arreglos y tablas.

**Recommended Readings**:

* [Visualizando Distribuciones Numéricas](https://www.inferentialthinking.com/chapters/07/2/Visualizing_Numerical_Distributions.html)
* [Functiones y Tablas](https://www.inferentialthinking.com/chapters/08/Functions_and_Tables.html)

Complete este notebook llenando las celdas que se indican. **Antes de comenzar, ejecute la celda de abajo para configurar el notebook importando algunas librerías que serán útiles.** Cada vez que inicie su servidor, necesitará ejecutar esta celda de nuevo.

Para todos los problemas en los que deba escribir explicaciones y oraciones, **deberá** proporcionar la respuesta en el espacio designado. Además, a lo largo de esta tarea y de todas las futuras, ¡asegúrese de no reasignar variables en todo el notebook! Por ejemplo, si usa `temperatura` en su respuesta a una pregunta, no la reasigne más adelante. De lo contrario, ¡reprobará las pruebas que creía que estaba aprobando anteriormente!


**Fecha de entrega:**

Esta tarea deberá ser entregada ** ___ a las ___**. Si quiere puntos extra, entréguela el (un día antes). Trabajo que se entregue por fuera del plazo no será recibido.


**Nota: Esta tarea tiene pruebas ocultas. Eso significa que aunque los exámenes indiquen que se aprobó al 100%, no significa que su calificación final será del 100%. Realizaremos más pruebas ocultas de corrección una vez que todos entreguen la tarea.**


Evite compartir respuestas o hacer copia con sus compañeros. Establezca en cambio discusiones grupales con otros estudiantes y docentes del curso.

Si desea contar con el tiempo necesario para pedir ayuda o asesoría en caso de estar atascado, se recomienda iniciar temprano con la realización de las tareas. Tenga en cuenta el cronograma del curso y de entregas.

## 1. San Diego lleno de Burritos

```python
# Ejecute esta celda para configurar el notebook. Por favor no la modifique.


# Estas líneas importan los mòdulos NumPy y Datascience.
import numpy as np
from datascience import *

# Estas líneas hacen algo de magia en las gráficas.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')

import warnings
warnings.simplefilter('ignore', FutureWarning)
warnings.filterwarnings("ignore")
```

Eunice, Ciara and Kanchana están tratando de usar la Ciencia de Datos para encontrar los mejores burritos en San Diego. Sus amugas Jessica y Sonia les dieron dos conjuntos de datos (o datasets) que contiene iunformación sobre cuántos establecimientos hay en el área de San Diego. Tomado y depurado de: https://www.kaggle.com/srcole/burritos-in-san-diego/data



La siguiente celda lee una tabla llamada `ratings`, la cual contiene los nombres de los restaurantes de burritos, su calificación en Yelp, su calificación en Google así como su calificación total. La calificacón total o `Overall` no es un promedio de los puntajes en `Yelp` y `Google`, pero es una calificación general de los clientes que fueron considerados en el estudio.

Esta también lee una tabla llamada `burritos_types`, la cual contiene los nombres de los restaurantes de burritos, los elementos en sus menùs y el costo de cada ítem del menú en el restaurante.

```python
# Solo corra esta celda
ratings = Table.read_table("ratings.csv")
ratings.show(5)
burritos_types = Table.read_table("burritos_types.csv").drop(0)
burritos_types.show(5)
```

**Pregunta 1.** Le puede resultar conveniente combinar la información de ambas tablas. Asigne `burritos` al resultado de unir las dos tablas, de forma que obtengamos una tabla con el puntaje de cada ítem del menú de cada restaurante. Cada ítem del menú tiene la misma calificación que el restaurante que lo vende. **(8 Puntos)**


*Nota:* No importa cuán tabla ingrese como el argumento del método de tabla. Cualquier orden funcionará en el calificador de las pruebas.


*Hint:* ¿Cuál función nos permite combinar información de dos tablas? Use la [referencia de Python](http://data8.org/fa23/reference/) si no lo tiene claro.



```python
burritos = ...
burritos.show(10)
```

```python
grader.check("q1_1")
```

<!-- BEGIN QUESTION -->

**Pregunta 2.** Echemos un vistazo a cómo se comparan los puntajes en Yelp se comparan con los puntajes en Google en la tabla `burritos`. Primero, asigne `yelp_google_tbl` a una tabla que solo contenga las columnas `Yelp` y `Google`. Luego, haga una gráfica de dispersión con los puntajes de Yelp en el eje x y con los de Google en el eje y. **(8 Points)**



```python
yelp_google_tbl = ...
...

# No cambie/edite/borre la siguiente línea.
# Para ayudarlo a sacar conclusiones, graficamos una línea recta en la gráfica (y=x)
plt.plot(np.arange(2.5,5,.5), np.arange(2.5,5,.5));
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Pregunta 3.** Observando la gráfica de dispersión que acabó de generar en la pregunta 1.2, ¿nota algún/algunos patrones o relaciones entre los puntajes de Yelp y Google (i.e. es uno de los dos puntajes consistentemente mayor que el otro)? Si es así, descríbalos **brevemente** en la siguiente celda. **(8 Puntos)**



_Escriba aquí su respuesta, reemplazando este texto._

<!-- END QUESTION -->


Aquí va un recordatorio sobre cómo `.group` funciona. Puede leerlo en el [libro](https://www.inferentialthinking.com/chapters/08/2/Classifying_by_One_Variable.html), o puede ver el vídeo de abajo (está en inglés).


Puede también usar el [Visualizador de Funciones de Tabla](http://data8.org/interactive_table_functions/) para obtener más experiencia con la función `.group`.

```python
# No se preocupe por lo que hace este código. Simplemente embebe un visualizador de Youtube en una celda de código.
from IPython.display import YouTubeVideo
YouTubeVideo("HLoYTCUP0fc")
```

**Pregunta 4.** En la tabla `burritos` hay muchos tipos de burritos de California. Kanchana quiere basar sus opciones de burritos en los puntajes y rankings. Para responder a esta pregunta, estamos tratando el puntaje de cada ítem en los menús como si fuera el mismo que el del restaurante, dado que no tenemos el puntaje individual de cada plato en esos restaurantes. No se preocupe por este hecho, pero vale la pena mencionarlo.

Cree una tabla que contenga dos columnas: la primera columna incluye **nombre de los burritos** y la segunda deberá contener el **puntaje total promedio** de ese burrito en los diferentes restaurantes. **En sus cálculos solo deberá comparar los burritos con la palaba "California" en su nombre.** Por ejemplo, hay burritos "California", burritos "Californa de desayuno", etc. **(9 Puntos)**


*Hint:* ¡“California” es un caso sensible (“California” no es lo mismo que “california”)!

*Nota:* La solucón es posible obtenerla en una sola línea. Sin embargo, siéntase libre de encontrar su solución en varias líneas y pasos. Solo asegúrese que al final asigne la tabal resultante a `california_burritos`.


```python
california_burritos = ...
california_burritos
```

```python
grader.check("q1_4")
```

**Pregunta 5.** Dada la nueva tabla `california_burritos`, Ciara puede encontrar el nombre del burrito Californa con el mejor puntaje promedio. Asigne `best_california_burrito` a una línea de código que le entregue la cadena de caracteres que represente el nombre del burrito California con el mejor puntaje promedio. Si varios burritos cumplen este criterio, solo deberá colocar cualquiera de ellos. **(8 Puntos)**


```python
best_california_burrito = ...
best_california_burrito
```

```python
grader.check("q1_5")
```

<!-- BEGIN QUESTION -->

**Question 6.** Eunice piensa que los burritos en San Diego son más económicos (y que saben mejor) que los burritos de Berkeley. Grafique un histograma que  permita visualizar la distribución de los costos de los burritos de San Diego en la tabla `burritos`. Use la variable `cost_bins` cuando haga el histograma, para hacerlo visualmente más informativo. **(8 Puntos)**


```python
cost_bins = np.arange(0, 15, 1) # No cambie esta línea
# Use, por favor, los bines que se le dan
...
```

<!-- END QUESTION -->

**Pregunta 7.** ¿Qué porcentaje de burritos en San Diego cuestan menos de $6 dólares? Asigne `burritos_less_than_six` a su respuesta, **la cual debe estar entre 0 y 100**. Usted deberá usar solamente le hnistograma del punto anterior para responder esta pregunta. No use líneas de cóndigo para buscar la respuesta en la tabla, solo use lo que observe en las alturas del histograma y use Python para evaluar la aritmética. **(8 Puntos)**


*Nota*:Su respuetsa no tiene que ser exacta, pero debe estar dentro de un intervalod e porcentajes que sea razonable.

*Hint:* Puede ser útil leer la [sección 7.2](https://inferentialthinking.com/chapters/07/2/Visualizing_Numerical_Distributions.html) del texto guía sobre el cálculo de porcentajes a partir de un histograma.

```python
burritos_less_than_six = ...
```

```python
grader.check("q1_7")
```

## 2. Salarios de Empleados de la Ciudad de San Francisco

Este ejercicio está diseñado para que ponga en práctica los métodos de tabla `.pivot` y `.group`. Aquí va un link a la [referencia de Python](http://data8.org/fa23/reference/) en caso que requiere refrescar algunos conceptos. El [visualizador de funciones de tabla](http://data8.org/interactive_table_functions/)también puede ser una herramienta útil.

Ejecute la celda de abajo para que vea una demostración sobre cómo usar pivot en una tabla.

```python
# # No se preocupe por lo que hace este código. Simplemente embebe un visualizador de Youtube en una celda de código.
from IPython.display import YouTubeVideo
YouTubeVideo("4WzXo8eKLAg")
```

La fuente de datos que usaremos en esta porción de la tarea es [provista públicamente](https://data.sfgov.org/City-Management-and-Ethics/Employee-Compensation/88g8-5mnd/data) por la Ciudad de San Francisco. Con ello, la hemos filtrado solo para extraer las columnas relevantes y restringido solo a datos del año calendario 2019. Ejecute la siguiente celda de código para cargar nuestros datos en una tabla llamada `full_sf`.

```python
full_sf = Table.read_table("sf2019.csv")
full_sf.show(5)
```

La tabla contiene una fila por cada uno de los 44525 empleados públicos de San Francisco en 2019.

Las primeras cuatro columnas describen el trabajo del empleado. Por ejemplo, el empleado en la tercera fila tiene un trabajo llamado "IS Business Analyst-Senior". Llamaremos a esta la *posición* o *título del trabajo* del empleado. El trabajo está en una familia de trabajos llamada Sistemas de Información o en inglés, Information Systems (dada la sigla IS en el inicio del título). La mayoría del tiempo usted estará trabajando con la columna `Job`.

Las siguientes tres columnas contienen la cantidad de dólares que fueron pagados al empleado como salario en el año calendario 2019, tiempos extras y beneficios. Note que el salario total del empleado no incluye las ganancias por horas extra.

La última columna contiene la compensación total pagada al empleado. Esto es la suma de las tres columnas previas

The last column contains the total compensation paid to the employee. It is the sum of the previous three columns:

$$\text{Compensación Total} = \text{Salario} + \text{Horas Extra} + \text{Beneficios}$$

Para esta tarea, estaremos usando las siguientes columnas:

1. `Organization Group`: Un grupo de departamentos. Por ejemplo el "Public Protection Org." incluye los departamentos de Policía, Bomberos, Protección al Mayor, etc.
2. `Department`: La unidad primaria organizacional usada por la Ciudad y el Condado de San Franciso.
3. `Job`: La posición específica que ese empleado ocupa.
4. `Total Compensation`: La suma del salario del empleado, horas extra y beneficios recibidos en 2019.

Ejecute la celda de abajo para seleccionar las cuatro columnas relevantes para nosotros, y cree una nueva tabla llamada `sf`.


```python
sf = full_sf.select("Job", "Department", "Organization Group",  "Total Compensation")
sf.show(5)
```

Queremos usar esta tabla para generar arreglos con los títulos del trabajo de los miembros de cada **Organization Group**.


**Pregunta 1.** Asigne `job_titles` a una tabla con dos columnas. La primer columnda se deberá llamar `Organization Group` y contiene el nombre de cada "Organization Group" listado cada uno una sola vez en esta columna. La segunda columna deberá llamarla `Jobs` con cada fila en dicha segunda columna conteniendo un *array* con los nombres de los títulos dle trabajo en cada "Organization Group".  No se preocupe si se repite muchas veces el mismo título del trabajo. **(9 Puntos)**



*Hint 1:* Piense en cómo funciona `group`: él recoge los valores dentro de un arreglo y aplica una función a ese arreglo. Hemos definido dos funciones abajo para usted, y necesitará utilizar una de ellas cuando use `group`.

*Hint 2:* Probablemente deba renombrar una de las dos columnas.

```python
# Escoja una de las dos funciones para llamar group.
def first_item(array):
    '''Regresa el primer ítem'''
    return array.item(0)

def full_array(array):
    '''Regresa el arreglo que se pasa'''
    return array

# Llame group usando una de las funciones de arriba cuando defina job_titles
job_titles = ...
job_titles
```

```python
grader.check("q2_1")
```

<!-- BEGIN QUESTION -->

**Pregunta 2.** Al momento, la columna `Job` en la tabla `sf` no está organizada (no tiene un orden particular). ¿Serán los arreglos que usted generó en la columna `Jobs` en la pregunta anterior la misma si la hubiésemos organizado alfabéticamente luego de generarlos? Para tener su respuesta correcta, su respuesta deberá explicar cómo el método `.group` funciona y cómo organizar la columna `Jobs` puede afectarlo. **(8 Points)**


*Nota:* Dos arreglos son **iguales** si contienen el mismo número de elementos y dichos elementos est+an ubicados en los mismos índices. Un ejemplo que dos arreglos que NO son el mismo:  `array([1,2]) != array([2,1])`.




_Escriba aquí su respuesta, reemplazando este texto._

<!-- END QUESTION -->

**Pregunta 3.** Asigne `department_ranges` a una tabla que contenga los departamentos como columnas, y los grupos de organización como las columnas. Los valores en las columnas deberán corresponder al rango total de compensación (o pago), donde el rango es la **diferencia entre la mayor y la menor compensación total en el departamento para ese grupo de organización**. **(9 Puntos)**

*Hint:* Primero tendrá que definir una nuneva función `compensation_range` que tome como entrada un arreglo de las compensaciones y regrese el rango de las compensaciones en dicho arreglo.


```python
# Defina primero compensation_range
...
    ...

department_ranges = ...
department_ranges
```

```python
grader.check("q2_3")
```

<!-- BEGIN QUESTION -->

**Pregunta 4.** ¿Por qué deberían algunos valores en la tabla `department_ranges` creada en la pregunta anterior ser iguales a `0`? **(8 Puntos)**


_Escriba aquí su respuesta, reemplazando este texto._

<!-- END QUESTION -->

**Pregunta 5.** Encuentre el número de departamentos que aparecen en la tabla `sf` que tienen un promedio de compensación total mayor que 125000 dólares. Asigne este valor a la variable `num_sobre_125k`. **(9 Puntos)**


*Nota:* Los nombres de las variables están pensados para guiarlo en los pasos intermedios y en el proceso en general. Siéntase libre de borrarlos si prefiere empezar desde cero, pero asegúrese que su pespuesta final sea `num_sobre_125k`.




```python
depts_and_comp = ...
departmento_promedio = ...
num_sobre_125k = ...
num_sobre_125k
```

```python
grader.check("q2_5")
```

¡Terminó la tarea 4!  

**Pasos importantes para la entrega:**
1. Corra las pruebas y verifique que todas se ejecuten.
2. Seleccione **Guardar Notebook** desde el menú **Archivo**, para luego **ejecutar la última celda**.
3. De click en el link para descargar el archivo zip.
4. Vaya a [Gradescope](https://www.gradescope.com/courses/576157) y cargue el archivo zip correspondiente a esta tarea.

**Es su responsabilidad asegurarse que su trabajo se guarde antes de correr la última celda.**

## Estudiantes de Fundamentos de Ciencia de Datos

¡Felicitaciones por terminar la tarea 3!

<img src="./gus.jpg" width="30%" alt="Close up picture of dog smiling."/>

Mascota de la semana: **Gus**


## Entrega

Asegúrese que ha ejecutado todas las celdas en su notebook en orden antes de ejecutar la celda de abajo. De esta forma todas las imágenes/gráficas apareceran en la salida. La celda de abajo va a generar un archivo con extensión .zip para que usted lo entregue. **¡Por favor guarde antes de exportar!**

```python
# Guarde primero su notebook, luego ejecute esta celda para obtener el archivo a
# entregar.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_1 (points: [0, 0])
```
>>> # Make sure you have all the columns from both tables
>>> set(["Name", "Menu_Item", "Yelp", "Google", "Overall", "Cost"]) == set(burritos.labels)
True
```
```
>>> burritos.num_rows == 212
True
```


Tests for q1_4 (points: [1, 1])
```
>>> # Number of columns should be 2
>>> california_burritos.num_columns == 2
True
```
```
>>> # Number of rows should be 19
>>> california_burritos.num_rows == 19
True
```


Tests for q1_5 (points: [0])
```
>>> # Your answer should be a string
>>> type(best_california_burrito) == str
True
```


Tests for q1_7 (points: [0])
```
>>> # Your answer should be between 0 and 100.
>>> 0 <= burritos_less_than_six <= 100
True
```


Tests for q2_1 (points: [0, 0, 0, 0])
```
>>> job_titles.num_columns
2
```
```
>>> job_titles.num_rows
6
```
```
>>> # Make sure that you have the correct column labels!
>>> np.asarray(job_titles.labels).item(1) != "Job full_array"
True
```
```
>>> # Make sure that you have the correct column labels!
>>> np.asarray(job_titles.labels).item(1) == "Jobs"
True
```


Tests for q2_3 (points: [0, 0, 2])
```
>>> # Double check that your salary_range function is correct
>>> compensation_range(make_array(5, 1, 20, 1000)) == 999
True
```
```
>>> # Make sure that the table has departments as the rows and organization groups as the columns.
>>> set(["Department", "Community Health", "Culture & Recreation", "General Administration & Finance", "Human Welfare & Neighborhood Development", "Public Protection", "Public Works, Transportation & Commerce"]) == set(department_ranges.labels)
True
```
```
>>> sum(department_ranges.column(1))
554179.0
```


Tests for q2_5 (points: [2, 2])
```
>>> # Make sure that your final answer is a number
>>> isinstance(num_over_125k, int)
True
```
```
>>> # Make sure that your answer makes sense given the sf table
>>> 0 <= num_over_125k <= 51
True
```
