<a href="https://colab.research.google.com/github/darioblandon/materials-fa23/blob/main/lab/lab08/lab08_td.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab08.ipynb")
```

<p><img alt="Colaboratory logo" height="140px" src="https://upload.wikimedia.org/wikipedia/commons/archive/f/fb/20161010213812%21Escudo-UdeA.svg" align="left" hspace="10px" vspace="0px"></p>

<h1> Fundamentos en ciencias de datos</h1>

# Laboratorio 8: Distribución normal y variabilidad de medias muestrales

¡Bienvenidos al Laboratorio 8!

En el laboratorio de hoy, aprenderemos sobre [la variabilidad de las medias muestrales](https://inferentialthinking.com/chapters/14/5/Variability_of_the_Sample_Mean.html), así como [la distribución normal](https://www.inferentialthinking .com/chapters/14/3/SD_and_the_Normal_Curve.html).

**Fecha límite de envío de laboratorio: viernes 27 de octubre a las 11 p.m.**

**Obtener ayuda en el laboratorio**: Siempre que te sientas estancado o necesites más aclaraciones, busca un GSI o un tutor, ¡y estarán encantados de ayudarte!

Estas son las políticas para obtener crédito completo:

1. Para los estudiantes inscritos en <span style="color: #BC412B">**Regular Labs**</span> en persona, recibirán crédito de laboratorio al **asistir** a la sección de laboratorio y lograr un progreso significativo en el cuaderno de laboratorio. **<span style="text-decoration: underline">debes</span>** enviar esta tarea a Gradescope.

2. Los estudiantes inscritos en <span style="color: #04888B">**Laboratorio de autoservicio**</span> recibirán crédito de laboratorio al **completar** el cuaderno, aprobar los casos de prueba y **enviándolo** a Gradescope antes de las **11:59 p. m. de la fecha de entrega**.

**Envío**: una vez que haya terminado, ejecute todas las celdas además de la última, seleccione Archivo > Guardar cuaderno y luego ejecute la celda final. El resultado contendrá un archivo zip que podrá utilizar para enviarlo en Gradescope.

Comencemos configurando las pruebas y las importaciones ejecutando la siguiente celda.

```python
# Ejecute esta celda, pero no la cambie.

# Estas líneas importan los módulos Numpy y Datascience.
import numpy as np
from datascience import *

# Estas líneas hacen algo de magia en la trama.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore')
```

# 1. Distribuciones normales

Cuando visualizamos la distribución de una muestra, a menudo nos interesa la media y la desviación estándar de la muestra (durante el resto de esta práctica de laboratorio, abreviaremos “desviación estándar” como “DE”). Estas dos estadísticas de muestra pueden darnos una visión general de la distribución, permitiéndonos saber dónde se ubica la distribución en la recta numérica y qué tan extendida está, respectivamente.

Queremos comprobar si los datos están relacionados linealmente, por lo que deberíamos mirar los datos.

Primero, carguemos la tabla de "nacimientos" de la conferencia, que es una muestra aleatoria grande de nacimientos en EE. UU. e incluye información sobre las parejas de madre e hijo.

```python
births = Table.read_table('baby.csv')
births
```

**Pregunta 1.1.** **Ahora, escriba una línea de código para trazar la distribución de las edades de la madre a partir de la tabla `nacimientos`.** No cambie la última línea, que trazará la media de la muestra en el distribución misma.


```python
...

# No cambies nada debajo de esta línea
plots.scatter(np.mean(births.column("Maternal Age")), 0, color='red', zorder=2, s=50);
plots.ylim(-0.005, 0.07);
```

En el gráfico anterior, podemos ver que la media es el centro de gravedad o punto de equilibrio de la distribución. Si corta la distribución de cartón y luego coloca el dedo en la media, la distribución se equilibrará perfectamente en su dedo. Dado que la distribución anterior está sesgada a la derecha (lo que significa que tiene una cola larga hacia la derecha), sabemos que la media de la distribución es mayor que la mediana, que es el punto "medio" de los datos. Por el contrario, si la distribución hubiera quedado asimétrica, sabemos que la media sería menor que la mediana.

**Pregunta 1.2.** Ejecute la siguiente celda para comparar la media (rojo) y la mediana (azul) de la distribución de edades de las madres.

```python
births.hist("Maternal Age")
plots.scatter(np.mean(births.column("Maternal Age")), 0, color='red', zorder=2, s=50, label="Mean")
plots.scatter(np.median(births.column("Maternal Age")), 0, color='blue', zorder=2, s=50, label="Median")
plots.legend()
plots.ylim(-0.005, 0.07);
```

También nos interesa la desviación estándar de las edades de la madre. La DE nos da una idea de cuán variables son las edades de las madres en torno a la edad promedio de las madres. Si la DE es grande, entonces las edades de las madres deberían extenderse en un rango amplio con respecto a la media. Si la DE es pequeña, entonces las edades de las madres deberían agruparse estrechamente en torno a la edad materna promedio.

**La DE de una matriz se define como la "media cuadrática de las desviaciones", donde las desviaciones son la diferencia entre los elementos originales y su promedio.**

Puede leer más sobre Variabilidad y SD en el [Capítulo 14.2](https://inferentialthinking.com/chapters/14/2/Variability.html). Este capítulo también profundiza un poco más en cómo se calcula la "media cuadrática de las desviaciones". ¡Hecho de la diversión! $\sigma$ (letra griega sigma) se usa para representar la SD y $\mu$ (letra griega mu) se usa para la media.

**Pregunta 1.3.** Asigne `age_mean` a la media de las edades maternas y `age_sd` a la DE de las edades maternas. Luego ejecute la celda a continuación para ver el ancho de una DE (azul) de la media de la muestra (rojo) trazada en el histograma de edades maternas.

*Pista*: `np.std` podría resultar útil aquí.

```python
age_mean = ...
age_sd = ...
births.hist("Maternal Age")

plots.scatter(age_mean, 0, color='red', zorder=2, s=50);
plots.scatter(age_mean+age_sd, 0, marker='^', color='blue', zorder=2, s=50);
plots.scatter(age_mean-age_sd, 0, marker='^', color='blue', zorder=2, s=50);
plots.ylim(-0.005, 0.07);
```

```python
grader.check("q1_3")
```

En este histograma, la desviación estándar no es fácil de identificar con solo mirar el gráfico.

Sin embargo, las distribuciones de algunas variables nos permiten detectar fácilmente la desviación estándar en el gráfico. Por ejemplo, si una muestra sigue una **[distribución normal](https://inferentialthinking.com/chapters/14/3/SD_and_the_Normal_Curve.html)**, la desviación estándar se detecta fácilmente en el punto de inflexión (el punto donde la curva comienza a cambiar la dirección de su curvatura) de la distribución.

**Pregunta 1.4.** Complete el siguiente código para examinar la distribución de las alturas maternas, que se distribuye aproximadamente normalmente. Graficaremos la desviación estándar en el histograma, como antes; observe dónde cae en el gráfico una desviación estándar (azul) de la media (rojo).


```python
height_mean = ...
height_sd = ...
births.hist("Maternal Height", bins=np.arange(55,75,1))

plots.scatter((height_mean), 0, color='red', zorder=2, s=50);
plots.scatter(height_mean+height_sd, 0, marker='^', color='blue', zorder=2, s=50);
plots.scatter(height_mean-height_sd, 0, marker='^', color='blue', zorder=2, s=50);
plots.ylim(-0.01, 0.16);
```

```python
grader.check("q1_4")
```

No siempre sabemos cómo se distribuirá una variable, y hacer suposiciones sobre si una variable seguirá o no una distribución normal es peligroso. Sin embargo, el **Teorema del límite central** define una distribución que siempre sigue una distribución normal. **La distribución de las *sumas* y *medias* de todas las muestras aleatorias grandes extraídas con reemplazo de una única distribución (independientemente de la forma original de la distribución) tendrá una distribución normal.** Recuerde que el teorema del límite central se refiere a la distribución de una *estadística* calculada a partir de una distribución, no de la distribución de la muestra o población original. Si esto le resulta confuso, ¡pregúntele a un miembro del personal!

La siguiente sección explorará las distribuciones de medias muestrales y verá cómo la desviación estándar de estas distribuciones depende del tamaño de las muestras.

# 2. Variabilidad de la media muestral

Según el [Teorema del límite central] (https://www.inferentialthinking.com/chapters/14/4/Central_Limit_Theorem.html), la distribución de probabilidad de la media o suma de una muestra aleatoria grande es aproximadamente normal. La curva de campana está centrada en la media poblacional. Algunas de las medias muestrales son más altas y otras más bajas, pero las desviaciones de la media poblacional son aproximadamente simétricas en ambos lados, como hemos visto repetidamente.

En nuestras simulaciones, también notamos que las medias de muestras más grandes tienden a agruparse más estrechamente alrededor de la media poblacional que las medias de muestras más pequeñas. En esta sección, cuantificaremos la [variabilidad de la media muestral] (https://www.inferentialthinking.com/chapters/14/5/Variability_of_the_Sample_Mean.html) y desarrollaremos una relación entre la variabilidad y el tamaño de la muestra.

Echemos un vistazo a los salarios de los empleados de la ciudad de San Francisco en 2014. El salario medio informado por el gobierno de la ciudad fue de aproximadamente $75,463.92.

*Nota: Si se queda atascado en alguna parte de esta práctica de laboratorio, consulte el [Capítulo 14 del libro de texto](https://www.inferentialthinking.com/chapters/14/Why_the_Mean_Matters.html).*

```python
salaries = Table.read_table('sf_salaries_2014.csv').select("salary")
salaries.show(5)
```

```python
salary_mean = np.mean(salaries.column('salary'))
print('Mean salary of San Francisco city employees in 2014: ', salary_mean.round(2))
```

```python
salaries.hist('salary', bins=np.arange(0, 250000+10000*2, 10000))
plots.scatter(salary_mean, 0, color='red', zorder=2, s=50)
plots.title('2014 salaries of city of SF employees')
plots.ylim(-0.0000005, 0.000014);
```

Claramente, la población *no* sigue una distribución normal. Tenga esto en cuenta a medida que avancemos en estos ejercicios.

Tomemos muestras aleatorias **con reemplazo** y observemos la distribución de probabilidad de la media muestral. Como es habitual, utilizaremos la simulación para obtener una aproximación empírica a esta distribución.

**Pregunta 2.1.** Defina una función `one_sample_mean`. Sus argumentos deben ser "tabla" (el nombre de una tabla), "etiqueta" (la etiqueta de la columna que contiene la variable) y "tamaño_muestra" (el número de empleados en la muestra). Debe probarse con reemplazo de la tabla y
devuelve la media de la columna "etiqueta" de la muestra.

*Nota:* La función debe hacer referencia a los argumentos y poder trabajar con cualquier tabla.

```python
def one_sample_mean(table, label, sample_size):
    new_sample = ...
    new_sample_mean = ...
    ...
```

```python
grader.check("q2_1")
```

**Pregunta 2.2.** Utilice `one_sample_mean` para definir una función `simulate_sample_mean`. Los argumentos son el nombre de la "tabla", la "etiqueta" de la columna que contiene la variable, el "tamaño_muestra" y el número de simulaciones ("repeticiones").

La función debe muestrear con reemplazo de la tabla y calcular la media de cada muestra. Debería guardar los medios de muestra en una matriz llamada "medios". El código restante de la función muestra un histograma empírico de las medias de la muestra.

*Nota:* Redondeamos todos los valores a 2 decimales en el código proporcionado porque estamos trabajando con salarios.

```python
"""Empirical distribution of random sample means"""

def simulate_sample_mean(table, label, sample_size, repetitions):

    means = make_array()

    for i in np.arange(repetitions):
        new_sample_mean = ...
        means = ...

    sample_means = Table().with_column('Sample Means', means)

    # Display empirical histogram and print all relevant quantities – don't change this!
    sample_means.hist(bins=20)
    plots.xlabel('Sample Means')
    plots.title('Sample Size {sample_size}; {repetitions} Resamples'.format(sample_size=sample_size, repetitions=repetitions))
    print("Sample size: ", sample_size)
    print("Population mean:", np.mean(table.column(label)).round(2))
    print("Average of sample means: ", np.mean(means).round(2))
    print("Population SD:", np.std(table.column(label)).round(2))
    print("SD of sample means:", np.std(means).round(2))
    return np.std(means)
```

Verifique con un compañero o asistente técnico que haya implementado correctamente la función anterior. Si no lo ha implementado correctamente, el resto del laboratorio no funcionará correctamente, por lo que este paso es crucial.

En la siguiente celda, crearemos una muestra de tamaño 100 a partir de "salarios" y la graficaremos usando nuestra nueva función "simulate_sample_mean".

  *Pista: Deberías ver una distribución similar a algo de lo que hemos estado hablando. Si no, verifique su función.*

```python
simulate_sample_mean(salaries, 'salary', 100, 10000)
plots.xlim(50000, 100000);
```

Notice that our distribution of *sample means* looks approximately normal! Did it matter if our original distribution was normally distributed?

**Pregunta 2.3.** Simule dos medias muestrales, una para una muestra de 400 salarios y otra para una muestra de 625 salarios. En cada caso, realiza 10.000 repeticiones. No se preocupe por las líneas `plots.xlim` o `plots.ylim`; solo se aseguran de que todos los gráficos tengan el mismo eje x e y, respectivamente.

```python
simulate_sample_mean(..., ..., ..., ...)
plots.xlim(50000, 100000)
plots.ylim(0, 0.0002)
plots.show();
print('\n')
simulate_sample_mean(..., ..., ..., ...)
plots.xlim(50000, 100000)
plots.ylim(0, 0.0002)
plots.show();
```

**Pregunta 2.4.** Asigne `q2_4` a una matriz de números correspondientes a afirmaciones verdaderas sobre las gráficas de 3.3.

*Pista*: Si una distribución tiene forma de campana, ¿qué tipo de curva sigue?

1. Vemos el teorema del límite central (CLT) en acción porque las distribuciones de las medias muestrales tienen forma de campana.
2. Vemos la Ley de los Promedios en acción porque las distribuciones de las medias muestrales se parecen a la distribución de la población.
3. Una de las condiciones para CLT es que tenemos que extraer una pequeña muestra aleatoria con reemplazo de la población.
4. Una de las condiciones para CLT es que tenemos que extraer una muestra aleatoria grande con reemplazo de la población.
5. Una de las condiciones para CLT es que la población debe estar distribuida normalmente.

```python
q2_4 = ...
```

```python
grader.check("q2_4")
```

**Pregunta 2.5.** Asigne `q2_5` a una matriz de números correspondientes a afirmaciones verdaderas sobre las gráficas de 2.3.

1. Ambos gráficos en 2.3 se centran aproximadamente en torno a la media poblacional.
2. Ambos gráficos en 2.3 se centran aproximadamente en torno a la media de una muestra particular.
3. La distribución de medias muestrales para un tamaño de muestra 625 tiene menos variabilidad que la distribución de medias muestrales para un tamaño de muestra 400.
4. La distribución de medias muestrales para un tamaño de muestra 625 tiene más variabilidad que la distribución de medias muestrales para un tamaño de muestra 400.


```python
q2_5 = ...
```

```python
grader.check("q2_5")
```

A continuación, veremos lo que sucede cuando tomamos un **número creciente de remuestreos de un tamaño de muestra fijo**. Observe qué número en el código cambia y qué permanece igual. ¿Cómo cambia la distribución de las medias remuestreadas?

```python
simulate_sample_mean(salaries, 'salary', 100, 500)
plots.xlim(50000, 100000)
plots.ylim(0, 0.00009);
```

```python
simulate_sample_mean(salaries, 'salary', 100, 1000)
plots.xlim(50000, 100000);
plots.ylim(0, 0.00009);
```

```python
simulate_sample_mean(salaries, 'salary', 100, 5000)
plots.xlim(50000, 100000);
plots.ylim(0, 0.00009);
```

```python
simulate_sample_mean(salaries, 'salary', 100, 10000)
plots.xlim(50000, 100000);
plots.ylim(0, 0.00009);
```

¿Qué notaste acerca de las distribuciones de medias muestrales en los cuatro histogramas anteriores? Discuta con sus compañeros o pregúntele a un miembro del personal.

**Pregunta 2.6.** Asigne la variable `SD_of_sample_means` al número entero correspondiente a su respuesta a la siguiente pregunta:

Cuando aumento el número de *nuevas muestras* que tomo, para un tamaño de muestra *fijo*, la SD de mis medias de muestra será...

1. aumentar
2. Disminución
3. Sigue igual
4. Varían ampliamente


```python
SD_of_sample_means = ...
```

```python
grader.check("q2_6")
```

**Pregunta 2.7.** Pensemos en cómo cambian las relaciones entre la DE de la población, la DE de la muestra y la DE de las medias de la muestra al variar el tamaño de la muestra. ¿Cual de los siguientes es verdadero? Asigne la variable `pop_vs_sample` a una matriz de números enteros que correspondan a declaraciones verdaderas.

*Pista 1:* La DE de la muestra es diferente de la DE de las medias de la muestra.

*Pista 2:* [Capítulo 14.5](https://inferentialthinking.com/chapters/14/5/Variability_of_the_Sample_Mean.html) podría ser útil para responder esta pregunta.

1. La DE de la muestra se hace más pequeña a medida que aumenta el tamaño de la muestra.
2. La DE de la muestra aumenta a medida que aumenta el tamaño de la muestra.
3. La DE de la muestra se vuelve más consistente con la DE de la población a medida que aumenta el tamaño de la muestra.
4. La DE de las medias muestrales se reduce al aumentar el tamaño de la muestra.
5. La DE de las medias muestrales aumenta a medida que aumenta el tamaño de la muestra.
6. La DE de las medias muestrales permanece igual al aumentar el tamaño de la muestra.

```python
pop_vs_sample = ...
```

```python
grader.check("q2_7")
```

**Pregunta 2.8.** ¿Existe una relación entre el tamaño de la muestra y la desviación estándar de las medias muestrales? Asigne `q2_8` al número correspondiente a la declaración que responde a esta pregunta.

*Pista:* [Capítulo 14.5](https://inferentialthinking.com/chapters/14/5/Variability_of_the_Sample_Mean.html) del libro de texto puede resultar útil.

1. La DE de las medias muestrales es inversamente proporcional a la raíz cuadrada del tamaño de la muestra.
2. La DE de las medias muestrales es directamente proporcional a la raíz cuadrada del tamaño de la muestra.

```python
q2_8 = ...
```

```python
grader.check("q2_8")
```

A lo largo de este laboratorio, hemos estado tomando muchas muestras aleatorias de una población. Sin embargo, todos estos principios son válidos para remuestreos con arranque de una sola muestra. El bootstrap funciona porque es como extraer datos de la población original, asumiendo que la muestra es representativa. Si su muestra original es relativamente grande, todas sus remuestras también serán relativamente grandes, por lo que la SD de las medias remuestreadas será relativamente pequeña.

**Para cambiar la variabilidad de la media de su muestra, tendría que cambiar el tamaño de la muestra original de la cual está tomando remuestras con arranque.**

-----
## ¡Hecho!

<img src="https://github.com/darioblandon/materials-fa23/blob/main/lab/lab08/oreo.jpeg?raw=1" alt="black and white doggie" width="300"/>

**Oreo** quiere recordarte que ¡ya estás a medio camino del final del semestre! ¡Puedes hacerlo!


---

¡Terminaste con el laboratorio!

**Información importante sobre el envío:**
- **Ejecuta todas las pruebas** y verifica que todas pasan
- **Guardar** desde el menú **Archivo**
- **Ejecute la celda final para generar el archivo zip**
- **Haga clic en el enlace para descargar el archivo zip**
- Luego, ve a [Gradescope](https://www.gradescope.com/courses/576157) y envía el archivo zip a la tarea correspondiente. El nombre de esta tarea es "Lab XX Autograder", donde XX es el número de laboratorio: 01, 02, 03, etc.

- Si terminas temprano en el laboratorio regular, **pídele a uno de los miembros del personal que te marque**.

**Es su responsabilidad asegurarse de que su trabajo esté guardado antes de ejecutar la última celda.**


## Envío

Asegúrese de haber ejecutado todas las celdas de su cuaderno en orden antes de ejecutar la siguiente celda, para que todas las imágenes/gráficos aparezcan en el resultado. La siguiente celda generará un archivo zip para que lo envíe. **¡Guarde antes de exportar!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_3 (points: None)
```
>>> round(age_mean, 3) == 27.228
True
```
```
>>> round(age_sd, 3) == 5.815
True
```


Tests for q1_4 (points: None)
```
>>> round(height_mean, 3) == 64.049
True
```
```
>>> round(height_sd, 3) == 2.525
True
```


Tests for q2_1 (points: None)
```
>>> np.random.seed(8)
>>> np.round(one_sample_mean(salaries, 'salary', 100), 3) == 76699.829
True
```


Tests for q2_4 (points: None)
```
>>> 
>>> import hashlib # This imports a hashing library for the autograder.
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> get_hash(set(np.array(q2_4).astype(int))) 
'3649655f87d3a44afa82f25eced24944'
```


Tests for q2_5 (points: None)
```
>>> 
>>> import hashlib # This imports a hashing library for the autograder.
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> get_hash(set(np.array(q2_5).astype(int)))
'b793dd27381f7b63cada88b749c0e023'
```


Tests for q2_6 (points: None)
```
>>> import hashlib # This imports a hashing library for the autograder.
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> get_hash(int(SD_of_sample_means))
'eccbc87e4b5ce2fe28308fd9f2a7baf3'
```


Tests for q2_7 (points: None)
```
>>> import hashlib # This imports a hashing library for the autograder.
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> get_hash(set(np.array(pop_vs_sample).astype(int)))
'1d919a653870c713cf8bd17b9c9c65ee'
```


Tests for q2_8 (points: None)
```
>>> 
>>> import hashlib # This imports a hashing library for the autograder.
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> get_hash(int(q2_8))
'c4ca4238a0b923820dcc509a6f75849b'
```
