<a href="https://colab.research.google.com/github/darioblandon/materials-fa23/blob/main/lab/lab02/lab02_td.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab02.ipynb")
```

<p><img alt="Colaboratory logo" height="140px" src="https://upload.wikimedia.org/wikipedia/commons/archive/f/fb/20161010213812%21Escudo-UdeA.svg" align="left" hspace="10px" vspace="0px"></p>

<h1> Fundamentos en ciencias de datos</h1>


#Laboratorio 2: Operaciones de tabla

¡Bienvenidos al Laboratorio 2! Esta semana, aprenderemos cómo importar un módulo y practicaremos operaciones de tablas. La [Referencia de Python](https://www.data8.org/fa23/reference/) tiene información que será útil para esta práctica de laboratorio.

**Lectura recomendada**:
  * [Introducción a las Tablas](https://www.inferentialthinking.com/chapters/03/4/Introduction_to_Tables)

Como recordatorio, estas son las políticas para obtener crédito completo:

1. Para los estudiantes inscritos en <span style="color: #BC412B">**Regular Labs**</span> en persona, recibirán crédito de laboratorio al **asistir** a la sección de laboratorio y lograr un progreso significativo en el cuaderno de laboratorio y enviándolo a Gradescope.

2. Para los estudiantes inscritos en <span style="color: #04888B">**Mega Lab**</span>, recibirán crédito de laboratorio al **completar** el cuaderno, aprobar los casos de prueba y ** enviándolo** a Gradescope antes de las **11 p. m. del viernes de la misma semana en que se publicó**.

**Envío**: una vez que haya terminado, ejecute todas las celdas además de la última, seleccione Archivo > Guardar cuaderno y luego ejecute la celda final. El resultado contendrá un archivo zip que podrá utilizar para enviarlo en Gradescope.

Comencemos configurando las pruebas e importaciones ejecutando la siguiente celda.

```python
# No cambies esta celda; simplemente ejecútelo.
import numpy as np
from datascience import *
```

# 1. Revisión: Los componentes básicos del código Python

Los dos bloques de construcción del código Python son *expresiones* y *declaraciones*. Una **expresión** es un fragmento de código que

* es autónomo, lo que significa que tendría sentido escribirlo en una línea por sí solo, y
* generalmente se evalúa como un valor.


Aquí hay dos expresiones que se evalúan como 3:

     3
     5 - 2
    
Un tipo importante de expresión es la **expresión de llamada**. Una expresión de llamada comienza con el nombre de una función y va seguida de los argumentos de esa función entre paréntesis. La función devuelve algún valor, según sus argumentos. Algunas funciones matemáticas importantes se enumeran a continuación.

| Función | Descripción |
|----------|-------------------------------------- -------------------------|
| `abdominales` | Devuelve el valor absoluto de su argumento |
| `máximo` | Devuelve el máximo de todos sus argumentos |
| `min` | Devuelve el mínimo de todos sus argumentos |
| `poder` | Eleva su primer argumento a la potencia de su segundo argumento |
| `redondo` | Redondea su argumento al entero más cercano |

Aquí hay dos expresiones de llamada que se evalúan como 3:

     abdominales(2 - 5)
     max(ronda(2.8), min(pow(2, 10), -1 * pow(2, 10)))

La expresión `2 - 5` y las dos expresiones de llamada dadas anteriormente son ejemplos de **expresiones compuestas**, lo que significa que en realidad son combinaciones de varias expresiones más pequeñas. `2 - 5` combina las expresiones `2` y `5` mediante resta. En este caso, `2` y `5` se denominan **subexpresiones** porque son expresiones que forman parte de una expresión más grande.

Una **declaración** es una línea completa de código. Algunas declaraciones son sólo expresiones. Las expresiones enumeradas anteriormente son ejemplos.

Otras declaraciones *hacen que algo suceda* en lugar de *tener un valor*. Por ejemplo, una **declaración de asignación** asigna un valor a un nombre.

Una buena forma de pensar en esto es que estamos **evaluando el lado derecho** del signo igual y **asignándolo al lado izquierdo**. Aquí hay algunas declaraciones de asignación:
    
     altura = 1,3
     el_número_cinco = abs(-5)
     diferencia_altura_absoluta = abs(altura - 1.688)

Una idea importante en programación es que se pueden construir cosas grandes e interesantes combinando muchas cosas simples y poco interesantes. La clave para comprender un código complicado es dividirlo en sus componentes simples.

Por ejemplo, suceden muchas cosas en la última declaración anterior, pero en realidad es solo una combinación de algunas cosas. Esta imagen describe lo que está pasando.

<img src="https://github.com/darioblandon/materials-fa23/blob/main/lab/lab02/statement.png?raw=1">

**Pregunta 1.1.** En la siguiente celda, asigne el nombre `new_year` al **número mayor** entre los dos números siguientes:

* el **valor absoluto** de $2^{6}-2^{11}-2^{5} - 7$, y
* $5 \times 13 \times 31 + 7$.

Intente utilizar solo una declaración (una línea de código). Asegúrese de verificar su trabajo ejecutando la celda de prueba después.

```python
new_year = ...
new_year
```

```python
grader.check("q11")
```

Le pedimos que use una línea de código en la pregunta anterior porque solo involucra operaciones matemáticas. Sin embargo, las preguntas de programación más complicadas requerirán más pasos. No siempre es una buena idea agrupar estos pasos en una sola línea porque puede hacer que el código sea más difícil de leer y de depurar.

Una buena práctica de programación implica dividir el código en pasos más pequeños y utilizar nombres apropiados. ¡Tendrás mucha práctica en el resto de este curso!

# 2. Importación de código

La mayor parte de la programación implica un trabajo muy similar al trabajo que se ha realizado antes. Dado que escribir código lleva mucho tiempo, es bueno confiar en el código publicado por otros cuando sea posible. En lugar de copiar y pegar, Python nos permite **importar módulos**. Un módulo es un archivo con código Python que tiene variables y funciones definidas. Al importar un módulo, podemos usar su código en nuestro propio cuaderno.

Python incluye muchos módulos útiles que están a sólo una "importación" de distancia. Veremos el módulo "matemáticas" como primer ejemplo. El módulo "math" es extremadamente útil para calcular expresiones matemáticas en Python.

Supongamos que queremos calcular con mucha precisión el área de un círculo con un radio de 5 metros. Para eso, necesitamos la constante $\pi$, que es aproximadamente 3,14. Convenientemente, el módulo "math" tiene "pi" definido para nosotros. Ejecute la siguiente celda para importar el módulo matemático:

```python
import math
radius = 5
area_of_circle = radius**2 * math.pi
area_of_circle
```

En el código anterior, la línea "importar matemáticas" importa el módulo matemático. Esta declaración crea un módulo y luego asigna el nombre "matemáticas" a ese módulo. Ahora podemos acceder a cualquier variable o función definida en `math` escribiendo el nombre del módulo seguido de un punto, luego seguido del nombre de la variable o función que queremos.

     <nombre del módulo>.<nombre>

**Pregunta 2.1.** El módulo `math` también proporciona el nombre `e` para la base del logaritmo natural, que es aproximadamente 2,71. Calcule $e^{\pi}-\pi$, dándole el nombre `near_twenty`.

*Recuerda: ¡También puedes acceder a `pi` desde el módulo `math`!*

```python
near_twenty = ...
near_twenty
```

```python
grader.check("q21")
```

## 2.1. Acceso a funciones

En la pregunta anterior, accedió a variables dentro del módulo "matemáticas".

**Los módulos** también definen **funciones**. Por ejemplo, "matemáticas" proporciona el nombre "suelo" para la función suelo. Habiendo importado ya "matemáticas", podemos escribir "math.floor(7.5)" para calcular el piso de 7.5. (Tenga en cuenta que la función piso devuelve el número entero más grande menor o igual a un número determinado).

**Pregunta 2.1.1.** Calcule el piso de pi usando `floor` y `pi` del módulo `math`. Asigne al resultado el nombre `floor_of_pi`.

```python
floor_of_pi = ...
floor_of_pi
```

```python
grader.check("q211")
```

Para su referencia, a continuación se muestran algunos ejemplos más de funciones del módulo "math".

Observe cómo diferentes funciones toman diferentes números de argumentos. A menudo, la [documentación](https://docs.python.org/3/library/math.html) del módulo proporcionará información sobre cuántos argumentos se requieren para cada función.

*Sugerencia: si presiona `shift+tab` mientras está al lado de la llamada de función, aparecerá la documentación para esa función.*

```python
# Calcular logaritmos (el logaritmo de 8 en base 2).
# El resultado es 3 porque 2 elevado a 3 es 8.
math.log(8, 2)
```

```python
# Calcular raíces cuadradas.
math.sqrt(5)
```

Hay varias formas de importar y acceder a código desde fuentes externas. El método que usamos anteriormente, `<importar <nombre_módulo>`, importa el módulo completo y requiere que usemos `<nombre_módulo>.<nombre>` para acceder a su código.

También podemos importar una constante o función específica en lugar del módulo completo. Tenga en cuenta que no es necesario utilizar el nombre del módulo de antemano para hacer referencia a ese valor en particular. Sin embargo, debe tener cuidado al reasignar los nombres de las constantes o funciones a otros valores.

```python
# Importando solo cos y pi desde matemáticas.
# No tenemos que usar "matemáticas" delante de cos o pi
from math import cos, pi
print(cos(pi))

# Sin embargo, tenemos que usarlo delante de otras funciones matemáticas.
math.log(pi)
```

O podemos importar cada función y valor de todo el módulo.

```python
# Por último, podemos importar todo desde matemáticas usando el *
# Una vez más, no tenemos que usar "matemáticas". antemano
from math import *
log(pi)
```

No se preocupe demasiado por qué tipo de importación utilizar. A menudo es una elección de estilo de codificación que queda en manos de cada programador. En este curso, siempre importará los módulos necesarios cuando ejecute la celda de configuración (como la primera celda de código en esta práctica de laboratorio).

¡Pasemos a practicar algunas de las operaciones de mesa que aprendiste en la conferencia!

# 3. Operaciones de mesa

La tabla `farmers_markets.csv` contiene datos sobre los mercados de agricultores en los Estados Unidos (datos asociados con [el USDA](https://apps.ams.usda.gov/FarmersMarketsExport/ExcelExport.aspx)). Cada fila representa uno de esos mercados.

Ejecute la siguiente celda para cargar la tabla `farmers_markets`. No habrá resultados; no se espera ningún resultado ya que la celda contiene una declaración de asignación. Una declaración de asignación no produce ningún resultado (no produce ningún valor).

```python
# Simplemente ejecuta esta celda

farmers_markets = Table.read_table('farmers_markets.csv')
```

Examinemos nuestra tabla para ver qué datos contiene.

**Pregunta 3.1.** Utilice el método `show` para mostrar las primeras 5 filas de `farmers_markets`.

*Nota:* Los términos "método" y "función" técnicamente no son lo mismo, pero para los propósitos de este curso, los usaremos indistintamente.

**Pista:** `tbl.show(3)` mostrará las primeras 3 filas de la tabla denominada `tbl`. Además, asegúrese de no llamar a `.show()` sin un argumento, ya que esto bloqueará su kernel.


```python
...
```

Observe que faltan algunos de los valores en esta tabla, como se indica con "nan". Esto significa que el valor no está disponible (por ejemplo, si no conocemos la dirección del mercado) o no es aplicable (por ejemplo, si el mercado no tiene una dirección). ¡También notarás que la tabla tiene una gran cantidad de columnas!

### `num_columnas`

La propiedad de la tabla `num_columns` devuelve el número de columnas de una tabla. (Una "propiedad" es solo un método al que no es necesario llamar agregando paréntesis).

Llamada de ejemplo: `tbl.num_columns` devolverá el número de columnas en una tabla llamada `tbl`

**Pregunta 3.2.** Utilice `num_columns` para encontrar el número de columnas en nuestro conjunto de datos de mercados de agricultores.

Asigne el número de columnas a `num_farmers_markets_columns`.


```python
num_farmers_markets_columns = ...
print("The table has", num_farmers_markets_columns, "columns in it!")
```

```python
grader.check("q32")
```

### `num_rows`

De manera similar, la propiedad `num_rows` te indica cuántas filas hay en una tabla.

```python
# Just run this cell

num_farmers_markets_rows = farmers_markets.num_rows
print("The table has", num_farmers_markets_rows, "rows in it!")
```

### `select`

La mayoría de las columnas tratan sobre productos concretos: si el mercado vende tofu, comida para mascotas, etc. Si no estamos interesados en esa información, la tabla resulta difícil de leer. Esto surge más de lo que se podría pensar, porque es posible que las personas que recopilan y publican datos no sepan de antemano qué querrán hacer con ellos.

En tales situaciones, podemos usar el método de tabla "seleccionar" para elegir solo las columnas que queremos en una tabla en particular. Se necesitan muchos argumentos. Cada uno debe ser el nombre de una columna de la tabla. Devuelve una nueva tabla con solo esas columnas. Las columnas están en el orden *en el que fueron enumeradas como argumentos*.

Por ejemplo, el valor de `farmers_markets.select("MarketName", "State")` es una tabla con solo el nombre y el estado de cada mercado de agricultores en `farmers_markets`.



**Pregunta 3.3.** Utilice `select` para crear una tabla con solo el nombre, ciudad, estado, latitud (`y`) y longitud (`x`) de cada mercado. Llame a esa nueva tabla "farmers_markets_locations".

*Sugerencia:* Asegúrese de ser exacto al usar nombres de columnas con `select`; ¡Compruebe dos veces las mayúsculas!

```python
farmers_markets_locations = ...
farmers_markets_locations
```

```python
grader.check("q33")
```

### `drop`

`drop` tiene el mismo propósito que `select`, pero elimina las columnas que usted proporciona en lugar de las que no proporciona. Al igual que "seleccionar", "soltar" devuelve una nueva tabla.

**Pregunta 3.4.** Supongamos que simplemente no desea las columnas `FMID` y `updateTime` en `farmers_markets`. Cree una tabla que sea una copia de `farmers_markets` pero que no incluya esas columnas. Llame a esa tabla "farmers_markets_ without_fmid".

```python
farmers_markets_without_fmid = ...
farmers_markets_without_fmid
```

```python
grader.check("q34")
```

Ahora supongamos que queremos responder algunas preguntas sobre los mercados de agricultores en Estados Unidos. Por ejemplo, ¿qué mercado(s) tiene la mayor longitud (dada por la columna "x")?

Para responder a esto, ordenaremos `farmers_markets_locations` por longitud.

```python
farmers_markets_locations.sort('x')
```

Vaya, eso no respondió nuestra pregunta porque ordenamos de menor a mayor longitud. Para ver las longitudes más grandes, tendremos que ordenar en orden inverso.

```python
farmers_markets_locations.sort('x', descending=True)
```

(The `descending=True` se llama *argumento opcional*. Tiene un valor predeterminado de `False`, por lo que cuando le dices explícitamente a la función `descending=True`, la función ordenará en orden descendente).

### `sort`

Algunos detalles sobre la clasificación:

1. El primer argumento para "ordenar" es el nombre de una columna por la que ordenar.
2. Si la columna tiene texto, "ordenar" ordenará alfabéticamente; si la columna tiene números, se ordenará numéricamente, ambos en orden ascendente de forma predeterminada.
3. El valor de `farmers_markets_locations.sort("x")` es una *copia* de `farmers_markets_locations`; la tabla `farmers_markets_locations` no se modifica. Por ejemplo, si llamamos a `farmers_markets_locations.sort("x")`, ejecutar `farmers_markets_locations` por sí solo aún devolvería la tabla sin clasificar.
4. Las filas siempre permanecen juntas cuando se ordena una tabla. No tendría sentido ordenar solo una columna y dejar las otras columnas como están. Por ejemplo, en este caso, si ordenáramos solo la columna "x", todos los mercados de agricultores terminarían con las longitudes incorrectas.

**Pregunta 3.5.** Cree una versión de `farmers_markets_locations` que esté ordenada por **latitud (`y`)**, con las latitudes más grandes primero. Llámelo "farmers_markets_locations_by_latitude".


```python
farmers_markets_locations_by_latitude = ...
farmers_markets_locations_by_latitude
```

```python
grader.check("q35")
```

Ahora digamos que queremos una tabla de todos los mercados de agricultores de California. La clasificación no nos ayudará mucho aquí porque California está más cerca del centro del conjunto de datos.

En su lugar, utilizamos el método de la tabla. `where`.

```python
california_farmers_markets = farmers_markets_locations.where('State', are.equal_to('California'))
california_farmers_markets
```

Ignora la sintaxis por el momento. En su lugar, intenta leer esa línea así:

> Asigne el nombre **`california_farmers_markets`** a una tabla cuyas filas son las filas de la tabla **`farmers_markets_locations`** **`donde`** están los **`"Estados"`** **` are equal to` `"California"`**.

### `where`

Ahora profundicemos un poco más en los detalles. `where` toma 2 argumentos:

1. El nombre de una columna. `where` busca filas donde los valores de esa columna cumplen algún criterio.
2. Un predicado que describe el criterio que debe cumplir la columna.

El predicado en el ejemplo anterior llamó a la función `are.equal_to` con el valor que queríamos, 'California'. Pronto veremos otros predicados.

`where` devuelve una tabla que es una copia de la tabla original, pero **con solo las filas que cumplen el predicado dado**.

**Pregunta 3.6.** Utilice `california_farmers_markets` para crear una tabla llamada `berkeley_markets` que contenga los mercados de agricultores en Berkeley, California.

```python
berkeley_markets = ...
berkeley_markets
```

```python
grader.check("q36")
```

Hasta ahora solo hemos estado usando `dónde` con el predicado que requiere encontrar que los valores en una columna sean *exactamente* iguales a un cierto valor. Sin embargo, hay muchos otros predicados. Hemos enumerado algunos de los más comunes a continuación, pero si desea una lista más completa, consulte la sección de predicados en la [hoja de referencia de Python] (https://www.data8.org/fa23/reference/#tablewhere-predicates).

|Predicate|Example|Result|
|-|-|-|
|`are.equal_to`|`are.equal_to(50)`|Find rows with values equal to 50|
|`are.not_equal_to`|`are.not_equal_to(50)`|Find rows with values not equal to 50|
|`are.above`|`are.above(50)`|Find rows with values above (and not equal to) 50|
|`are.above_or_equal_to`|`are.above_or_equal_to(50)`|Find rows with values above 50 or equal to 50|
|`are.below`|`are.below(50)`|Find rows with values below 50|
|`are.between`|`are.between(2, 10)`|Find rows with values above or equal to 2 and below 10|
|`are.between_or_equal_to`|`are.between_or_equal_to(2, 10)`|Find rows with values above or equal to 2 and below or equal to 10|

## 4. Analizando un conjunto de datos

Ahora que está familiarizado con las operaciones de tablas, ¡respondamos una pregunta interesante sobre un conjunto de datos!

Ejecute la siguiente celda para cargar la tabla `imdb`. Contiene información sobre las 250 películas mejor calificadas en IMDb.

```python
# Simplemente ejecuta esta celda

imdb = Table.read_table('imdb.csv')
imdb
```

A menudo, queremos realizar múltiples operaciones (clasificar, filtrar u otras) para convertir una tabla que tenemos en algo más útil. Puede realizar estas operaciones una por una, por ejemplo.

```
first_step = original_tbl.where(“col1”, are.equal_to(12))
second_step = first_step.sort(‘col2’, descending=True)
```
Sin embargo, dado que el valor de la expresión `original_tbl.where(“col1”, are.equal_to(12))` es en sí mismo una tabla, puedes simplemente llamar a un método de tabla en él:
```
original_tbl.where(“col1”, are.equal_to(12)).sort(‘col2’, descending=True)
```
Debe organizar su trabajo de la manera que tenga más sentido para usted, utilizando nombres informativos para las tablas intermedias que cree.

**Pregunta 4.1.** Cree una tabla de películas estrenadas entre 2010 y 2015 (inclusive) con calificaciones superiores a 8. La tabla solo debe contener las columnas `Title` and `Rating`, **en ese orden**.

Asigne la tabla al nombre `above_eight`.

*Pista:* Piensa en los pasos que debes seguir e intenta ordenarlos en un orden que tenga sentido. ¡Siéntete libre de crear tablas intermedias para cada paso, pero asegúrate de asignar a tu tabla final el nombre "above_eight"!


```python
above_eight = ...
above_eight
```

```python
grader.check("q41")
```

**Pregunta 4.2.** Utilice `num_rows` (y aritmética) para encontrar la *el numero* de películas en el conjunto de datos que se estrenaron entre 1900 y 1999, y la *el numero* de películas en el conjunto de datos que se estrenaron en el año 2000 o después.

Asigne "num_en_siglo_20" a el número de películas en el conjunto de datos que se estrenaron entre 1900 y 1999, y "num_en_siglo_21" a el número de películas en el conjunto de datos que se estrenaron en el año 2000 o después.

*Pista:* El *número* de películas estrenadas en el siglo XX es el *número* de películas estrenadas en el siglo XX, dividida por el *número total* de películas.


```python
num_movies_in_dataset = ...
num_in_20th_century = ...
num_in_21st_century = ...
proportion_in_20th_century = ...
proportion_in_21st_century = ...
print("Proportion in 20th century:", proportion_in_20th_century)
print("Proportion in 21st century:", proportion_in_21st_century)
```

```python
grader.check("q42")
```

## 5. Resumen

Para su referencia, aquí hay una tabla de todas las funciones y métodos que vimos en esta práctica de laboratorio. ¡Aprenderemos más métodos para agregar a esta tabla la próxima semana!

|Nombre|Ejemplo|Propósito|
|-|-|-|
|`sort`|`tbl.sort("N")`|Crear una copia de una tabla ordenada por los valores de una columna|
|`where`|`tbl.where("N", are.above(2))`|Crea una copia de una tabla con solo las filas que coinciden con algún *predicado*|
|`num_rows`|`tbl.num_rows`|Calcular el número de filas en una tabla|
|`num_columns`|`tbl.num_columns`|Calcular el número de columnas en una tabla|
|`select`|`tbl.select("N")`|Crear una copia de una tabla con solo algunas de las columnas|
|`drop`|`tbl.drop("N")`|Crear una copia de una tabla sin algunas de las columnas|

**¡Lego** quiere felicitarte por terminar el Laboratorio 2!

<img src="https://github.com/darioblandon/materials-fa23/blob/main/lab/lab02/lego.png?raw=1" alt="Picture of brown cavalier poodle dog" width="300"/>


---

¡Terminaste con el laboratorio!

**Información importante sobre el envío:**
- **Ejecuta todas las pruebas** y verifica que todas pasan
- **Guardar** desde el menú **Archivo**
- **Ejecute la celda final para generar el archivo zip**
- **Haga clic en el enlace para descargar el archivo zip**
- Luego, ve a [Gradescope](https://www.gradescope.com/courses/576157) y envía el archivo zip a la tarea correspondiente. El nombre de esta tarea es "Lab XX Autograder", donde XX es el número de laboratorio: 01, 02, 03, etc.

- Si terminas temprano en el laboratorio regular, **pídele a uno de los miembros del personal que te marque**.

**Es su responsabilidad asegurarse de que su trabajo esté guardado antes de ejecutar la última celda.**


## Envío

Asegúrese de haber ejecutado todas las celdas de su cuaderno en orden antes de ejecutar la siguiente celda, para que todas las imágenes/gráficos aparezcan en el resultado. La siguiente celda generará un archivo zip para que lo envíe. **¡Guarde antes de exportar!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q11 (points: None)
```
>>> new_year
2023
```


Tests for q21 (points: None)
```
>>> round(near_twenty, 8)
19.99909998
```


Tests for q211 (points: None)
```
>>> import math
>>> floor_of_pi == 3
True
```


Tests for q32 (points: None)
```
>>> num_farmers_markets_columns == 59
True
```


Tests for q33 (points: None)
```
>>> sorted(farmers_markets_locations.labels)
['MarketName', 'State', 'city', 'x', 'y']
```
```
>>> farmers_markets_locations.num_rows == 8546
True
```


Tests for q34 (points: None)
```
>>> farmers_markets_without_fmid.num_columns
57
```
```
>>> print(sorted(farmers_markets_without_fmid.labels))
['Bakedgoods', 'Beans', 'Cheese', 'Coffee', 'County', 'Crafts', 'Credit', 'Eggs', 'Facebook', 'Flowers', 'Fruits', 'Grains', 'Herbs', 'Honey', 'Jams', 'Juices', 'Location', 'Maple', 'MarketName', 'Meat', 'Mushrooms', 'Nursery', 'Nuts', 'Organic', 'OtherMedia', 'PetFood', 'Plants', 'Poultry', 'Prepared', 'SFMNP', 'SNAP', 'Seafood', 'Season1Date', 'Season1Time', 'Season2Date', 'Season2Time', 'Season3Date', 'Season3Time', 'Season4Date', 'Season4Time', 'Soap', 'State', 'Tofu', 'Trees', 'Twitter', 'Vegetables', 'WIC', 'WICcash', 'Website', 'WildHarvested', 'Wine', 'Youtube', 'city', 'street', 'x', 'y', 'zip']
```


Tests for q35 (points: None)
```
>>> type(farmers_markets_locations_by_latitude) == tables.Table
True
```
```
>>> # HINT: Check the order of your table. 
>>> np.isclose(round(farmers_markets_locations_by_latitude.column('y').item(1), 4), round(64.8459, 4))
True
```


Tests for q36 (points: None)
```
>>> berkeley_markets.num_rows == 3
True
```
```
>>> list(berkeley_markets.column('city')) == ['Berkeley', 'Berkeley', 'Berkeley']
True
```


Tests for q41 (points: None)
```
>>> type(above_eight) == tables.Table
True
```
```
>>> above_eight.num_rows == 20
True
```
```
>>> # Make sure your columns are in the right order!
>>> # First column should be 'Title', second column should be 'Rating'
>>> print(above_eight.sort(0).take([17]))
Title       | Rating
Toy Story 3 | 8.3
```


Tests for q42 (points: None)
```
>>> proportion_in_20th_century == 0.684
True
```
```
>>> proportion_in_21st_century == 0.316
True
```
