<a href="https://colab.research.google.com/github/darioblandon/materials-fa23/blob/main/lab/lab06/lab06.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab06.ipynb")
```

<p><img alt="Colaboratory logo" height="140px" src="https://upload.wikimedia.org/wikipedia/commons/archive/f/fb/20161010213812%21Escudo-UdeA.svg" align="left" hspace="10px" vspace="0px"></p>

<h1> Fundamentos en ciencias de datos</h1>



# Laboratorio 6: Examinando el toque terapéutico

¡Bienvenidos al Laboratorio 6!


**Obtener ayuda en el laboratorio**: Siempre que te sientas estancado o necesites más aclaraciones, busca un GSI o un tutor, ¡y estarán encantados de ayudarte!

Estas son las políticas para obtener crédito completo:

1. Para los estudiantes inscritos en <span style="color: #BC412B">**Regular Labs**</span> en persona, recibirán crédito de laboratorio al **asistir** a la sección de laboratorio y lograr un progreso significativo en el cuaderno de laboratorio. **<span style="text-decoration: underline">debes</span>** enviar esta tarea a Gradescope.

2. Los estudiantes inscritos en <span style="color: #04888B">**Laboratorio de autoservicio**</span> recibirán crédito de laboratorio al **completar** el cuaderno, aprobar los casos de prueba y **enviándolo** a Gradescope antes de las **11:59 p. m. de la fecha de entrega**.

**Envío**: una vez que haya terminado, ejecute todas las celdas además de la última, seleccione Archivo > Guardar cuaderno y luego ejecute la celda final. El resultado contendrá un archivo zip que podrá utilizar para enviarlo en Gradescope.

Comencemos configurando las pruebas y las importaciones ejecutando la siguiente celda.

Después de una introducción tan extensa a la programación para la ciencia de datos, finalmente pasamos a la sección del curso donde podemos aplicar nuestras nuevas habilidades para responder preguntas reales.

En este laboratorio, usaremos técnicas de prueba que se introdujeron en la conferencia para probar la idea del toque terapéutico, la idea de que algún practicante puede sentir y masajear su campo de energía humano.


```python
# Ejecute esta celda, pero no la cambie.

# Estas líneas importan los módulos Numpy y Datascience.
import numpy as np
from datascience import *

# Estas líneas hacen algo de magia en la trama
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)
from matplotlib import patches
from ipywidgets import interact, interactive, fixed
import ipywidgets as widgets

```

## 1. ¿Qué es el Toque Terapéutico?

El Toque Terapéutico (TT) es la idea de que todos pueden sentir el Campo de Energía Humana (HEF) alrededor de los individuos. Quienes practican TT han descrito los HEF de diferentes personas como "cálidos como gelatina" y "táctiles como caramelo".

La TT fue una técnica popular utilizada a lo largo del siglo XX que se promocionó como una excelente manera de equilibrar la salud de una persona. Ciertos profesionales afirman que tienen la capacidad de sentir el HEF y masajearlo para promover la salud y la relajación de las personas.

### Emily Rosa

[Emily Rosa](https://en.wikipedia.org/wiki/Emily_Rosa) era una estudiante de 4º de primaria que estaba muy familiarizada con el mundo de la TT, gracias a sus padres, que eran a la vez médicos y escépticos de la TT.

Para su proyecto de feria de ciencias de cuarto grado, Emily decidió probar si los practicantes de TT realmente podían interactuar con el HEF de una persona. Más tarde publicó su trabajo en TT, convirtiéndose en la persona más joven en publicar un artículo de investigación en una revista médica revisada por pares.

### El experimento de Emily

El experimento de Emily fue limpio, simple y efectivo. Debido a las ocupaciones de sus padres en el campo médico, tenía amplio acceso a personas que afirmaban ser practicantes de TT.

Emily tomó 21 practicantes de TT y los usó para su experimento científico. Tomaría a un practicante de TT y le pediría que extendiera sus manos a través de una pantalla (a través de la cual no puede ver). Emily estaría del otro lado y lanzaría una moneda justa. Dependiendo de cómo cayera la moneda, extendía su mano izquierda o su mano derecha. El practicante de TT tendría entonces que responder qué mano extendió Emily. Si un practicante realmente pudiera interactuar con el HEF de una persona, se esperaría que respondiera correctamente.

En total, a través de 210 muestras, el practicante eligió la mano correcta el 44% de las veces.

El objetivo principal de Emily aquí era probar si las conjeturas de los practicantes de TT eran aleatorias, como el lanzamiento de una moneda. En la mayoría de los experimentos médicos, esta es la norma. **Queremos probar si el tratamiento tiene un efecto o no, *no* si el tratamiento realmente funciona o no.**

Ahora comenzaremos a formular este experimento en términos de la terminología que aprendimos en este curso.

**Pregunta 1.1**: Describa el [modelo] de Emily (https://inferentialthinking.com/chapters/11/1/Assessing_a_Model.html) sobre la probabilidad de que los practicantes de TT elijan la mano correcta. ¿Qué modelo alternativo pretende desacreditar su modelo?

Si puede, consulte con sus compañeros, en el foro de discusión o con los asistentes técnicos/inteligentes de su laboratorio para llegar a una conclusión.

_Escriba su respuesta aquí, reemplazando este texto._

**Pregunta 1.2:** Recuerde que el practicante obtuvo la respuesta correcta el 44% (0,44) de las veces. Según el modelo de Emily, en promedio, ¿qué proporción de veces esperamos que el practicante adivine la mano correcta? Asegúrate de que tu respuesta sea un número entre 0 y 1.

```python
expected_proportion_correct = ...
expected_proportion_correct
```

```python
grader.check("q1_2")
```

El objetivo ahora es ver si nuestra desviación de esta proporción esperada de respuestas correctas se debe a algo más que al azar.

**Pregunta 1.3:** Generalmente utilizamos una estadística para ayudar a determinar hacia qué modelo apunta la evidencia. ¿Qué estadística podemos usar para comparar los resultados del modelo de Emily con lo observado? Asigne `valid_stat` a una matriz de números enteros que representan estadísticas de prueba que Emily puede usar:

1. La diferencia entre el porcentaje correcto esperado y el porcentaje correcto real
2. La diferencia absoluta entre el porcentaje correcto esperado y el porcentaje correcto real
3. La suma del porcentaje correcto esperado y el porcentaje correcto real

**NOTA:** ¡Asegúrese de usar `make_array` para crear su matriz de números enteros!

> *Sugerencia*: ¿Cuál debería ser el dominio (posibles valores de x) para la distribución de nuestras estadísticas de prueba?

```python
valid_stat = ...
valid_stat
```

```python
grader.check("q1_3")
```

**Pregunta 1.4:** ¿Por qué la estadística de la Pregunta 1.3 es la opción adecuada para comparar resultados en el experimento de Emily? ¿Cómo se relaciona con los modelos que definió en la Pregunta 1.1?

_Escriba su respuesta aquí, reemplazando este texto._

**Pregunta 1.5:** Defina la función "estadística" que toma una proporción esperada y una proporción real, y devuelve el valor de la estadística elegida en la Pregunta 1.3. Suponga que el argumento toma proporciones, pero devuelva su respuesta como un porcentaje.

*Pista:* Recuerde que estamos pidiendo un **porcentaje**, no una proporción.

```python
def statistic(expected_prop, actual_prop):
    ...
```

```python
grader.check("q1_5")
```

**Pregunta 1.6:** Utilice la función recién definida para calcular la estadística observada en el experimento de Emily.

```python
observed_statistic = ...
observed_statistic
```

```python
grader.check("q1_6")
```

**¿Esta estadística observada es consistente con lo que esperamos ver bajo el modelo de Emily?**

Para responder a esta pregunta, debemos simular el experimento como si el modelo de Emily fuera correcto y calcular nuestra estadística para cada simulación.

### `proporciones_de_muestra`

`sample_proportions` se puede utilizar para realizar muestras aleatorias de varias categorías cuando se conoce la proporción de puntos de datos que se espera que caigan en cada categoría. `sample_proportions` toma dos argumentos: el tamaño de la muestra y una serie de proporciones correspondientes a cada categoría de la población (debe sumar 1).

Considere lanzar una moneda justa, donde los dos resultados (la moneda sale cara y la moneda sale cruz) ocurren con la misma probabilidad. Esperamos que la mitad de todos los lanzamientos de monedas arrojen cara y la otra mitad de todos los lanzamientos de monedas arrojen cruz.

Ejecute la siguiente celda para ver la simulación de 10 lanzamientos de una moneda justa. Sea el primer elemento de "coin_proportions" la proporción de caras y el segundo elemento de "coin_proportions" la proporción de cruces.

*Observa lo que sucede cuando ejecutas esta celda varias veces: la proporción de lanzamientos de monedas que salen cara y cruz parece cambiar, ya que estás simulando lanzar 10 monedas cada vez.*

```python
coin_proportions = make_array(0.5, 0.5)
ten_flips = sample_proportions(10, coin_proportions)
ten_flips
```

`sample_proportions` devuelve una matriz que tiene la misma longitud que la matriz de proporciones que se pasa. Contiene la proporción de cada categoría que aparece en la muestra.

En nuestro ejemplo, el primer elemento de "ten_flips" es la proporción simulada de caras y el segundo elemento de "ten_flips" es la proporción simulada de cruces.

```python
simulated_proportion_heads = ten_flips.item(0)
simulated_proportion_tails = ten_flips.item(1)

print("In our simulation, " + str(simulated_proportion_heads) + " of flips were heads and " \
      + str(simulated_proportion_tails) + " of flips were tails.")
```

**Pregunta 1.7:** Para comenzar a simular, debemos comenzar creando una representación del modelo de Emily para usarla en nuestra simulación. Esta será una matriz con dos elementos. El primer elemento debe ser la proporción de veces que un practicante de TT elige la mano correcta, asumiendo que el modelo de Emily era correcto. El segundo elemento debe ser la proporción de veces, bajo el mismo supuesto, que el practicante de TT elige la mano incorrecta. Asigne `model_proportions` a esta matriz.

Después de esto, podemos simular 210 elecciones de manos, como Emily evaluó en la vida real, y encontrar una estadística única para resumir este ejemplo de la simulación. Utilice la función `sample_proportions` y asigne la **proporción de elecciones de manos correctas** (de 210) a `simulation_proportion_correct`. Por último, utilice su función "estadística" para asignar "una_estadística" al valor de la estadística para esta simulación.

*Sugerencia:* El uso de `sample_proportions` se puede encontrar en la [Referencia de Python](https://www.data8.org/fa23/reference/).


```python
# Esto guarda el estado aleatorio de nuestro código para que podamos
# generamos los mismos números cada vez que ejecutamos el código.
# Por favor no cambie esta línea.

np.random.seed(16)

model_proportions = ...
simulation_proportion_correct = ...
one_statistic = ...
one_statistic
```

```python
grader.check("q1_7")
```

**Pregunta 1.8:** Veamos ahora cómo es realmente la distribución de las estadísticas según el modelo de Emily.

Defina la función `simulación_y_estadística` para incluir la matriz `model_proportions` y la proporción esperada de veces que un practicante de TT adivinaría una mano correctamente según el modelo de Emily. La función debería simular que Emily realiza el experimento 210 veces y devolver la estadística de esta simulación.

*Pista:* Esto debería seguir el mismo patrón que el código que hiciste en el problema anterior.

_Escriba su respuesta aquí, reemplazando este texto._

```python
def simulation_and_statistic(model_proportions, expected_proportion_correct):
    '''Simulates 210 TT hand choices under Emily’s model.
    Returns one statistic from the simulation.'''
    ...
```

```python
num_repetitions = 1000

simulated_statistics = ...

for ... in ...:
    ...
```

```python
grader.check("q1_8")
```

Veamos la distribución de las estadísticas simuladas bajo el modelo de Emily y comparemos visualmente dónde se encuentra la estadística observada en relación con las estadísticas simuladas.

```python
t = Table().with_column('Simulated Statistics', simulated_statistics)
t.hist()
plt.scatter(observed_statistic, 0, color='red', s=100, zorder=2);
```

Podemos argumentar visualmente si creemos que la estadística observada es consistente con el modelo de Emily. En este caso, dado que valores mayores de la estadística de prueba sugieren el modelo alternativo (donde la probabilidad de adivinar la mano correcta es algo distinta del 50%), podemos formalizar nuestro análisis encontrando qué proporción de estadísticas simuladas fueron tan grandes o mayores que nuestras observadas. estadística de prueba (el área a la derecha o a la derecha de la estadística de prueba observada). Si esta área es lo suficientemente pequeña, declararemos que los datos observados son inconsistentes con nuestro modelo simulado. Aquí está el [enlace](https://inferentialthinking.com/chapters/11/1/Assessing_a_Model.html) a la sección del libro de texto.

**Pregunta 1.9:** Calcule la proporción de estadísticas simuladas en la Pregunta 1.8 mayor o igual a la estadística observada.

*Pista:* El uso de `np.count_nonzero` se puede encontrar [aquí](http://data8.org/fa23/reference).

```python
proportion_greater_or_equal = ...
proportion_greater_or_equal
```

```python
grader.check("q1_9")
```

Por convención, a menudo comparamos la proporción que acabamos de calcular con 0,05. Si la proporción de estadísticas simuladas mayores o iguales a las estadísticas observadas es suficientemente pequeña (menor o igual a 0,05), entonces esto es evidencia en contra del modelo de Emily. Conceptualmente, se puede pensar en este caso como el caso en el que menos del 5% de los valores simulados están tan o más lejos de lo que esperábamos. Si este no es el caso, no tenemos ningún motivo para dudar del modelo de Emily.

Esto debería ayudarte a sacar tus propias conclusiones sobre el experimento de Emily Rosa.

El contacto terapéutico dejó de utilizarse después de este experimento, que finalmente fue aceptado en una de las principales revistas médicas. Los practicantes de TT respondieron y acusaron a Emily y su familia de alterar los resultados, mientras que algunos afirmaron que el mal humor espiritual de Emily hacia el contacto terapéutico hacía difícil leer su HEF. Sea lo que sea, el experimento de Emily es un ejemplo clásico de cómo cualquiera, con los recursos adecuados, puede probar lo que quiera.

**Pregunta 1.10:** Ahora, tómate un tiempo para reflexionar sobre las preguntas a continuación y luego discútelas con tus compañeros o echa un vistazo a las discusiones en la publicación educativa de este laboratorio.

1. ¿Son los datos más consistentes con el modelo de Emily (los practicantes adivinaban al azar)?
2. ¿Qué significa esto en términos del experimento de Emily? ¿Las respuestas de los practicantes de TT siguen un modelo de probabilidad par o hay algo más en juego?

¿Hablaste con tus compañeros o miraste el foro de discusión? (Verdadero Falso)


```python
peer_talk = ...
peer_talk
```

```python
grader.check("q1_10")
```

## ¡CASI LLEGAMOS!

<img src="https://github.com/darioblandon/materials-fa23/blob/main/lab/lab06/teddy.jpeg?raw=1" alt="Cute dog standing in a hallway with tile floors" width="300"/>

**Teddy** quiere hacerte saber que eres genial.


---

¡Terminaste con el laboratorio!

**Información importante sobre el envío:**
- **Ejecuta todas las pruebas** y verifica que todas pasan
- **Guardar** desde el menú **Archivo**
- **Ejecute la celda final para generar el archivo zip**
- **Haga clic en el enlace para descargar el archivo zip**
- Luego, ve a [Gradescope](https://www.gradescope.com/courses/576157) y envía el archivo zip a la tarea correspondiente. El nombre de esta tarea es "Lab XX Autograder", donde XX es el número de laboratorio: 01, 02, 03, etc.

- Si terminas temprano en el laboratorio regular, **pídele a uno de los miembros del personal que te marque**.

**Es su responsabilidad asegurarse de que su trabajo esté guardado antes de ejecutar la última celda.**

## Envío

Asegúrese de haber ejecutado todas las celdas de su cuaderno en orden antes de ejecutar la siguiente celda, para que todas las imágenes/gráficos aparezcan en el resultado. La siguiente celda generará un archivo zip para que lo envíe. **¡Guarde antes de exportar!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_10 (points: None)
```
>>> peer_talk == True
True
```


Tests for q1_2 (points: None)
```
>>> 0 <= expected_proportion_correct <= 1
True
```
```
>>> expected_proportion_correct == 0.5
True
```


Tests for q1_3 (points: None)
```
>>> type(valid_stat) == np.ndarray
True
```
```
>>> int(sum(valid_stat))
2
```


Tests for q1_5 (points: None)
```
>>> int(round(statistic(.5,.5) + statistic(.4,.1),1))
30
```
```
>>> int(statistic(.4,.1) - statistic(.1,.4))
0
```


Tests for q1_6 (points: None)
```
>>> int(round(observed_statistic,2))
6
```


Tests for q1_7 (points: None)
```
>>> len(model_proportions) % 2 == 0
True
```
```
>>> len(np.unique(model_proportions))
1
```
```
>>> sum(model_proportions) == 1
True
```
```
>>> type(simulation_proportion_correct) == float
True
```
```
>>> round(simulation_proportion_correct, 2)
0.49
```
```
>>> round(one_statistic, 2) - 0.95 < 0.05
True
```


Tests for q1_8 (points: None)
```
>>> len(simulated_statistics) == 1000
True
```
```
>>> np.all(simulated_statistics <= 30)
True
```
```
>>> np.all(simulated_statistics >= 0)
True
```
```
>>> 0 <= simulation_and_statistic(model_proportions, expected_proportion_correct) <= 25
True
```


Tests for q1_9 (points: None)
```
>>> 0 <= proportion_greater_or_equal <= 1
True
```
```
>>> proportion_greater_or_equal*1000 == np.count_nonzero(simulated_statistics >= observed_statistic)
True
```
