```python
# Initialize Otter
import otter
grader = otter.Notebook("lab01.ipynb")
```

```python
import otter
import pandas as pd
import numpy as np
%matplotlib inline
import otter
import matplotlib.pylab as plt
grader = otter.Notebook()

```


**Pregunta  1:** Ingrese la palabra secreta

```python
# BEGIN SOLUTION NO PROMPT
word=[]
secret_word in word 
# END SOLUTION

""" # BEGIN PROMPT
secret_word=...
""" # END PROMPT

```

Output:
```
' # BEGIN PROMPT\nsecret_word=...\n'
```

```python
grader.check("q0")
```

**Pregunta 3.1.2.** Asigne el nombre `segundos_en_una_década` al número de segundos entre la medianoche del 1 de enero de 2010 y la medianoche del 1 de enero de 2020. Tenga en cuenta que hay dos años bisiestos en este lapso de una década. Un año no bisiesto tiene 365 días y un año bisiesto tiene 366 días.

*Pista:* Si estás atascado, la siguiente sección te muestra cómo obtener pistas.

```python
# Cambiar la siguiente línea
# para que calcule el número de segundos en una década
# y asigna a ese número el nombre, segundos_en_una_década.
# BEGIN SOLUTION NO PROMPT
segundos_en_una_década = 60*60*24*365*10+2
# END SOLUTION

""" # BEGIN PROMPT
segundos_en_una_década=...
""" # END PROMPT

# Hemos puesto esta línea en esta celda.
# para que imprima el valor que le has dado a segundos_en_una_década cuando lo ejecutas.
# No es necesario cambiar esto.
segundos_en_una_década
```

Output:
```
315360002
```

```python
grader.check("q3_1_2")
```

**Pregunta 3.3.1.** Complete el código en la siguiente celda para completar los datos del experimento.

*Pista:* No se requiere ningún cálculo; simplemente complete los datos del párrafo anterior.

```python
# t, la duración de la caída en el experimento, en segundos.
# Complete esto.
# BEGIN SOLUTION NO PROMPT
time = 1.2
estimated_distance_m = 1.6*time**2/2
# END SOLUTION

""" # BEGIN PROMPT
time = ...
estimated_distance_m=...
""" # END PROMPT
# La distancia estimada que realmente cayó el martillo, en metros.
# Complete esto.
estimated_distance_m
```

Output:
```
1.152
```

```python
grader.check("q3_3_1")
```

**Pregunta 3.3.2.** Ahora, completa el código en la siguiente celda para calcular la diferencia entre las distancias previstas y estimadas (en metros) a las que cayó el martillo en este experimento.

Esto solo significa traducir la fórmula anterior ($\frac{1}{2}G\frac{M}{R^2}t^2$) al código Python. Tendrás que reemplazar cada variable en la fórmula matemática con el nombre que le dimos a ese número en el código Python.

*Pista:* Intente utilizar variables que ya definió en la pregunta 3.3.1

```python
# Primero, hemos anotado los valores de las 3 constantes universales.
# que aparecen en la fórmula de Newton.

# G, la constante universal que mide la fuerza de la gravedad.
gravity_constant = 6.674 * 10**-11

# M, la masa de la luna, en kilogramos.
moon_mass_kg = 7.34767309 * 10**22

# R, el radio de la luna, en metros.
moon_radius_m = 1.737 * 10**6


# La distancia que debería haber caído el martillo
# durante la duración de la caída, en metros,
# según la ley de gravedad de Newton.
# El texto de arriba describe la fórmula.
# para esta distancia dada por la ley de Newton.
# **USTED COMPLETA ESTA PARTE.**

# BEGIN SOLUTION NO PROMPT
predicted_distance_m = 0.5*gravity_constant*moon_mass_kg/moon_radius_m**2*1.2**2
# END SOLUTION

""" # BEGIN PROMPT
predicted_distance_m = ...
""" # END PROMPT

# Aquí hemos calculado la diferencia.
# entre la distancia de caída prevista y la distancia que realmente medimos.
# Si ha completado el código anterior, esto debería funcionar.

difference = predicted_distance_m - estimated_distance_m
difference
```

Output:
```
0.018223694659304845
```

```python
grader.check("q3_3_1")
```

**Pregunta 4.1.1.** Complete la línea `num_avenues_away = ...` en la siguiente celda para que la celda calcule la distancia que debe caminar Chunhua y le dé el nombre `manhattan_distance`. Todo lo demás ya lo hemos completado. **Utilice la función `abs`.** Además, asegúrese de ejecutar la celda de prueba después para probar su código.

```python
# Aquí está el número de calles de distancia:
num_streets_away = abs(42-34)

# Calcule la distancia de avenidas de manera similar:
# BEGIN SOLUTION NO PROMPT
num_avenues_away =  abs(7-10)
# END SOLUTION

""" # BEGIN PROMPT
num_avenues_away = ...
""" # END PROMPT

street_length_m = 80
avenue_length_m = 274

# Ahora calculamos la distancia total que debe caminar Chunhua.
manhattan_distance = street_length_m*num_streets_away + avenue_length_m*num_avenues_away

# Hemos incluido esta línea para que veas la distancia que has calculado
# cuando ejecutas esta celda.
# No es necesario cambiarlo, pero puedes hacerlo si lo deseas.
manhattan_distance
```

Output:
```
1462
```

```python
grader.check("q4_1_1")
```

**Pregunta 5.1.** Digamos que la longitud de Genghis es de 16,7 pulgadas. En la siguiente celda, use "abs" para calcular el valor absoluto de la diferencia entre la longitud de Genghis y la longitud promedio del gato. Asigne a ese valor el nombre `genghis_distance_from_average_in`.

```python
# Reemplazar... con una expresión
# para calcular el valor absoluto
# de la diferencia entre la longitud de Genghis (16,7 pulgadas) y la longitud promedio de un gato.
# BEGIN SOLUTION NO PROMPT
genghis_distance_from_average_in= abs(21.7-16.7 )
# END SOLUTION

""" # BEGIN PROMPT
genghis_distance_from_average_in = ...
""" # END PROMPT

genghis_distance_from_average_in
```

Output:
```
5.0
```

```python
grader.check("q51")
```

**Pregunta 5.1.1.** Dadas las longitudes de los gatos de Yanay, Hummus, Gatkes y Zeepty, escribe una expresión que calcule la diferencia más pequeña entre cualquiera de las tres longitudes. Su expresión no debe tener números, solo llamadas a funciones y los nombres "hummus", "gatkes" y "zeepty". Asigne al valor de su expresión el nombre `min_length_difference`.*

```python
hummus =  24.5 # Hummus is 24.5 inches long
gatkes = 19.7 # Gatkes is 19.7 inches long
zeepty = 15.8 # Zeepty is 15.8 inches long

# BEGIN SOLUTION NO PROMPT
hummus_gatkes_length_diff=abs(hummus-gatkes )
hummus_zeepty_length_diff=abs(hummus-zeepty)
zeepty_gatkes_length_diff=abs(gatkes-zeepty)
min_length_difference = min(hummus_gatkes_length_diff,hummus_zeepty_length_diff,zeepty_gatkes_length_diff)
# END SOLUTION

""" # BEGIN PROMPT
min_length_difference = ...
""" # END PROMPT
min_length_difference
```

Output:
```
3.8999999999999986
```

```python
grader.check("q5_1_1")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q0 (points: None)
```
>>> (secret_word in secret_word_List) == True
False
```


Tests for q3_1_2 (points: None)
```
>>> isinstance(segundos_en_una_década, int) == True
True
```
```
>>> import numpy as np
>>> np.isclose(segundos_en_una_década, 315360000, rtol=1e-08) == True
True
```


Tests for q3_3_1 (points: None)
```
>>> isinstance(estimated_distance_m, float) == True
True
```
```
>>> import numpy as np
>>> np.isclose(estimated_distance_m, 1.13, rtol=1000.0) == True
True
```
```
>>> isinstance(difference, float) == True
True
```
```
>>> import numpy as np
>>> np.isclose(difference, 0.0182236, rtol=0.0001) == True
True
```


Tests for q4_1_1 (points: None)
```
>>> isinstance(difference, float) == True
True
```
```
>>> manhattan_distance == 1462
True
```


Tests for q51 (points: None)
```
>>> isinstance(genghis_distance_from_average_in, float) == True
True
```
```
>>> genghis_distance_from_average_in == abs(21.7 - 16.7)
True
```


Tests for q5_1_1 (points: None)
```
>>> isinstance(min_length_difference, float) == True
True
```
```
>>> np.isclose(min_length_difference, 3.89, rtol=0.01) == True
True
```
