```python
# Initialize Otter
import otter
grader = otter.Notebook("lab05.ipynb")
```

<p><img alt="Colaboratory logo" height="140px" src="https://upload.wikimedia.org/wikipedia/commons/archive/f/fb/20161010213812%21Escudo-UdeA.svg" align="left" hspace="10px" vspace="0px"></p>

<h1> Fundamentos en ciencias de datos</h1>

# Laboratorio 5: Simulaciones

¡Bienvenidos!

Repasaremos [iteración](https://www.inferentialthinking.com/chapters/09/2/Iteration.html) y [simulaciones](https://www.inferentialthinking.com/chapters/09/3/Simulation.html), así como introducir el concepto de [aleatoriedad](https://www.inferentialthinking.com/chapters/09/Randomness.html).

Los datos utilizados en esta práctica de laboratorio contendrán datos salariales y otras estadísticas de los jugadores de baloncesto de la temporada 2014-2015 de la NBA. Estos datos se recopilaron de los siguientes sitios de análisis deportivos: [Basketball Reference](http://www.basketball-reference.com) y [Spotrac](http://www.spotrac.com).

**Obtener ayuda en el laboratorio**: Siempre que te sientas estancado o necesites más aclaraciones, busca un GSI o un tutor, ¡y estarán encantados de ayudarte!

Estas son las políticas para obtener crédito completo:

1. Para los estudiantes inscritos en <span style="color: #BC412B">**Regular Labs**</span> en persona, recibirán crédito de laboratorio al **asistir** a la sección de laboratorio y lograr un progreso significativo en el cuaderno de laboratorio. **<span style="text-decoration: underline">debes</span>** enviar esta tarea a Gradescope.

2. Los estudiantes inscritos en <span style="color: #04888B">**Laboratorio de autoservicio**</span> recibirán crédito de laboratorio al **completar** el cuaderno, aprobar los casos de prueba y **enviándolo** a Gradescope antes de las **11:59 p. m. de la fecha de entrega**.

**Envío**: una vez que haya terminado, ejecute todas las celdas además de la última, seleccione Archivo > Guardar cuaderno y luego ejecute la celda final. El resultado contendrá un archivo zip que podrá utilizar para enviarlo en Gradescope.

Comencemos configurando las pruebas y las importaciones ejecutando la siguiente celda.

```python
# Ejecute esta celda, pero no la cambie.

# Estas líneas importan los módulos Numpy y Datascience.
import numpy as np
from datascience import *

# These lines do some fancy plotting magic
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
```

## 1. Nachos and Conditionals

En Python, el booleano es un tipo de datos con solo dos valores posibles: "Verdadero" y "Falso". Las expresiones que contienen operadores de comparación como `<` (menor que), `>` (mayor que) y `==` (igual a) se evalúan como valores booleanos. ¡A continuación encontrará una lista de operadores de comparación comunes!

<img src="https://github.com/darioblandon/materials-fa23/blob/main/lab/lab05/comparisons.png?raw=1" alt="Chart of comparison operators">

Ejecute la siguiente celda para ver un ejemplo de un operador de comparación en acción.

```python
3 > (1 + 1)
```

Incluso podemos asignar el resultado de una operación de comparación a una variable. Tenga en cuenta que `==` y `=` **no** son lo mismo.

```python
result = 10 / 2 == 5
result
```

Así como los operadores aritméticos se pueden aplicar a cada elemento de una matriz, los operadores de comparación también se pueden usar en matrices para comparar una matriz completa con algún valor. El resultado de esta comparación es una matriz de valores booleanos.

```python
make_array(1, 5, 7, 8, 3, -1) > 3
```

Un día, cuando llegas a casa después de una larga semana, ¡ves un plato caliente de nachos esperando en la mesa del comedor! Digamos que siempre que saques un nacho del tazón, tendrá solo **queso**, solo **salsa**, **tanto** queso como salsa, o **ni** queso ni salsa (un triste tortilla chip de hecho).

Intentemos simular tomar nachos del tazón al azar usando la función `np.random.choice(...)`.

### `np.random.choice`

`np.random.choice` elige un elemento al azar de la matriz dada. Es igualmente probable que elija cualquiera de los elementos. Ejecute la siguiente celda varias veces y observe cómo cambian los resultados. _Consejo:_ Para seguir ejecutando una celda varias veces, puedes usar el método abreviado de teclado `ctrl` + `return`.

```python
nachos = make_array('cheese', 'salsa', 'both', 'neither')
np.random.choice(nachos)
```

Para repetir este proceso varias veces, pase un int `n` como segundo argumento para devolver `n` opciones aleatorias diferentes. De forma predeterminada, `np.random.choice` muestra **con reemplazo** y devuelve una *matriz* de elementos. Muestreo **con reemplazo** significa que después de extraer un elemento, se reemplaza de nuevo al lugar donde se está tomando el muestreo y se puede extraer nuevamente en el futuro. Si tomamos muestras `n` veces con reemplazo, cada vez, cada elemento tiene la misma probabilidad de ser seleccionado.

Ejecute la siguiente celda para ver un ejemplo de muestreo con reemplazo 10 veces de la matriz `nachos`.

```python
np.random.choice(nachos, 10)
```

Para contar la cantidad de veces que se elige aleatoriamente un determinado tipo de nacho, podemos usar `np.count_nonzero`

### `np.count_nonzero`

`np.count_nonzero` cuenta el número de valores distintos de cero que aparecen en una matriz. Cuando se pasa una matriz de valores booleanos a través de la función, ésta contará el número de valores "Verdaderos" (recuerde que en Python, **"Verdadero" se codifica como 1 y "Falso" se codifica como 0.**)

Ejecute la siguiente celda para ver un ejemplo que usa `np.count_nonzero`.

```python
np.count_nonzero(make_array(True, False, False, True, True))
```

**Pregunta 1.1** Supongamos que tomamos diez nachos al azar y almacenamos los resultados en una matriz llamada `ten_nachos` como se hace a continuación. **Encuentre la cantidad de nachos solo con queso usando el código** (no ingrese manualmente la respuesta final).

*Pista:* Nuestra solución implica un operador de comparación (por ejemplo, `==`, `<`, ...) y el método `np.count_nonzero`.

```python
ten_nachos = make_array('neither', 'cheese', 'both', 'both', 'cheese', 'salsa', 'both', 'neither', 'cheese', 'both')
number_cheese = ...
number_cheese
```

```python
grader.check("q11")
```

**Declaraciones condicionales**

Una declaración condicional es una declaración de varias líneas que permite a Python elegir entre diferentes alternativas según el valor de verdad de una expresión.

Aquí hay un ejemplo básico.

```
def sign(x):
    if x > 0:
        return 'Positive'
    else:
        return 'Negative'
```

Si la entrada `x` es mayor que `0`, devolvemos la cadena `'Positiva'`. De lo contrario, devolvemos "Negativo".

Si queremos probar varias condiciones a la vez, utilizamos el siguiente formato general.

```
if <if expression>:
    <if body>
elif <elif expression 0>:
    <elif body 0>
elif <elif expression 1>:
    <elif body 1>
...
else:
    <else body>
```

Sólo se evaluará el cuerpo de la primera expresión condicional que sea verdadera. Cada expresión `if` y `elif` se evalúa y considera en orden, comenzando desde arriba. `elif` solo se puede utilizar si lo precede una cláusula `if`. Tan pronto como se encuentra un valor verdadero, se ejecuta el cuerpo correspondiente y se omite el resto de la declaración condicional. Si ninguna de las expresiones `if` o `elif` es verdadera, entonces se ejecuta el `else body`.

Para obtener más ejemplos y explicaciones, consulte la sección sobre declaraciones condicionales [aquí](https://inferentialthinking.com/chapters/09/1/Conditional_Statements.html).

**Pregunta 1.2** Complete la siguiente declaración condicional para que la cadena `'Más por favor'` se asigne a la variable `say_please` si el número de nachos con queso en `ten_nachos` es menor que `5`. Utilice la instrucción if para hacer esto (no reasigne directamente la variable `say_please`).

*Pista*: Deberías usar `number_cheese` de la Pregunta 1.

```python
say_please = '?'

if ...:
    say_please = 'More please'
say_please
```

```python
grader.check("q12")
```

**Pregunta 1.3** Escribe una función llamada `nacho_reaction` que devuelva una reacción (como una cadena) basada en el tipo de nacho pasado como argumento. Utilice la siguiente tabla para hacer coincidir el tipo de nacho con la reacción adecuada.

<img src="https://github.com/darioblandon/materials-fa23/blob/main/lab/lab05/nacho_reactions.png?raw=1" alt="Table of nacho type with their respective reactions">

*Pista:* Si no pasas la prueba, revisa la ortografía de tus reacciones.

```python
def nacho_reaction(nacho):
    if nacho == "cheese":
        return ...
    ... :
        ...
    ... :
        ...
    ... :
        ...

spicy_nacho = nacho_reaction('salsa')
spicy_nacho
```

```python
grader.check("q13")
```

**Pregunta 1.4** Cree una tabla `ten_nachos_reactions` que consista en los nachos en `ten_nachos` así como las reacciones para cada uno de esos nachos. Las columnas deben llamarse `Nachos` y `Reacciones`.

*Sugerencia:* Considere usar el método `apply`, que devuelve una matriz.


```python
ten_nachos_tbl = Table().with_column('Nachos', ten_nachos)
ten_nachos_reactions = ...
ten_nachos_reactions
```

```python
grader.check("q14")
```

**Pregunta 1.5** Usando el código, encuentre el número de '¡Guau!' reacciones para los nachos en `ten_nachos_reactions`.

```python
number_wow_reactions = ...
number_wow_reactions
```

```python
grader.check("q15")
```

## 2. Simulaciones y bucles For
Usando una declaración "para", podemos realizar una tarea varias veces. Esto se conoce como iteración. La estructura general de un bucle for es:

`for <placeholder> in <array>:` seguido de líneas de código sangradas que se repiten para cada elemento de la `matriz` sobre la que se itera. Puede leer más sobre los bucles for [aquí](https://www.inferentialthinking.com/chapters/09/2/Iteration.html).

**NOTA:** A menudo usamos `i` como `marcador de posición` en nuestros ejemplos de clases, ¡pero puedes ponerle cualquier nombre! Algunos ejemplos se pueden encontrar a continuación.

Un uso de la iteración es recorrer un conjunto de valores. Por ejemplo, podemos imprimir todos los colores del arcoíris.

```python
rainbow = make_array("red", "orange", "yellow", "green", "blue", "indigo", "violet")

for color in rainbow:
    print(color)
```

Podemos ver que la parte sangrada del bucle `for`, conocida como cuerpo, se ejecuta una vez para cada elemento en `rainbow`. El nombre "color" se asigna al siguiente valor en "arco iris" al comienzo de cada iteración. Tenga en cuenta que el nombre "color" es arbitrario; fácilmente podríamos haberlo llamado de otra manera. Cualquiera que sea el nombre que elijamos, debemos usarlo de manera consistente durante todo el ciclo "for".

```python
for another_name in rainbow:
    print(another_name)
```

Sin embargo, en general nos gustaría que el nombre de la variable fuera algo informativo.

**Pregunta 2.1** En la siguiente celda, cargamos el texto de _Orgullo y prejuicio_ de Jane Austen, lo dividimos en palabras individuales y almacenamos estas palabras en una matriz `p_and_p_words`. Usando un bucle "for", asigne "longer_than_five" al número de palabras de la novela que tengan más de 5 letras.

*Pista*: Puedes encontrar el número de letras de una palabra con la función `len`.

*Pista*: ¿Cómo puedes usar `longer_than_five` para realizar un seguimiento del número de palabras que tienen más de cinco letras?

```python
austen_string = open('Austen_PrideAndPrejudice.txt', encoding='utf-8').read()
p_and_p_words = np.array(austen_string.split())

longer_than_five = ...

for ... in ...:
    ...
longer_than_five
```

```python
grader.check("q21")
```

Otra forma en que podemos usar bucles "for" es repetir líneas de código muchas veces. Recuerde la estructura de un bucle `for`:

`for <placeholder> in <array>:` seguido de líneas de código sangradas que se repiten para cada elemento de la matriz sobre la que se itera.

A veces, no nos importa cuál es el valor del marcador de posición. En su lugar, aprovechamos el hecho de que el bucle `for` se repetirá tantas veces como la longitud de nuestra matriz. En la siguiente celda, iteramos a través de una matriz de longitud 5 e imprimimos "¡Hola, mundo!" en cada iteración, pero no necesitamos usar el marcador de posición "i" en el cuerpo de nuestro bucle "for".

```python
for i in np.arange(5):
    print("Hello, world!")
```

**Pregunta 2.2** Usando una simulación con 10,000 intentos, asigne `num_diferente` al **número** de veces, en 10,000 intentos, que dos palabras elegidas uniformemente al azar (con reemplazo) de Orgullo y Prejuicio tienen diferentes longitudes.

*Pista 1*: ¿Qué función utilizamos en la sección 1 para tomar muestras al azar con reemplazo de una matriz?

*Pista 2*: Recuerde que `!=` comprueba la no igualdad entre dos elementos.


```python
trials = 10000
num_different = ...

for ... in ...:
    ...
num_different
```

```python
grader.check("q22")
```

## 3. Muestreo de datos de baloncesto

Ahora presentaremos el tema del muestreo, que discutiremos con más profundidad en las conferencias de esta semana. Lo guiaremos a través de este código, pero si desea leer más sobre diferentes tipos de ejemplos antes de intentar esta pregunta, puede consultar la [sección 10 del libro de texto](https://www.inferentialthinking.com/chapters/10/Sampling_and_Empirical_Distributions.html).

Ejecute la siguiente celda para cargar los datos de jugadores y salarios que usaremos para nuestro muestreo.


```python
player_data = Table().read_table("player_data.csv")
salary_data = Table().read_table("salary_data.csv")
full_data = salary_data.join("PlayerName", player_data, "Name")

# The show method immediately displays the contents of a table.
# This way, we can display the top of two tables using a single cell.
player_data.show(3)
salary_data.show(3)
full_data.show(3)
```

En lugar de obtener datos de cada jugador (como en las tablas cargadas arriba), imagine que obtuvimos datos solo de un subconjunto más pequeño de jugadores. Para 492 jugadores, no es tan descabellado esperar ver todos los datos, pero normalmente no tenemos tanta suerte.

Si queremos hacer estimaciones sobre una determinada propiedad numérica de la población, es posible que tengamos que hacerlo basándose sólo en una muestra más pequeña. La propiedad numérica de la población se conoce como **parámetro** y la estimación se conoce como **estadística** (por ejemplo, la media o la mediana). El hecho de que estas estimaciones sean útiles o no depende a menudo de cómo se recopiló la muestra. Hemos preparado algunos conjuntos de datos de muestra de ejemplo para ver cómo se comparan con el conjunto de datos completo de la NBA. Más adelante te pediremos que crees tus propias muestras para ver cómo se comportan.

Para ahorrar escritura y aumentar la claridad de su código, empaquetaremos el código de análisis en algunas funciones. Esto será útil en el resto de la práctica de laboratorio, ya que necesitaremos crear histogramas y recopilar estadísticas resumidas a partir de esos datos repetidamente.

Hemos definido la función "histogramas" a continuación, que toma una tabla con las columnas "Edad" y "Salario" y dibuja un histograma para cada una. Utiliza anchos de bin de 1 año para "Edad" y $1,000,000 para "Salario".

```python
def histograms(t):
    ages = t.column('Age')
    salaries = t.column('Salary')/1000000
    t1 = t.drop('Salary').with_column('Salary', salaries)
    age_bins = np.arange(min(ages), max(ages) + 2, 1)
    salary_bins = np.arange(min(salaries), max(salaries) + 1, 1)
    t1.hist('Age', bins=age_bins, unit='year')
    plt.title('Age distribution')
    t1.hist('Salary', bins=salary_bins, unit='million dollars')
    plt.title('Salary distribution')

histograms(full_data)
print('Two histograms should be displayed below')
```

**Pregunta 3.1**. Cree una función llamada `compute_statistics` que tome una tabla que contenga una columna "Edad" y una columna "Salario" y:
- Dibuja un histograma de edades.
- Dibuja un histograma de salarios.
- Devuelve una matriz de dos elementos que contiene la edad promedio y el salario promedio (en ese orden)

¡Puedes llamar a la función `histograms` para dibujar los histogramas!


```python
def compute_statistics(age_and_salary_data):
    ...
    age = ...
    salary = ...
    ...


full_stats = compute_statistics(full_data)
full_stats
```

```python
grader.check("q31")
```

### Muestreo aleatorio simple
Un enfoque más justificable es tomar muestras aleatorias y uniformes de los jugadores. En una **muestra aleatoria simple (SRS) sin reemplazo**, nos aseguramos de que cada jugador sea seleccionado como máximo una vez. Imagínese escribir el nombre de cada jugador en una tarjeta, poner las tarjetas en una caja y barajar la caja. Luego, saque las tarjetas una por una y déjelas a un lado, deteniéndose cuando se alcance el tamaño de muestra especificado.

### Producir muestras aleatorias simples
A veces es útil tomar muestras aleatorias incluso cuando tenemos datos de toda la población. Nos ayuda a comprender la precisión del muestreo.

### `sample`

El método de tabla  `sample` produce una muestra aleatoria de la tabla. De forma predeterminada, extrae aleatoriamente **con reemplazo** de las filas de una tabla. El muestreo con reemplazo significa que para cualquier fila seleccionada al azar, existe la posibilidad de que se pueda seleccionar nuevamente si tomamos muestras varias veces. `Sample` toma el tamaño de la muestra como argumento y devuelve una **tabla** con solo las filas que se seleccionaron. Esto difiere de `np.random.choice`, que toma una matriz y genera un valor aleatorio de la matriz.

Ejecute la siguiente celda para ver un ejemplo de llamada a `sample()` con un tamaño de muestra de 5, con reemplazo.

```python
# Just run this cell

salary_data.sample(5)
```

El argumento opcional `with_replacement=False` se puede pasar a través de `sample()` para especificar que la muestra debe extraerse sin reemplazo.

Ejecute la siguiente celda para ver un ejemplo de llamada a `sample()` con un tamaño de muestra de 5, sin reemplazo.

```python
# Just run this cell

salary_data.sample(5, with_replacement=False)
```

**Pregunta 3.2** Produzca una muestra aleatoria simple **sin** reemplazo del tamaño **44** de `full_data`. Luego, ejecute su análisis nuevamente usando la función `compute_statistics` que definió anteriormente. Ejecute la celda varias veces para ver cómo cambian los histogramas y las estadísticas en diferentes muestras.

- ¿Cuánto cambia la edad promedio entre muestras?
- ¿Qué pasa con el salario medio?

(Para su información: srs = muestra aleatoria simple, wor = sin reemplazo)

_Escriba su respuesta aquí, reemplazando este texto._

<!-- BEGIN QUESTION -->



```python
my_small_srswor_data = ...
my_small_stats = ...
my_small_stats
```

<!-- FIN DE PREGUNTA -->

## 4. Más práctica de muestreo aleatorio

Más práctica para el muestreo aleatorio utilizando `np.random.choice`.

### Simulaciones y bucles For (cont.)

**Pregunta 4.1** Podemos usar `np.random.choice` para simular múltiples pruebas.

Stephanie decide jugar a tirar un dado estándar de seis caras, donde su puntuación en cada tirada está determinada por la cara que sale. Quiere saber cuál sería su puntuación total si lanzara el dado 1000 veces. Escribe un código que simule su puntuación total después de 1000 tiradas.

*Pista:* Primero decide los posibles valores que puedes tomar en el experimento (valores puntuales en este caso). Luego use `np.random.choice` para simular las tiradas de Stephanie. Finalmente, suma las tiradas para obtener la puntuación total de Stephanie.

```python
possible_point_values = ...
num_tosses = 1000
simulated_tosses = ...
total_score = ...
total_score
```

```python
grader.check("q41")
```

### Muestreo aleatorio simple (cont.)

**Pregunta 4.2** Como en la pregunta anterior, analice varias muestras aleatorias simples de tamaño 100 de `full_data` utilizando la función `compute_statistics`.
- ¿Las formas del histograma parecen cambiar más o menos entre muestras de 100 que entre muestras de tamaño 44?
- ¿Están los promedios muestrales y los histogramas más cercanos a sus valores/forma reales para la edad o el salario? ¿Qué esperabas ver?

_Escriba su respuesta aquí, reemplazando este texto._

<!-- BEGIN QUESTION -->



```python
my_large_srswor_data = ...
my_large_stats = ...
my_large_stats
```

<!-- FIN DE PREGUNTA -->

---

<img src="https://github.com/darioblandon/materials-fa23/blob/main/lab/lab05/winston.jpeg?raw=1" alt="Picture of an adorable cat named Winston" width="300"/>

**Winston** está muy feliz de que hayas terminado el laboratorio.


---
¡Terminaste con el laboratorio!

**Información importante sobre el envío:**
- **Ejecuta todas las pruebas** y verifica que todas pasan
- **Guardar** desde el menú **Archivo**
- **Ejecute la celda final para generar el archivo zip**
- **Haga clic en el enlace para descargar el archivo zip**
- Luego, ve a [Gradescope](https://www.gradescope.com/courses/576157) y envía el archivo zip a la tarea correspondiente. El nombre de esta tarea es "Lab XX Autograder", donde XX es el número de laboratorio: 01, 02, 03, etc.

- Si terminas temprano en el laboratorio regular, **pídele a uno de los miembros del personal que te marque**.

**Es su responsabilidad asegurarse de que su trabajo esté guardado antes de ejecutar la última celda.**

## Envío

Asegúrese de haber ejecutado todas las celdas de su cuaderno en orden antes de ejecutar la siguiente celda, para que todas las imágenes/gráficos aparezcan en el resultado. La siguiente celda generará un archivo zip para que lo envíe. **¡Guarde antes de exportar!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q11 (points: None)
```
>>> number_cheese == 3
True
```


Tests for q12 (points: None)
```
>>> say_please == 'More please'
True
```


Tests for q13 (points: None)
```
>>> nacho_reaction('salsa')
'Spicy!'
```
```
>>> nacho_reaction('cheese')
'Cheesy!'
```
```
>>> nacho_reaction('both')
'Wow!'
```
```
>>> nacho_reaction('neither')
'Meh.'
```


Tests for q14 (points: None)
```
>>> # One or more of the reaction results could be incorrect
>>> np.count_nonzero(ten_nachos_reactions.column('Reactions') == make_array('Meh.', 'Cheesy!', 'Wow!', 'Wow!', 'Cheesy!', 'Spicy!', 'Wow!', 'Meh.', 'Cheesy!', 'Wow!')) == 10
True
```


Tests for q15 (points: None)
```
>>> 2 < number_wow_reactions < 6
True
```
```
>>> # Incorrect value for number_wow_reactions
>>> number_wow_reactions == 4
True
```


Tests for q21 (points: None)
```
>>> longer_than_five == 35453
True
```


Tests for q22 (points: None)
```
>>> 8100 <= num_different <= 9100
True
```


Tests for q31 (points: None)
```
>>> stats = compute_statistics(full_data)
>>> plt.close()
>>> plt.close()
>>> round(float(stats[0]), 2) == 26.54
True
```
```
>>> stats = compute_statistics(full_data)
>>> plt.close()
>>> plt.close()
>>> round(float(stats[1]), 2) == 4269775.77
True
```


Tests for q41 (points: None)
```
>>> 3000 <= total_score <= 4000
True
```
