<a href="https://colab.research.google.com/github/darioblandon/materials-fa23/blob/main/lab/lab07/lab07.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab07.ipynb")
```

<p><img alt="Colaboratory logo" height="140px" src="https://upload.wikimedia.org/wikipedia/commons/archive/f/fb/20161010213812%21Escudo-UdeA.svg" align="left" hspace="10px" vspace="0px"></p>

<h1> Fundamentos en ciencias de datos</h1>


# Laboratorio 7: Great British Bake Off (Prueba A/B)

¡Bienvenido al Laboratorio 7 de Data 8! El laboratorio de esta semana se centrará en las pruebas A/B utilizando datos del siempre popular programa de televisión británico, [*The Great British Bake Off*](https://en.wikipedia.org/wiki/The_Great_British_Bake_Off).

**Obtener ayuda en el laboratorio**: Siempre que te sientas estancado o necesites más aclaraciones, busca un GSI o un tutor, ¡y estarán encantados de ayudarte!

Estas son las políticas para obtener crédito completo:

1. Para los estudiantes inscritos en <span style="color: #BC412B">**Regular Labs**</span> en persona, recibirán crédito de laboratorio al **asistir** a la sección de laboratorio y lograr un progreso significativo en el cuaderno de laboratorio. **<span style="text-decoration: underline">debes</span>** enviar esta tarea a Gradescope.

2. Los estudiantes inscritos en <span style="color: #04888B">**Laboratorio de autoservicio**</span> recibirán crédito de laboratorio al **completar** el cuaderno, aprobar los casos de prueba y **enviándolo** a Gradescope antes de las **11:59 p. m. de la fecha de entrega**.

**Envío**: una vez que haya terminado, ejecute todas las celdas además de la última, seleccione Archivo > Guardar cuaderno y luego ejecute la celda final. El resultado contendrá un archivo zip que podrá utilizar para enviarlo en Gradescope.

Comencemos configurando las pruebas y las importaciones ejecutando la siguiente celda.

```python
# Ejecute esta celda para configurar el cuaderno, pero no la cambie.

# Estas líneas importan los módulos Numpy y Datascience.
import numpy as np
from datascience import *

# Estas líneas hacen algo de magia en la trama.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', (FutureWarning, np.VisibleDeprecationWarning))

```

## 0. Encuesta de mitad de semestre

Nos gustaría que complete la encuesta de mitad de semestre antes de continuar con esta tarea.

Complete la encuesta a continuación e ingrese la frase secreta que se muestra al final del formulario cuando lo envía. ¡Asigne esta frase a `mid_secret` como una cadena en la celda a continuación!

Encuentre la encuesta [aquí](https://docs.google.com/forms/d/e/1FAIpQLSeDAOTEU8ILuP53naX__xa8MZ0ioa8gnyYacpv38E024H63sA/viewform?usp=sf_link)

```python
mid_secret = ...
```

```python
grader.check("q0")
```

## 1. Pruebas A/B

La prueba A/B es una forma de prueba de hipótesis que le permite hacer comparaciones entre dos distribuciones. También podemos referirnos a una prueba A/B como prueba de permutación.

Casi nunca se te pedirá explícitamente que realices una prueba A/B. Asegúrese de poder identificar situaciones en las que la prueba sea apropiada y sepa cómo implementar correctamente cada paso. A menudo, utilizamos una prueba A/B para determinar si dos muestras provienen o no de la misma distribución subyacente.

**Pregunta 1.1.** Las siguientes afirmaciones son los pasos de una prueba de hipótesis A/B presentada en *orden aleatorio*:

1. Elija una estadística de prueba (normalmente la diferencia de medias entre dos categorías)

2. Mezcle las etiquetas de la muestra original, encuentre su estadística de prueba simulada y repita muchas veces.

3. Encuentre el valor del estadístico de prueba observado.

4. Calcule el valor p basándose en las estadísticas de prueba observadas y simuladas.

5. Definir un modelo nulo y alternativo.

6. Utilice el valor p y el límite del valor p para sacar una conclusión sobre la hipótesis nula.

Asigne `ab_test_order` a una matriz de números enteros que contiene el orden correcto de una prueba A/B, donde el primer elemento de la matriz es el primer paso de una prueba A/B y el último elemento de la matriz es el último paso de una prueba A/B.

```python
ab_test_order = ...
```

```python
grader.check("q1_1")
```

**Pregunta 1.2.** Si la hipótesis nula de una prueba A/B es correcta, ¿el orden de las etiquetas debería afectar las diferencias de medias entre cada grupo? ¿Por qué mezclamos etiquetas en una prueba A/B? Si se encuentra en una sección de laboratorio, confirme su respuesta con un vecino o miembro del personal antes de continuar.


_Escriba su respuesta aquí, reemplazando este texto._

## 2. Los grandes británicos hornean

>"The Great British Bake Off (a menudo abreviado como Bake Off o GBBO) es una competencia de repostería de la televisión británica, producida por Love Productions, en la que un grupo de panaderos aficionados compiten entre sí en una serie de rondas, intentando impresionar a un grupo. de jueces con sus habilidades para hornear" [Wikipedia](https://en.wikipedia.org/wiki/The_Great_British_Bake_Off)

Por cada semana de la competencia, los jueces asignan a un concursante el título de "Star Baker". Al final, cada temporada se corona a un ganador. Usando esta información, nos gustaría investigar cómo ganar premios Star Baker afecta las probabilidades de ganar una temporada del programa.

**Pregunta 2.1.** Queremos saber si ganar más premios Star Baker ___causa___ un cambio en la probabilidad de ganar la temporada. ¿Por qué no es suficiente comparar las tarifas de los panaderos estrella entre ganadores y perdedores?

_Escriba su respuesta aquí, reemplazando este texto._

### Ejecutando un experimento

Realizaremos la siguiente prueba de hipótesis para determinar la asociación entre ganar y la cantidad de premios Star Baker. La población que estamos examinando son todos los concursantes de las temporadas 2 a 11 de GBBO. Vamos a utilizar las siguientes hipótesis nula y alternativa:

**Hipótesis nula:** La distribución de los premios Star Baker entre los concursantes que ganaron su temporada y los concursantes que no ganaron su temporada es la misma.

**Hipótesis alternativa:** Los concursantes que ganen su temporada del programa ganarán más premios Star Baker en promedio.

Nuestra hipótesis alternativa está relacionada con nuestra sospecha de que los concursantes que ganan más premios Star Baker son más hábiles, por lo que tienen más probabilidades de ganar la temporada.

**Pregunta 2.2.** ¿Deberíamos utilizar una prueba A/B para probar estas hipótesis? En caso afirmativo, ¿cuál es nuestro grupo "A" y cuál es nuestro grupo "B"?

_Escriba su respuesta aquí, reemplazando este texto._

Verifique sus respuestas con sus vecinos o un miembro del personal antes de pasar a la siguiente sección.

La siguiente tabla de "panaderos" describe la cantidad de premios de panadero estrella que ganó cada concurso y si ganaron o no su temporada ("1" si ganaron, "0" si no ganaron). Los datos se agregaron manualmente de Wikipedia para las temporadas 2 a 11 del programa. Aleatorizamos el orden de las filas para no estropear el resultado del espectáculo.

```python
bakers = Table.read_table("star_bakers.csv")
bakers.show(3)
```

**Pregunta 2.3.** Cree una nueva tabla llamada `medias` que contenga el número medio de premios de panadero estrella para los panaderos que no ganaron (`ganaron==0`) y los panaderos que ganaron (`ganaron==1`). ). La tabla debe tener los nombres de las columnas "ganado" y "significan los premios Star Baker".

```python
means = ...
means
```

```python
grader.check("q2_3")
```

**Pregunta 2.4.** Visualice la distribución de los premios Star Baker para los ganadores y no ganadores como histogramas superpuestos. Debe utilizar los contenedores que le proporcionamos.

Sugerencia: querrá utilizar el argumento de grupo de `tbl.hist`. Para producir varios histogramas superpuestos basados en valores únicos en una columna determinada, podemos hacer algo como `tbl.hist(..., group=<col_name>, bins=...)`. Esto graficará un histograma para cada valor único en la columna especificada, todo en un solo gráfico.

```python
useful_bins = np.arange(0, 7)
...
```

**Pregunta 2.5.** Queremos determinar si existe una diferencia entre la distribución de los premios Star Baker entre ganadores y no ganadores.

¿Cuál debería ser el estadístico de prueba? ¿Qué valores de este estadístico de prueba respaldan la nula y qué valores respaldan la alternativa? **Asigne `test_option` al número correspondiente a la estadística de prueba correcta.**

1. Valor absoluto de la diferencia de medias entre ambos grupos; los valores altos apoyan el nulo
2. Valor absoluto de la diferencia entre las medias entre ambos grupos; los valores bajos apoyan el nulo
3. Premios Star Baker promedio para los ganadores: premios Star Baker promedio para los no ganadores; los valores altos apoyan el nulo
4. Premios Star Baker promedio para los ganadores: premios Star Baker promedio para los no ganadores; los valores bajos apoyan el nulo

Antes de continuar, confirma tu respuesta con un compañero o en los foros de discusión.

_Pista:_ Deberías pensar en qué medidas utilizamos para describir una distribución.


```python
test_option = ...
```

```python
grader.check("q2_5")
```

**Pregunta 2.6.** Establezca la "diferencia observada" en el estadístico de prueba observado utilizando la tabla de "medias".

```python
observed_difference = ...
observed_difference
```

```python
grader.check("q2_6")
```

**Pregunta 2.7.** Dada una tabla como `panaderos`, una columna de etiqueta `label_col` y una columna de valores `val_col`, escriba una función que calcule la estadística de prueba adecuada.

*Pista:* Asegúrese de tener en cuenta la direccionalidad de nuestra hipótesis alternativa.


```python
def find_test_stat(tbl, label_col, val_col):
    ...

find_test_stat(bakers, "won", "star baker awards")
```

```python
grader.check("q2_7")
```

Cuando ejecutamos una simulación para pruebas A/B, volvemos a muestrear **mezclando las etiquetas** de la muestra original. Si la hipótesis nula es cierta y las distribuciones de premios del panadero estrella son las mismas, esperamos que la diferencia en la media de los premios del panadero estrella no cambie cuando se cambian las etiquetas "ganado".

**Pregunta 2.8.** Escriba una función `simulate_and_test_statistic` para calcular una prueba de nuestra prueba A/B. Su función debería ejecutar una simulación y devolver una estadística de prueba.

```python
def simulate_and_test_statistic(tbl, labels_col, values_col):
    ...

simulate_and_test_statistic(bakers, "won", "star baker awards")
```

```python
grader.check("q2_8")
```

**Pregunta 2.9.** Simule 5000 pruebas de nuestra prueba A/B y almacene las estadísticas de la prueba en una matriz llamada "differences".

```python
# Esta celda puede tardar un par de segundos en ejecutarse
differences = make_array()

...

differences
```

```python
grader.check("q2_9")
```

Ejecute la celda a continuación para ver un histograma de sus estadísticas de prueba simuladas trazadas con su estadística de prueba observada.

```python
Table().with_column('Difference Between Group Means', differences).hist(bins=20)
plots.scatter(observed_difference, 0, color='red', s=30, zorder=2)
plots.ylim(-0.1, 1.35);
```

**Pregunta 2.10.** Encuentre el valor p para su prueba y asígnelo a `empirical_p`.

```python
empirical_p = ...
empirical_p
```

```python
grader.check("q2_10")
```

**Pregunta 2.11.** Utilizando un límite de valor P del 5%, saque una conclusión sobre las hipótesis nula y alternativa. Describe tus hallazgos usando un lenguaje simple y no técnico. ¿Qué le dice su análisis sobre la asociación entre los premios Star Baker y ganar? ¿Qué puede afirmar sobre la causalidad a partir de su análisis estadístico? Confirme su respuesta con un compañero, instructor o en los foros de discusión.


_Escriba su respuesta aquí, reemplazando este texto._

## ¡Todo listo!

**Ozzie** aparece para felicitarte por terminar el laboratorio 7.

<img src="https://github.com/darioblandon/materials-fa23/blob/main/lab/lab07/ozzie.png?raw=1" alt="pcture of a black furry dog sitting on pavement" width="300"/>


---

¡Terminaste con el laboratorio!

**Información importante sobre el envío:**
- **Ejecuta todas las pruebas** y verifica que todas pasan
- **Guardar** desde el menú **Archivo**
- **Ejecute la celda final para generar el archivo zip**
- **Haga clic en el enlace para descargar el archivo zip**
- Luego, ve a [Gradescope](https://www.gradescope.com/courses/576157) y envía el archivo zip a la tarea correspondiente. El nombre de esta tarea es "Lab XX Autograder", donde XX es el número de laboratorio: 01, 02, 03, etc.

- Si terminas temprano en el laboratorio regular, **pídele a uno de los miembros del personal que te marque**.

**Es su responsabilidad asegurarse de que su trabajo esté guardado antes de ejecutar la última celda.**


## Envío

Asegúrese de haber ejecutado todas las celdas de su cuaderno en orden antes de ejecutar la siguiente celda, para que todas las imágenes/gráficos aparezcan en el resultado. La siguiente celda generará un archivo zip para que lo envíe. **¡Guarde antes de exportar!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q0 (points: None)
```
>>> mid_secret == "bing chilling"
True
```


Tests for q1_1 (points: None)
```
>>> len(ab_test_order) == 6
True
```
```
>>> np.allclose(np.array(ab_test_order)[:3] % 2 == 1, True) # Check the first three elements in your array.
True
```
```
>>> np.allclose(np.array(ab_test_order)[3:] % 2 == 0, True) # Check the last three elements in your array.
True
```
```
>>> 
>>> # This imports a hashing library for the autograder.
>>> import hashlib
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> get_hash(np.array(ab_test_order).astype(int))
'a7196ed0f271c873d9750cb92422d911'
```


Tests for q2_10 (points: None)
```
>>> 0 <= empirical_p < 0.05
True
```


Tests for q2_3 (points: None)
```
>>> means.num_rows
2
```
```
>>> np.round(min(means.column("star baker awards mean")), 2) == 0.65
True
```
```
>>> np.round(max(means.column("star baker awards mean")), 2) == 1.5
True
```


Tests for q2_5 (points: None)
```
>>> test_option == 4
True
```


Tests for q2_6 (points: None)
```
>>> isinstance(observed_difference, float)
True
```
```
>>> float(round(observed_difference, 3))
0.848
```


Tests for q2_7 (points: None)
```
>>> np.isclose(round(find_test_stat(bakers, "won", "star baker awards"), 3) - 0.848, 0)
True
```


Tests for q2_8 (points: None)
```
>>> test_stat = round(simulate_and_test_statistic(bakers, "won", "star baker awards"), 3)
>>> -2 < test_stat < 2
True
```
```
>>> np.random.seed(1)
>>> test_stat2 = simulate_and_test_statistic(bakers, "won", "star baker awards")
>>> np.round(test_stat2, 3) == -0.023 or np.round(test_stat2, 3) == -0.132
True
```


Tests for q2_9 (points: None)
```
>>> len(differences)
5000
```
```
>>> abs(np.average(differences)) < 0.05 # On average, your test statistic should be close to 0
True
```
```
>>> all(differences == differences.item(0)) == False # Make sure all of the test statistics are different
True
```
