```python
# Initialize Otter
import otter
grader = otter.Notebook("hw03.ipynb")
```

# Tarea 3: Manipulación de Tablas y Visualización


Por favor complete este notebook llenando las celdas que sde indican. Antes de comenzar, ejecute la celda previa para cargar las pruebas.


**Recursos útiles:**
- [Referencia de Python](http://data8.org/fa23/reference/): Hoja de trucos y métodos útiles de métodos de arreglos y tablas.

**Lectura recomendada**:
* [Visualización](https://inferentialthinking.com/chapters/07/Visualization.html)

Para todos los problemas en los que deba escribir explicaciones y oraciones, **deberá** proporcionar la respuesta en el espacio designado. Además, a lo largo de esta tarea y de todas las futuras, ¡asegúrese de no reasignar variables en todo el notebook! Por ejemplo, si usa `temperatura` en su respuesta a una pregunta, no la reasigne más adelante. De lo contrario, ¡reprobará las pruebas que creía que estaba aprobando anteriormente!



**Fecha de entrega:**

Esta tarea deberá ser entregada ** ___ a las ___**. Si quiere puntos extra, entréguela el (un día antes). Trabajo que se entregue por fuera del plazo no será recibido.


**Nota: Esta tarea tiene pruebas ocultas. Eso significa que aunque los exámenes indiquen que se aprobó al 100%, no significa que su calificación final será del 100%. Realizaremos más pruebas ocultas de corrección una vez que todos entreguen la tarea.**


Evite compartir respuestas o hacer copia con sus compañeros. Establezca en cambio discusiones grupales con otros estudiantes y docentes del curso.

Si desea contar con el tiempo necesario para pedir ayuda o asesoría en caso de estar atascado, se recomienda iniciar temprano con la realización de las tareas. Tenga en cuenta el cronograma del curso y de entregas.


```python
# No cambie esta celda; solo ejecútela.

import numpy as np
from datascience import *
import warnings
warnings.simplefilter('ignore', FutureWarning)

# These lines do some fancy plotting magic.\n",
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
```

## 1. Desempleo


La gran recesión del 2008-2009 fue un período de una caída económica observada a nivel global, con escala y temporalidad variando de país a país. En los Estados Unidos, resultó en un rápido aumento del desempleo que afetó industrias y poblaciones de diferentes campos.


La reserva federal del Banco de San Luis publicó datos sobre los empleos en Estados Unidos. Abajo, hemos cargado datos sobre el desempleo en Estados Unidos. Hay muchas formas de definir el desempleo, y en nuestro dataset incluye dos nociones de tasa de desempleo:



1.   *Índice de Desempleo (o NEI por sus siglas en inglés)*: De las personas en capacidad de trabajar y que están buscando un trabajo de tiempo completo, el porcentaje que no puede encontrar un empleo.


2. *NEI-PTER*: De las personas en capacidad de trabajar y que están buscando por un trabajo de tiempo completo, el porcentaje que no puede encontrar un empleo *o* que tienen un trabajo por horas. El último grupo es llamado "Part-Time for Economic Reasons" (o Tiempo Parcial por Razoners Económicas).



La fuente de los datos la encuentran [aquí](https://fred.stlouisfed.org/categories/33509).

**Pregunta 1.** Los datos están en un archivo CSV llamado `unemployment.csv`. Cargue dicho archivo en una tabla llamada `desempleo`. **(4 Puntos)**

_Hint:_ Luego de cargar la tabla `desempleo`, esta debe verse así:

<img src="unemployment.png" width="20%"/>


```python
desempleo = ...
desempleo
```

```python
grader.check("q1_1")
```

**Pregunta 2.** Organice los datos en orden descendente por NEI, llamando la tabla organizada `por_nei`. Cree otra tabla llamada `por_nei_pter` que organice los datos organizados de forma descendiente pero usando en cambio el índice NEI-PTER. **(4 Puntos)**

```python
por_nei = ...
por_nei_pter = ...
```

```python
grader.check("q1_2")
```

```python
# Ejecute esta celda para verificar su tabla por_nei. No hay necesidad de cambiar el código.
por_nei.show(5)
```

```python
# Ejecute esta celda para verificar su tabla por_nei_pter. No hay necesidad de cambiar el código.
por_nei_pter.show(5)
```

**Pregunta 3.** Usando `take`, asigne `mayor_nei` a una tabla que contenga los datos de los 11 trimestres en los que el NEI fue mayor.

`mayor_nei` debe estar organizado en orden descendente según el índice `NEI`. Note que cada columda de `desempleo` representa un trimestre. **(4 Puntos)**


```python
mayor_nei = ...
mayor_nei
```

```python
grader.check("q1_3")
```

**Pregunta 4.** Se cree que muchas personas pasaron a tener empleos de tiempo parcial o "Part-Time for Economic Reasons" en la "Gran Recesión" del 2008-2009. NEI-PTER es el porcentaje de personas que son desempleados (incluyendo los NEI) más el porcentaje de personas que son PTER.

Calcule un arreglo que contenga el porcentaje de personas que son PTER en cada trimestre. (El primer elemento del arreglo deberá corresponder a la primer fila de `desempleo`, y así sucesivamente.) **(4 Puntos)**


*Nota:* Use la tabla original `desempleo` para esto.


```python
pter = ...
pter
```

```python
grader.check("q1_4")
```

**Pregunta 5.** Agregue `pter` como una columna de `desempleo` (nombre la columna `PTER`) y organice la tabla resultante usando dicha columna en orden descendente. Nombre la tabla resultante `por_pter`

Trate de hacer esto en una sola línea de código, si puede hacerlo. **(4 Puntos)**


```python
por_pter = ...
por_pter
```

```python
grader.check("q1_5")
```

**Pregunta 6.** Genere una gráfica de línea de PTER en función del tiempo. Para hacerlo, cree una nueva tabla llamada `pter_tiempo` con las mismas columnas que la tabla `desempleo` con la adición de dos nuevas columnas: `Año` y `PTER` usando el arreglo `año` y el arreglo `pter`, respectivamente. Luego, genere una gráfica de línea usando uno de los métodos de tabla que aprendió en clase.

El orden de las columnas importa para la corrección de las pruebas. Asegúrese que `Año` venga antes de `PTER`. **(4 Puntos)**


*Nota:* Cuando cree `pter_tiempo`, que no sea solo añadir la columna `año` a la tabla`por_pter`. Por favor siga las instrucciones de la pregunta anterior.


```python
año = 1994 + np.arange(por_pter.num_rows)/4
pter_tiempo = ...
...
plt.ylim(0,2); # No cambie esta línea
```

```python
grader.check("q1_6")
```

**Pregunta 7.** ¿Dónde las tasas de PTER fueron altas durante la recesión (es decir, donde las tasas de PTER son altas entre a lo largo de los años 2008 hasta 2011)? Asigne `altoPTER` a `True` si usted cree que las tasas PTER fueron altas en este período, o `False` si usted cree que no. **(4 Points)**


```python
altoPTER = ...
```

```python
grader.check("q1_7")
```

## 2. Tasas de Nacimiento

La siguiente tabla muestra la población estimada de cada estado de los Estados Unidos, según un censos del 1 de julio del 2015 y del 1 de julio del 2016. Las últimas cuatro columnas describen los componentes del cambio estimado en la población durante este intervalo de tiempo. **Para todas las siguientes preguntas, suponga que la palabra "estados" hace referencia a las 52 filas incluyendo Puerto Rico y el Distrito de Columbia.**

Los datos fueron tomados de [aquí](http://www2.census.gov/programs-surveys/popest/datasets/2010-2016/national/totals/nst-est2016-alldata.csv). (Nota: Sie l archivo no descarga cuando le da click en el link, copie y pegue el enlace en la barra de búsqueda de su navegador). Si quiere leer más acerca de la descripción de las diferentes columnas, haga click [aquí](http://www2.census.gov/programs-surveys/popest/datasets/2010-2015/national/totals/nst-est2015-alldata.pdf).

Los datos que descarga están ujn poco desordenados, por lo que la siguiente celda los organiza un poco.

```python
# No cambie esta celda, solo ejecútela
pop = Table.read_table('nst-est2016-alldata.csv').where('SUMLEV', 40).select([1, 4, 12, 13, 27, 34, 62, 69])
pop = pop.relabeled('POPESTIMATE2015', '2015').relabeled('POPESTIMATE2016', '2016')
pop = pop.relabeled('BIRTHS2016', 'BIRTHS').relabeled('DEATHS2016', 'DEATHS')
pop = pop.relabeled('NETMIG2016', 'MIGRATION').relabeled('RESIDUAL2016', 'OTHER')
pop = pop.with_columns("REGION", np.array([int(region) if region != "X" else 0 for region in pop.column("REGION")]))
pop.set_format([2, 3, 4, 5, 6, 7], NumberFormatter(decimals=0)).show(5)
```

**Question 1.** Asigne `us_tasa_nacimientos` a la tasa de nacimientos anual total durante este intervalo de tiempo. La tasa de nacimientos anual por un período de un año es el total de nacimientos en ese período como una proporción del tamaño total de la población al iniciar el período de tiempo. **(4 Puntos)**

_Hint:_ Recuerde que cada fila en la tabla `pop` se refiere a un estado, no al total de todo Estados Unidos.



```python
us_tasa_nacimientos = ...
us_tasa_nacimientos
```

```python
grader.check("q2_1")
```

**Pregunta 2.** Asigne `mudanzas`al número de estados en los cuales el **valor absolut** de la **tasa anual de migración** fue mayor del 1%. La tasa anual de migración por eun período de un año es el númeor neto de migraciones (hacia adentro y afuera) como una proporción del tamaño de la población al inicio de ese período. La columna `MIGRATION` contiene una estimación de la migración anual neta por estado. **(4 Puntos)**

*Hint*: `tasa_migracion` deberá ser una tabla y `mudanzas`un número.


```python
tasa_migracion = ...
mudanzas = ...
mudanzas
```

```python
grader.check("q2_2")
```

**Pregunta 3.** Asigne `nacimientos_oeste` al número total de nacimientos que ocurrieron en la región 4 (oeste de los Estados Unidos). **(4 Puntos)**

*Hint:* Asegúrese de verificar el tipo de valores en la columna `REGION` y fíltrelos apropiada,ente (todos los tipos deben coincidir).

```python
nacimientos_oeste = ...
nacimientos_oeste
```

```python
grader.check("q2_3")
```

**Pregunta 4.** En la siguiente pregunta, creará una visualización para entender la relación entre las tasas de nacimiento y de muertes. La tasa de muerte anual por un período de un año es el número de muertes en dicho período como una proporción del tamaño de la población en el inicio de ese período de tiempo.

¿Cuál visualización es más apropiada para ver si existe una relación entre las tasas de nacimientos y muertes anuales en los diferentes estados de los Estados Unidos?

1. Gráfico de líneas
2. Gráfico de barras
3. Gráfico de dispersión

Asigne `visualizacion` en la celda siguiente al número correspondiente con la correcta visualización. **(4 Puntos)**


```python
visualizacion = ...
```

```python
grader.check("q2_4")
```

<!-- BEGIN QUESTION -->

**Pregunta 5.** En la siguiente celda de código, cree una visualización que nos permita determinar si hay una relación entre las tasas de nacimiento y muertes durante este intervalo de tiempo. Puede ser de ayuda crear una tabla intermedia que contenga las tasas de nacimiento y las de muerte para cada estado. **(4 Puntos)**

Para su consideración:

- ¿Cuál tipo de gráfica nos ayudará a ilustrar la asociación entre 2 variables?
- ¿Cómo podría manipular una cierta tabla para generar su gráfica?
- CRevise esta [lectura recomendada](https://inferentialthinking.com/chapters/07/Visualization.html) para esta tarea.


```python
# En esta celda use tasa_nacimientos y tasa_muertes para generar su visualización
tasa_nacimiento_2015 = pop.column('BIRTHS') / pop.column('2015')
tasa_muertes_2015 = pop.column('DEATHS') / pop.column('2015')
...
```

<!-- END QUESTION -->

**Pregunta 6.** Verdadero o Falso: Existe una relación entre las tasas de nacimientos y muertes durante este período de tiempo.

Asigne `relacion` a `Verdadero` o `Falso` en la siguiente celda. **(4 Puntos)**


```python
relacion = ...
```

```python
grader.check("q2_6")
```

## 3. Uber

**Nota:** Recomendamos leer el [capítulo 7.2](https://inferentialthinking.com/chapters/07/2/Visualizing_Numerical_Distributions.html) del libro guía antes de llegar a la pregunta 3.

Abajo cargamos tablas que contienen 200000 viajes de Uber en semana en las áreas metropolitanas de Manila y Boston con datos del proyecto [Uber Movement](https://movement.uber.com). Las columnas `sourceid` y `dstid` contienen códigos correspondientes a los lugares de inicio y de destino de cada viaje. La columna `hod` contiene códigos que corresponden a la hora del día en la que el viaje tuvo lugar. La columna `ride time` contiene la duración del viaje en minutos.

```python
boston = Table.read_table("boston.csv")
manila = Table.read_table("manila.csv")
print("Tabla de Boston")
boston.show(4)
print("Tabla de Manila")
manila.show(4)
```

<!-- BEGIN QUESTION -->

**Pregunta 1.** Genere un hiztograma que permita visualizar la distribución de los tiempos de viaje en Boston usando los bines dados en `bines_iguales`. **(4 Puntos)**

*Hint:* Vea el [capítulo 7.2](https://inferentialthinking.com/chapters/07/2/Visualizing_Numerical_Distributions.html) si no sabe cómo especificar los bines.

```python
bines_iguales = np.arange(0, 120, 5)
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Pregunta 2.** Ahora, genere un histograma en el que se visualixe la distribución de todos los tiempos de viaje en Manila usando los bines dados. **(4 Puntos)**


```python
bines_iguales = np.arange(0, 120, 5)
...

# ¡No borre la siguiente línea!
plt.ylim(0, 0.05);
```

<!-- END QUESTION -->

**Pregunta 3.** Vamos a echar un vistazo más detallado a la etiqueta del eje y. Asigne `significado_unidad` a un entero (1, 2, 3) que corrasponda a la "unidad" en "Percent per unit".  **(4 Puntos)**

1. minuto  
2. tiempo de viaje
3. segundo


```python
significado_unidad = ...
significado_unidad
```

```python
grader.check("q3_3")
```

**Pregunta 4.** Asigne `boston_menor_15` y `manila_menor_15` al porcentaje de viajes que tuvieron una diración menor a 15 minutos en sus respectivas áreas metropolitanas. Use las variables de altura que se le dan a continuación para calcular los porcentajes. Su solución solo deberá usar las variables de altura, números y operaciones matemáticas. **No** deberá ingresar a las tablas `boston` y `manila` de ninguna manera. **(4 Points)**



> ***Nota:***  las variables de altura representan la altura del bin que describen.


```python
boston_menor_5_bin_altura = 1.2
manila_menor_5_bin_altura = 0.6
boston_5_a_menor_10_bin_altura = 3.2
manila_5_a_menor_10_bin_altura = 1.4
boston_10_a_menor_15_bin_altura = 4.9
manila_10_a_menor_15_bin_altura = 2.2

boston_menor_15 = ...
manila_menor_15 = ...

boston_menor_15, manila_menor_15
```

```python
grader.check("q3_4")
```

**Pregunta 5.** Miremos con mayor detenimiento a la distribución de tiempos de los viajes en Boston. Asigne `boston_mediana_bin` a un entero (1, 2, 3 ó 4) que corresponda al bin que contenga la mediana del tiempo. **(4 Puntos)**

Let's take a closer look at the distribution of ride times in Boston. Assign `boston_median_bin` to an integer (1, 2, 3, or 4) that corresponds to the bin that contains the median time. **(4 Points)**

1. 0-8 minutos  
2. 8-14 minutos  
3. 14-20 minutos  
4. 20-40 minutos  

*Hint:* La mediana de una lista ordenada tiene la mitad de los elementos a la derecha y la mitad a la izquierda.


```python
boston_mediana_bin = ...
boston_mediana_bin
```

```python
grader.check("q3_5")
```

<!-- BEGIN QUESTION -->

**Pregunta 6.** Identifique una diferencia entre los histogramas, en términos de sus propiedades estadísticas. ¿Qué puede decir sobre el promedio y/o la desviación de cada histograma? **(4 Puntos)**

*Hint*: La mejor forma de hacerlo es comparar los dos histogramas (de los puntos 3.1 and 3.2) visualmente.


_Escriba aquí su respuesta, reemplazando este texto._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Pregunta 7.** ¿Por qué su respuesta en la pregunta 6 es correcta? Basado en las siguientes lecturas, ¿por qué las distribuciones de Boston y Manila son diferentes? **(4 Puntos)**

- [Lectura caso Boston](https://www.climatestotravel.com/climate/united-states/boston)
- [Lectura caso Manila](https://manilafyi.com/why-is-manila-traffic-so-bad/)

*Hint:* Trate de pensar en factores externos de ambas ciudades que pueden estar causando dichas diferencias. Las lecturas le dan algunos facvtores potenciales -- Trate de conectarlos con los datos de tiempos de viaje.



_Escriba aquí su respuesta, reemplazando este texto._

<!-- END QUESTION -->

## 4. Histograms

Considere la siguiente gráfica de dispersión:

![Alt text](scatter.png "Scatter plot showing data points for the variables 'x' and 'y'. The data are symmetric about the x-axis centered at 0 and symmetric about the y-axis centered at 0, but with no data in the [-0.5, 0.5] range on the y-axis.")

Los ejes de la gráfica representan valores de dos variables: $x$ y $y$.


Suponga que tiene una tabla llamada `t` que contiene dos columnas en ella

- `x`: una columna que contuene los valores de x de los puntos en la gráfica de dispersión.
- `y`: una columna que contuene los valores de y de los puntos en la gráfica de dispersión.

Abajo se le dan tres histogramas. Uno corresponde a la columna `x`, uno corresponde a la columna `y`, y otro a ninguna de las dos columnas.

**Histogram A:**

![Alt text](var3.png "Symmetrical, bell-shaped histogram centered around 0")

**Histogram B:**

![Alt text](var1.png "Symmetrical histogram with two peaks at -1 and 1 but no data around 0")

**Histogram C:**

![Alt text](var2.png "Asymmetrical histogram with a peak around -0.5 and a right skew")

**Pregunta 1.** Suponga que usted corre `t.hist('x')`. ¿Cuál histograma debería producir este código? Asigne `histograma_columna_x` a 1, 2, ó 3. **(5 Puntos)**

1. Histograma A
2. Histograma B
3. Histograma C


```python
histograma_columna_x = ...
```

```python
grader.check("q4_1")
```

<!-- BEGIN QUESTION -->

**Pregunta 2.** Dé al menos una razón por la cual eligió el histograma de la Pregunta 1. **Asegúrese de indicar claramente cuál histograma eligió** (ejemplo: "Yo elegí el histograma A porque..."). **(5 Puntos)**


_Escriba aquí su respuesta, reemplazando este texto._

<!-- END QUESTION -->

**Pregunta 3.** Suponga que usted corre `t.hist('y')`. ¿Cuál histograma debería producir este código? Asigne `histograma_columna_y` a 1, 2, ó 3. **(5 Puntos)**

1. Histograma A
2. Histograma B
3. Histograma C



```python
histograma_columna_y = ...
```

```python
grader.check("q4_3")
```

<!-- BEGIN QUESTION -->

**Pregunta 4.** Dé al menos una razón por la cual eligió el histograma de la Pregunta 3. **Asegúrese de indicar claramente cuál histograma eligió** (ejemplo: "Yo elegí el histograma A porque..."). **(5 Puntos)**


_Escriba aquí su respuesta, reemplazando este texto._

<!-- END QUESTION -->

¡Terminó la tarea 1!

**Información sobre la entrega:** Asegúrese de correr todas las pruebas y verificar que se ejectuen correctamente. Aquí, seleccione **Guardar Notebook** En el menú de **Archivo** y **corra la  última celda** para luego darle click en descargar el archivo zip. Con el archivo, vaya a [Gradescope](https://www.gradescope.com/courses/576157) y cargue el archivo zip correspondiente a esta tarea. El nombre de esta tarea es "HW 01 Autograder". **Es su responsabilidad asegurarse que su trabajo se guarde antes de correr la última celda**.

## Estudiantes de Fundamentos de Ciencia de Datos

¡Felicitaciones por terminar la tarea 3!

<img src="./lego.png" width="30%" alt="Close up picture of dog smiling."/>

Perro de la semana: **Lego**



## Entrega

Asegúrese que ha ejecutado todas las celdas en su notebook en orden antes de ejecutar la celda de abajo. De esta forma todas las imágenes/gráficas apareceran en la salida. La celda de abajo va a generar un archivo con extensión .zip para que usted lo entregue. **¡Por favor guarde antes de exportar!**

```python
# Guarde primero su notebook, luego ejecute esta celda para obtener el archivo a
# entregar.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_1 (points: [4])
```
>>> unemployment.select('Date', 'NEI', 'NEI-PTER').take(0)
Date       | NEI     | NEI-PTER
1994-01-01 | 10.0974 | 11.172
```


Tests for q1_2 (points: [2, 2])
```
>>> by_nei.take(0)
Date       | NEI     | NEI-PTER
2009-10-01 | 10.9698 | 12.8557
```
```
>>> by_nei_pter.take(0)
Date       | NEI     | NEI-PTER
2009-10-01 | 10.9698 | 12.8557
```


Tests for q1_3 (points: [2])
```
>>> greatest_nei.take(0)
Date       | NEI     | NEI-PTER
2009-10-01 | 10.9698 | 12.8557
```


Tests for q1_4 (points: [1, 3])
```
>>> # It looks like you subtracted in the wrong order.
>>> round(pter.item(6), 4) != -1.1282
True
```
```
>>> round(pter.item(6), 4)
1.1282
```


Tests for q1_5 (points: [4])
```
>>> by_pter.take(0)
Date       | NEI     | NEI-PTER | PTER
2009-07-01 | 10.8089 | 12.7404  | 1.9315
```


Tests for q1_6 (points: [0])
```
>>> pter_over_time.take(0)
Date       | NEI     | NEI-PTER | Year | PTER
1994-01-01 | 10.0974 | 11.172   | 1994 | 1.0746
```


Tests for q1_7 (points: [0])
```
>>> highPTER == True or highPTER == False
True
```


Tests for q2_1 (points: [0])
```
>>> 0 < us_birth_rate < 1
True
```


Tests for q2_2 (points: [0])
```
>>> 0 < movers <= 52
True
```


Tests for q2_3 (points: [0])
```
>>> 5e5 < west_births < 1e6
True
```


Tests for q2_4 (points: [0, 0])
```
>>> # visualization should be assigned to an integer.
>>> type(visualization) == int
True
```
```
>>> # Make sure visualization is assigned to 1, 2 or 3.
>>> 1 <= visualization <= 3
True
```


Tests for q2_6 (points: [0])
```
>>> type(assoc) is bool
True
```


Tests for q3_3 (points: [0])
```
>>> 1 <= unit_meaning <= 3
True
```


Tests for q3_4 (points: [0, 0])
```
>>> boston_under_15 >= 0 and boston_under_15 <= 100
True
```
```
>>> manila_under_15 >= 0 and manila_under_15 <= 100
True
```


Tests for q3_5 (points: [0])
```
>>> # Make sure boston_median_bin is assigned to 1, 2, 3, or 4.
>>> 1 <= boston_median_bin <= 4
True
```


Tests for q4_1 (points: [0, 1])
```
>>> # Make sure you assign histogram_column_x to either 1 or 2!
>>> type(histogram_column_x) == int
True
```
```
>>> histogram_column_x == 1 or histogram_column_x == 2 or histogram_column_x == 3
True
```


Tests for q4_3 (points: [0, 1])
```
>>> # Make sure you assign histogram_column_y to either 1 or 2!
>>> type(histogram_column_y) == int
True
```
```
>>> # Make sure histogram_column_y is assigned to 1, 2 or 3.
>>> histogram_column_y == 1 or histogram_column_y == 2 or histogram_column_y == 3
True
```
