```python
# Initialize Otter
import otter
grader = otter.Notebook("hw05.ipynb")
```

# Tarea 5: Aplicando Funciones e Iteraciones

Por favor complete este notebook llenando las celdas que se especifican. Antes de comenzar, ejecute la celda anterior para cargar las pruebas.

**Recursos 칰tiles:**
- [Referencia de Python](http://data8.org/fa23/reference/): Hoja de trucos y m칠todos 칰tiles de m칠todos de arreglos y tablas.

**Lecturas recomendadas**:

*[Gu칤a de pensamiento tabular](https://drive.google.com/file/d/1D9PXUgDLXjFBXZoyGVSKpzk0nKQziDTP/view?usp=sharing)
* [Aplicando Funciones](https://www.inferentialthinking.com/chapters/08/1/Applying_a_Function_to_a_Column.html)
* [Condicionales](https://www.inferentialthinking.com/chapters/09/1/Conditional_Statements.html)
* [Iteraciones](https://www.inferentialthinking.com/chapters/09/2/Iteration.html)

Complete este notebook llenando las celdas que se indican. **Antes de comenzar, ejecute la celda de abajo para configurar el notebook importando algunas librer칤as que ser치n 칰tiles.** Cada vez que inicie su servidor, necesitar치 ejecutar esta celda de nuevo.

Para todos los problemas en los que deba escribir explicaciones y oraciones, **deber치** proporcionar la respuesta en el espacio designado. Adem치s, a lo largo de esta tarea y de todas las futuras, 춰aseg칰rese de no reasignar variables en todo el notebook! Por ejemplo, si usa `temperatura` en su respuesta a una pregunta, no la reasigne m치s adelante. De lo contrario, 춰reprobar치 las pruebas que cre칤a que estaba aprobando anteriormente!


**Fecha de entrega:**

Esta tarea deber치 ser entregada ** ___ a las ___**. Si quiere puntos extra, entr칠guela el (un d칤a antes). Trabajo que se entregue por fuera del plazo no ser치 recibido.


**Nota: Esta tarea tiene pruebas ocultas. Eso significa que aunque los ex치menes indiquen que se aprob칩 al 100%, no significa que su calificaci칩n final ser치 del 100%. Realizaremos m치s pruebas ocultas de correcci칩n una vez que todos entreguen la tarea.**


Evite compartir respuestas o hacer copia con sus compa침eros. Establezca en cambio discusiones grupales con otros estudiantes y docentes del curso.

Si desea contar con el tiempo necesario para pedir ayuda o asesor칤a en caso de estar atascado, se recomienda iniciar temprano con la realizaci칩n de las tareas. Tenga en cuenta el cronograma del curso y de entregas.

**Question 1.** The DATA C8 Fall 2023 Midterm Exam will take place on **Friday, October 13th from 7PM - 9PM PT**. Please complete this form so that we can best accommodate you for the midterm.
_**All students are required to fill out this form**_. The deadline to submit this form is **Saturday, October 7th by 11:59PM.** The link can be found below. **(1 Point)**

- [Fall 2023 Midterm Accommodations Survey](https://forms.gle/havE742QdpzJxvvq7)

Assign `secret_phrase` to the secret phrase given at the end of the accommodations survey. Make sure the phrase is in quotes (i.e. is a string)!


```python
secret_phrase = ...
```

```python
grader.check("q0_1")
```

## 1. Temporada de F칰tbol 2021

```python
# Ejecute esta celda para configurar el notebook. No la cambie o modifique.

# Estas l칤neas importan los m칩dulos Numpy y Datascience.


import numpy as np
from datascience import *

# Estas l칤neas hacen algo de magia en las gr치ficas.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)
```

James est치 tratando de analizar cu치n bien se ha deseempe침ado el equipo de f칰tbol de California en la temporada del 2021. Un juego de f칰tbol est치 dividido en cuatro per칤odos, llamados cuartos. El n칰mero de puntos que el equipo anot칩 en cada cuarto y el n칰mero de puntos que anot칩 su oponente en cada cuarto est치n en la tabla llamada `cal_fb.csv`.

```python
# Solo ejecute esta celda
# Lectura del archivo cal_fb csv
games = Table().read_table("cal_fb.csv")
games.show()
```

Comencemos encontrando el total de puntos que cada equipo anot칩 en un juego.


**Pregunta 1.** Escriba una funci칩n llamada `sum_scores`. Esta deber치 tomar cuatro argumentos, donde cada argumento representa un n칰mero entero con el puntaje de cada equipo en cada cuarto. Esta deber치 regresar el puntaje total de cada equipo para ese juego. **(2 Puntos)**

*Hint:* No sobrepiense esta pregunta.



```python
def sum_scores(..., ..., ..., ...):
    '''Regresa el puntaje total anotado por un equipo en los cuartos'''
    ...

sum_scores(14, 7, 3, 0) #No cambie esta l칤nea
```

```python
grader.check("q1_1")
```

**Pregunta 2.** Cree una nueva tabla `final_scores` con tres columnas en este orden espec칤fico: `Opponent`, `Cal Score`, `Opponent Score`. Usted deber치 crear  las columnas `Cal Score` y `Opponent Score`. Use la funci칩n `sum_scores` que defini칩 antes para resolver este problema. **(5 Puntos)**

*Hint:* Si quiere aplicar una funci칩n que toma m칰ltiples argumentos, puede pasar m칰ltuples nombres de columnas como argumentos en `tbl.apply()`. Los valores de las columnas ser치n entregados a los correspondientes argumentos de la funci칩n. Eche un vistazo a la Referencia de Python y a la demostraci칩n de la clase para conocer la sintaxis.

*Nota:* Si est치 encontrando problemas creando `final_scores`, revise que `cal_scores` y `opp_scores` entreguen los resultados que usted desee. Si est치 obteniendo un error `TypeError`, revise la [Referencia de Python ](https://www.data8.org/fa23/reference/) para ver si los inputs/outputs de la funci칩n sean lo que usted est치 esperando.



```python
cal_scores = ...
opp_scores = ...
final_scores = ...
final_scores
```

```python
grader.check("q1_2")
```


Podemos obtener filas espec칤ficas de una tabla. Puede usar `tbl.row(n)` para obtener la `n`-칠sima fila de una tabla. `row.item("column_name")` le permitir치 seleccionar el elemento que corresponda a la `column_name` en una fila particular. Ac치 va un ejemplo:



```python
# Solo ejecute esta celda

games.row(9) # <-- esta l칤nea entrega un objeto con las filas.
```

```python
# Solo ejecute esta celda

games.row(9).item("Cal 4Q") # <-- esto entregar치 un elemento (por ejemplo un entero) de un objeto fila.
```

**Pregunta 3.** Queremos verpara un partido en particular, as칤 California haya perdido o no. Escriba una funci칩n llamada `did_cal_lose`. Esta deber치 recibir como argumento: un **objeto fila** de la tabla `final_scores`.Esta deber치 regresar `True` si el punjtaje del equipo de California fue menor que le del oponente, y `False` si fue mayor. **(5 Puntos)**

*Nota 1*: "Objeto fila" significa una fila de la tabla extra칤da (detr치s de escenas) usando `tbl.row(index)`que contiene todos los datos de una fila en espec칤fico. **No** es el 칤ndice de una fila. No intente llamar a `final_scores.row(row)` dentro de la funci칩n.


*Nota 2*: Si sigue confundido sobre los objetos fila,. intente imprimir `final_scores.row(1)` en una nueva celda para visualizar lo que entrega. Esta l칤nea de c칩digo est치 est치 tomando el objeto fila ubicado en el 칤ndice 1 de la tabla `final_scores` y la entrega. Cuando usted lo muestra en una celda, ver치 que no est치 ubicado en una tabla, sino que es un objeto aparte por s칤 mismo.


```python
def did_cal_lose(row):
    ...

did_cal_lose(final_scores.row(1)) #NO CAMBIE ESTA LINEA
```

```python
grader.check("q1_3")
```

**Pregunta 4.** James quiere ver c칩mo se desempe침칩 el equipo de California contra cada oponente durante la temporada 2021. Usando la tabla `final_scores`:

1. Asigne `results_array` a un arreglo de valores `True` y `False` de acuerdo a si el equipo de California perdi칩 o no.
2. Agregue `results_array` a la tabla `final_scores` en una columna llamada `Results` y asigne esta tabla a `final_scores_with_results`.
3. Luego, respectivamente, asigne el n칰mero de victorias o derrotas del equipo de California a `cal_wins` y `cal_losses`.


**(10 Puntos)**

*Hint*: `tbl.apply()` puede ser 칰til. Vaya a la [Referencia de Python](https://www.data8.org/fa23/reference/) si no est치 seguro de c칩mo funciona.


```python
results_array = ...
final_scores_with_results = ...
cal_losses = ...
cal_wins = ...

# No borre o cambie la siguiente l칤nea:
print(f"En la temporada 2021, el equipo de California gan칩 {cal_wins} partidos y perdi칩 {cal_losses} partidos. Go Bears! 游냩")
```

```python
grader.check("q1_4")
```

## 2. Unrolling Loops (Opcional)

**Esta secci칩n de la tarea 5 es opcional. H치gala por su propia pr치ctica, pero no ser치 incorporada a la nota**



"Desenrollar" un ciclo `for` significa escribir manualmente todo el c칩digo que este ejecuta. El resultado es c칩digo que hace lo mismo que el ciclo, pero sin la ayuda de la estructura del ciclo. TPor ejemplo, para el siguiente ciclo:

    for num in np.arange(3):
        print("El n칰mero es", num)

La versi칩n desenrollada se ve as칤:

    print("El n칰mero es", 0)
    print("El n칰mero es", 1)
    print("El n칰mero es", 2)


Desenrollar un ciclo `for` es una gran forma de entender lo que hace el ciclo en cada paso. En este ejercicio practicar치 desenrollando un ciclo `for`.


En la pregunta de abajo, escriba un c칩digo que haga lo mismo que el c칩digo dado, pero con los ciclos `for` desenrollados. Es una buena idea ejecutar su respuesta y el c칩digo original para verificar que hagan lo mismo.

**Pregunta opcional 1.** Desenrolle el siguiente c칩digo


```python
for joke_iteration in range(3):
    print("Knock. " * (joke_iteration + 1))
    print("Who's there?")
    print("Banana.")
    print("Banana who?")
print("Knock, knock.")
print("Who's there?")
print("Orange.")
print("Orange who?")
print("Orange you glad I didn't say banana?")
```

```python
...
```

춰Termin칩 la tarea 5!  

**Pasos importantes para la entrega:**
1. Corra las pruebas y verifique que todas se ejecuten.
2. Seleccione **Guardar Notebook** desde el men칰 **Archivo**, para luego **ejecutar la 칰ltima celda**.
3. De click en el link para descargar el archivo zip.
4. Vaya a [Gradescope](https://www.gradescope.com/courses/576157) y cargue el archivo zip correspondiente a esta tarea.

**Es su responsabilidad asegurarse que su trabajo se guarde antes de correr la 칰ltima celda.**



## Estudiantes de Fundamentos de Ciencia de Datos

춰Felicitaciones por terminar la tarea 5!

<img src="./mel.jpeg" width="30%" alt="Close up picture of dog smiling, sitting in leaves."/>

Mascota de la semana: **Mel**

## Entrega

Aseg칰rese que ha ejecutado todas las celdas en su notebook en orden antes de ejecutar la celda de abajo. De esta forma todas las im치genes/gr치ficas apareceran en la salida. La celda de abajo va a generar un archivo con extensi칩n .zip para que usted lo entregue. **춰Por favor guarde antes de exportar!**

```python
# Guarde primero su notebook, luego ejecute esta celda para obtener el archivo a
# entregar.
grader.export(pdf=False, run_tests=True)
```

Tests for q0_1 (points: [0])
```
>>> # Make sure that your answer is a string
>>> type(secret_phrase) == str 
True
```


Tests for q1_1 (points: [1, 1])
```
>>> sum_scores(2, 3, 6, 1)
12
```
```
>>> sum_scores(-2,3,5,-10)
-4
```


Tests for q1_2 (points: [0, 0])
```
>>> final_scores.num_columns
3
```
```
>>> set(['Opponent', 'Cal Score', 'Opponent Score']) == set(final_scores.labels)
True
```


Tests for q1_3 (points: [0, 2])
```
>>> 'did_cal_lose' in globals()
True
```
```
>>> did_cal_lose(final_scores.row(1))
True
```


Tests for q1_4 (points: [0, 0])
```
>>> 0 <= cal_wins <= 10
True
```
```
>>> 0 <= cal_losses <= 10
True
```
