```python
# Initialize Otter
import otter
grader = otter.Notebook("lab03.ipynb")
```

<p><img alt="Colaboratory logo" height="140px" src="https://upload.wikimedia.org/wikipedia/commons/archive/f/fb/20161010213812%21Escudo-UdeA.svg" align="left" hspace="10px" vspace="0px"></p>

<h1> Fundamentos en ciencias de datos</h1>

# Laboratorio 3: Tipos de datos y matrices
¡Bienvenidos al Laboratorio 3!

Hasta ahora, hemos usado Python para manipular números y trabajar con tablas. Pero necesitamos analizar los tipos de datos para profundizar nuestra comprensión de cómo trabajar con datos en Python.

En esta práctica de laboratorio, primero verá cómo representar y manipular otro tipo fundamental de datos: el texto. Un fragmento de texto se llama *cadena* en Python. También verás cómo trabajar con *matrices* de datos. Una matriz podría contener todos los números entre 0 y 100 o todas las palabras del capítulo de un libro. Por último, creará tablas y practicará su análisis con su conocimiento de las operaciones de tablas. La [Referencia de Python](https://www.data8.org/fa23/reference/) tiene información que será útil para esta práctica de laboratorio.

**Fecha límite de envío de laboratorio (solo Mega Lab): viernes 8 de septiembre a las 11 p.m.**

**Obtener ayuda en el laboratorio**: Siempre que te sientas estancado o necesites más aclaraciones, busca un GSI o un tutor, ¡y estarán encantados de ayudarte!

Estas son las políticas para obtener crédito completo:

1. Para los estudiantes inscritos en <span style="color: #BC412B">**Regular Labs**</span> en persona, recibirán crédito de laboratorio al **asistir** a la sección de laboratorio y lograr un progreso significativo en el cuaderno de laboratorio. **<span style="text-decoration: underline">debes</span>** enviar esta tarea a Gradescope.

2. Los estudiantes inscritos en <span style="color: #04888B">**Laboratorio de autoservicio**</span> recibirán crédito de laboratorio al **completar** el cuaderno, aprobar los casos de prueba y **enviándolo** a Gradescope antes de las **11:59 p. m. de la fecha de entrega**.

**Envío**: una vez que haya terminado, ejecute todas las celdas además de la última, seleccione Archivo > Guardar cuaderno y luego ejecute la celda final. El resultado contendrá un archivo zip que podrá utilizar para enviarlo en Gradescope.

Comencemos configurando las pruebas y las importaciones ejecutando la siguiente celda.

```python
# Just run this cell
import numpy as np
import math
from datascience import *
```

# 1. Texto
La programación no se trata sólo de números. El texto es uno de los tipos de datos más comunes utilizados en los programas.

El texto está representado por un **valor de string** en Python. La palabra "string" es un término de programación para una secuencia de caracteres. Una cadena puede contener un solo carácter, una palabra, una oración o un libro completo.

Para distinguir los datos de texto del código real, delimitamos las cadenas poniéndolas entre comillas. Las comillas simples (`'`) y las comillas dobles (`"`) son válidas, pero los tipos de comillas de apertura y cierre deben coincidir. El contenido puede ser cualquier secuencia de caracteres, incluidos números y símbolos.

Hemos visto cadenas antes en declaraciones `print` A continuación, se pasan dos cadenas diferentes como argumentos a la función`print` function.

```python
print("I <3", 'Data Science')
```

Así como se pueden dar nombres a los números, se pueden dar nombres a los valores de cadena. No es necesario que los nombres y las cadenas sean similares de ningún modo. Se puede asignar cualquier nombre a cualquier cadena.

```python
one = 'two'
plus = '*'
print(one, plus, one)
```

**Pregunta 1.1.** Yuri Gagarin fue la primera persona en viajar por el espacio exterior. Cuando salió de su cápsula al aterrizar en la Tierra, [según se informa](https://en.wikiquote.org/wiki/Yuri_Gagarin) tuvo la siguiente conversación con una mujer y una niña que vieron el aterrizaje:

     La mujer preguntó: "¿Será que has venido del espacio exterior?"
     Gagarin respondió: "¡De hecho, sí!".

La siguiente celda contiene código sin terminar. Complete los `...`s para que imprima esta conversación *exactamente* como aparece arriba.

Sugerencia: examine cómo se define `mujer_cita` para incluir las comillas dobles.


```python
mujer_preguntando = ...
cita_mujer = '"¿Será que has venido del espacio exterior?"'
respuesta_de_gagarin = 'Gagarin respondió:'
cita_de_gagarin = ...

print(mujer_preguntando, cita_mujer)
print(respuesta_de_gagarin, cita_de_gagarin)
```

```python
grader.check("q11")
```

## 1.1. Métodos de string

Las cadenas se pueden transformar usando **métodos**. Recuerde que métodos y funciones no son lo mismo. Aquí está la sección del libro de texto sobre métodos de cadena: [4.2.1 Métodos de cadena](https://inferentialthinking.com/chapters/04/2/1/String_Methods.html).

Aquí hay un bosquejo de cómo llamar a métodos en una cadena:

     <expresión que se evalúa como una cadena>.<nombre del método>(<argumento>, <argumento>, ...)
    
Un ejemplo de un método de cadena es "replace", que reemplaza todas las instancias de alguna parte de la cadena original (o una *subcadena*) con una nueva cadena.

     <cadena original>.replace(<subcadena antigua>, <subcadena nueva>)
    
`replace` devuelve (evalúa) una nueva cadena, dejando la cadena original sin cambios.
    
Intente predecir el resultado de este ejemplo y luego ejecute la celda.

```python
# Reemplazar una letra
bag = 'bag'
print(bag.replace('g', 't'), bag)
```

Puede llamar funciones sobre los resultados de otras funciones. Por ejemplo, `max(abs(-5), abs(3))` se evalúa como 5. De manera similar, puede llamar a métodos según los resultados de otras llamadas a métodos o funciones.

Es posible que ya hayas notado una diferencia entre funciones y métodos: una función como `max` no requiere un `.` antes de ser llamada, pero un método de cadena como `replace` sí. La referencia de Python a la que se hace referencia en el párrafo inicial puede resultar útil en los casos en los que no esté seguro de cómo llamar a una función o método.

```python
# Calling replace on the output of another call to replace
'train'.replace('t', 'ing').replace('in', 'de')
```

Here's a picture of how Python evaluates a "chained" method call like that:

<img src="chaining_method_calls.png"/>

**Question 1.1.1.** Use `replace` to transform the string `'hitchhiker'` into `'matchmaker'`. Assign your result to `new_word`.


```python
new_word = ...
new_word
```

```python
grader.check("q111")
```

Hay muchos más métodos de cadenas en Python, pero la mayoría de los programadores no memorizan sus nombres ni cómo usarlos. En el "mundo real", la gente normalmente simplemente busca documentación y ejemplos en Internet. Una [lista completa de métodos de cadena](https://docs.python.org/3/library/stdtypes.html#string-methods) aparece en la documentación del lenguaje Python. [Stack Overflow](http://stackoverflow.com) tiene una enorme base de datos de preguntas respondidas que a menudo demuestran cómo utilizar estos métodos para lograr diversos fines. El material cubierto en estos recursos que aún no se haya presentado en esta clase estará fuera de alcance.

## 1.2. Conversión hacia y desde cadenas

Las cadenas y los números son *tipos* diferentes de valores, incluso cuando una cadena contiene los dígitos de un número. Por ejemplo, evaluar la siguiente celda provoca un error porque no se puede agregar un número entero a una cadena.

```python
8 + "8"
```

Las cadenas y los números son *tipos* diferentes de valores, incluso cuando una cadena contiene los dígitos de un número. Por ejemplo, evaluar la siguiente celda provoca un error porque no se puede agregar un número entero a una cadena.

|Función |Descripción|
|-|-|
|`int`|Convierte una cadena de dígitos o un valor flotante en un valor entero ("int")|
|`float`|Convierte una cadena de dígitos (quizás con un punto decimal) o un int en un valor decimal ("flotante")|
|`str`|Convierte cualquier valor en una cadena|

Intente predecir qué tipo de datos y valor se evalúa "ejemplo", luego ejecute la celda.

```python
example = 8 + int("10") + float("8")

print(example)
print("Este ejemplo devolvió un" + str(type(example)) + "!")
```

Supongamos que está escribiendo un programa que busca fechas en un texto y desea que su programa encuentre la cantidad de tiempo que transcurrió entre dos años que identificó. No tiene sentido restar dos textos, pero primero puedes convertir el texto que contiene los años en números.

**Pregunta 1.2.1.** Complete el siguiente código para calcular el número de años que transcurrieron entre "un_año" y "otro_año". No escriba simplemente los números "1618" y "1648" (o "30"); Utilice una función de conversión para convertir los datos de texto dados en números.


```python
# Algunos datos de texto:
un_año = "1618"
otro_año = "1648"

# Complete la siguiente línea. Tenga en cuenta que no podemos simplemente escribir:
# otro_año - un_año
# Si no ves por qué, intenta ver qué sucede cuando
# escribe eso aquí.
diferencia =...
diferencia
```

```python
grader.check("q121")
```

## 1.3. Pasar cadenas a funciones

Los valores de cadena, como los números, pueden ser argumentos de funciones y pueden ser devueltos por funciones.

La función `len` (derivada de la palabra "longitud") toma una sola cadena como argumento y devuelve el número de caracteres (incluidos espacios) en la cadena.

Tenga en cuenta que no cuenta *palabras*. `len("un pequeño paso para el hombre")` se evalúa como 22 caracteres, no 5 palabras.

**Pregunta 1.3.1.** Utilice `len` para encontrar la cantidad de caracteres en la cadena larga en la siguiente celda. Los caracteres incluyen cosas como espacios y puntuación. Asigne `sentence_length` a ese número.

(La cadena es la primera oración de la traducción al inglés de la [Declaración de los Derechos del Hombre] en francés (http://avalon.law.yale.edu/18th_century/rightsof.asp).)

```python
una_frace_larga = "Los representantes del pueblo francés, organizados en Asamblea Nacional, creyendo que la ignorancia, el olvido o el menosprecio de los derechos del hombre son la única causa de las calamidades públicas y de la corrupción de los gobiernos, han decidido exponer en una declaración solemne de los derechos naturales, inalienables y sagrados del hombre, a fin de que esta declaración, estando constantemente ante todos los miembros del cuerpo social, les recuerde continuamente sus derechos y deberes; a fin de que los actos del poder legislativo , así como los del poder ejecutivo, puedan compararse en todo momento con los objetos y fines de todas las instituciones políticas y así puedan ser más respetados, y, por último, para que los agravios de los ciudadanos, basados en lo sucesivo en motivos simples y principios indiscutibles, tenderán al mantenimiento de la constitución y redundarán en la felicidad de todos."
longitud_oración = ...
longitud_oración
```

```python
grader.check("q131")
```

# 2. Matrices

Las computadoras son más útiles cuando puedes usar una pequeña cantidad de código para *realizar la misma acción* en *muchas cosas diferentes*.

Por ejemplo, en el tiempo que le lleva calcular la propina del 18% en la factura de un restaurante, una computadora portátil puede calcular el 18% de propina por cada factura de restaurante pagada por cada ser humano en la Tierra ese día. (¡Eso es si eres bastante rápido haciendo aritmética mentalmente!)

**Las matrices** son la forma en que colocamos muchos valores en un solo lugar para poder operar con ellos como un grupo. Por ejemplo, si `billions_of_numbers` es una matriz de números, la expresión

     .18 *billions_of_numbers

proporciona una nueva matriz de números que contiene el resultado de multiplicar **cada número** en `billones_de_números` por 0,18. Las matrices no se limitan a números; También podemos poner todas las palabras de un libro en una serie de cadenas.

Concretamente, una matriz es una **colección de valores del mismo tipo**.

## 2.1. Hacer matrices

Primero, aprendamos cómo ingresar valores manualmente en una matriz. Por lo general, no es así como funcionan los programas. Normalmente, creamos matrices cargándolas desde una fuente externa, como un archivo de datos.

Para crear una matriz a mano, llame a la función `make_array`. Cada argumento que pase a `make_array` estará en la matriz que devuelve. Ejecute esta celda para ver un ejemplo:

```python
make_array(0.125, 4.75, -1.3)
```

Cada valor de una matriz (en el caso anterior, los números 0,125, 4,75 y -1,3) se denomina *elemento* de esa matriz.

Las matrices en sí mismas también son valores, al igual que los números y las cadenas. Eso significa que puedes asignarlos a nombres o usarlos como argumentos para funciones. Por ejemplo, `len(<some_array>)` devuelve el número de elementos en `some_array`.

**Pregunta 2.1.1.** Haz una matriz que contenga los números 0, 1, -1 y $\pi$, en ese orden. Nómbrelo "números_interesantes".

*Pista:* ¿Cómo obtuviste el valor $\pi$ en el laboratorio 2? Puedes consultarlo de la misma manera aquí. El módulo "matemáticas" se importó en la parte superior de este cuaderno.

```python
números_interesantes = ...
números_interesantes
```

```python
grader.check("q211")
```

**Pregunta 2.1.2.** Cree una matriz que contenga las cinco cadenas `"Hello"`, `","`, `" "`, `"world"` y `"!"`. (El tercero es un espacio único entre comillas). Nómbralo `hello_world_components`.

*Nota:* Si evalúa `hello_world_components`, notará información adicional además de su contenido: `dtype='<U5'`. Esa es simplemente la forma extremadamente críptica que tiene NumPy de decir que los tipos de datos en la matriz son cadenas.


```python
hello_world_components = ...
hello_world_components
```

```python
grader.check("q212")
```

### `np.arange`
Las matrices las proporciona un paquete llamado [NumPy](http://www.numpy.org/) (pronunciado "NUM-pie"). El paquete se llama "numpy", pero es estándar cambiarle el nombre a "np" por brevedad. Puedes hacerlo con:

     importar numpy como np

Muy a menudo, en la ciencia de datos, queremos trabajar con muchos números que estén espaciados uniformemente dentro de algún rango. NumPy proporciona una función especial para esto llamada "arange". La línea de código `np.arange(start, stop, step)` se evalúa como una matriz con todos los números comenzando en `start` y contando `step`, deteniéndose **antes** de llegar a `stop`.

Ejecute las siguientes celdas para ver algunos ejemplos.

```python
# Esta matriz comienza en 1 y cuenta hasta 2
# y luego se detiene antes de las 6
np.arange(1, 6, 2)
```

```python
# Esta matriz no contiene 9
# porque np.arange se detiene *antes* de alcanzar el valor de parada
np.arange(4, 9, 1)
```

**Pregunta 2.1.3.** Importe `numpy` como `np` y luego use `np.arange` para crear una matriz con múltiplos de 99 desde 0 hasta (**e incluyendo**) 9999. (Entonces sus elementos son 0, 99, 198, 297, etc.)


```python
...
multiples_of_99 = ...
multiples_of_99
```

```python
grader.check("q213")
```

## 2.2. Trabajar con elementos individuales de matrices ("indexación")
Trabajemos con un conjunto de datos más interesante. La siguiente celda crea una matriz llamada "montos_población" que incluye las poblaciones mundiales estimadas de cada año desde **1950** hasta aproximadamente el presente. (Las estimaciones provienen del sitio web de la Oficina del Censo de EE. UU.).

En lugar de escribir los datos manualmente, los hemos cargado desde un archivo en su computadora llamado `world_population.csv`. ¡Aprenderá cómo hacerlo más adelante en esta práctica de laboratorio!

```python
cantidades_población = Table.read_table("población_mundial.csv").columna("Población")
cantidades_poblacion
```

Así es como obtenemos el primer elemento de "cantidades_población", que es la población mundial en el primer año del conjunto de datos, 1950.

```python
population_amounts.item(0)
```

El valor de esa expresión es el número 2.557.628.654 (alrededor de 2.500 millones), porque eso es lo primero en la matriz `cantidades_población`.

Observe que escribimos `.item(0)`, no `.item(1)`, para obtener el primer elemento. Esta es una convención extraña en informática. **0 se llama *índice* del primer elemento.** Es el número de elementos que aparecen *antes* de ese elemento. Entonces 3 es el índice del cuarto elemento.

Aquí hay algunos ejemplos más. En los ejemplos, hemos dado nombres a las cosas que obtenemos de `población_cantidades`. Lea y ejecute cada celda.

```python
# El decimotercer elemento de la matriz es la población.
# en 1962 (que es 1950 + 12).
población_1962 = cantidades_población.item(12)
población_1962
```

```python
# El elemento 66 es la población en 2015.
población_2015 = cantidades_población.item(65)
población_2015
```

```python
# La matriz tiene sólo 66 elementos, por lo que esto no funciona.
# (No hay ningún elemento con otros 66 elementos antes).
población_2016 = cantidades_población.item(66)
población_2016
```

Dado que `make_array` devuelve una matriz, podemos llamar a `.item(3)` en su salida para obtener su cuarto elemento, tal como "encadenamos" llamadas al método `replace` anteriormente.

```python
make_array(-1, -3, 4, -2).item(3)
```

**Pregunta 2.2.1.** Establezca "población_1973" como la población mundial en 1973, obteniendo el elemento apropiado de "cantidades_población" usando "item".

```python
población_1973 = ...
población_1973
```

```python
grader.check("q221")
```

## 2.3. Hacer algo con cada elemento de una matriz
Las matrices son principalmente útiles para realizar la misma operación muchas veces, por lo que no tenemos que usar `.item` y trabajar con elementos individuales con frecuencia.

### Redondeo
Aquí hay una pregunta simple que podríamos hacer sobre la población mundial:

> ¿Qué tamaño tenía la población cada año, redondeado al millón más cercano?

El redondeo se utiliza a menudo con números grandes cuando no necesitamos tanta precisión en nuestros números. Un ejemplo de esto es cuando presentamos datos en tablas y visualizaciones.

Podríamos intentar responder nuestra pregunta usando la función `round` que está integrada en Python y el método `item` que acabas de ver.

**Nota:** la función `redondear` toma dos argumentos: el número que se va a redondear y el número de decimales a los que se va a redondear. El segundo argumento se puede considerar como cuántos pasos hacia la derecha o hacia la izquierda se mueve desde el punto decimal. Los números negativos nos dicen que nos muevamos hacia la izquierda y los números positivos nos dicen que nos muevamos hacia la derecha. Entonces, si tenemos "redondear (1234.5, -2)", significa que debemos movernos dos lugares hacia la izquierda y luego convertir todos los números a la derecha de este lugar en ceros. Esto generaría el número 1200.0. Por otro lado, si tenemos "redondear (6.789, 1)", debemos movernos un lugar a la derecha y luego convertir todos los números a la derecha de este lugar en ceros. Esto generaría el número 6,8.

```python
population_1950_magnitude = round(population_amounts.item(0), -6)
population_1951_magnitude = round(population_amounts.item(1), -6)
population_1952_magnitude = round(population_amounts.item(2), -6)
population_1953_magnitude = round(population_amounts.item(3), -6)
```

Pero esto es tedioso y realmente no aprovecha el hecho de que estamos usando una computadora.

En cambio, NumPy proporciona su propia versión de "redondo" que redondea cada elemento de una matriz. Toma dos argumentos: una única matriz de números y el número de decimales a redondear. Devuelve una matriz de la misma longitud, donde el primer elemento del resultado es el primer elemento del argumento redondeado, y así sucesivamente.

**Pregunta 2.3.1.** Utilice `np.round` para calcular la población mundial de cada año, redondeada al millón más cercano (6 ceros). Asigne al resultado (una matriz de 66 números) el nombre "población_redondeada". Tu código debe ser muy corto.

```python
población_redondeada = ...
población_redondeada
```

```python
grader.check("q231")
```

Lo que acabas de hacer se llama aplicación **elemental** de `np.round`, ya que `np.round` opera por separado en cada elemento de la matriz al que se llama. Aquí hay una imagen de lo que está pasando:

<img src="array_round.jpg">

The textbook's [section](https://www.inferentialthinking.com/chapters/05/1/Arrays)  on arrays has a useful list of NumPy functions that are designed to work elementwise, like `np.round`.

### Aritmética
La aritmética también funciona por elementos en matrices, lo que significa que si realiza una operación aritmética (como resta, división, etc.) en una matriz, Python realizará la operación en cada elemento de la matriz individualmente y devolverá una matriz con todos los resultados. Por ejemplo, puedes dividir todas las cifras de población por mil millones para obtener cifras en miles de millones:

```python
población_en_billones = cantidades_población / 1000000000
población_en_millones
```

Puedes hacer lo mismo con la suma, resta, multiplicación y exponenciación (`**`).

## 3. Creando tablas

Una matriz es útil para describir un único atributo de cada elemento de una colección. Por ejemplo, digamos que nuestra colección son todos los estados de EE. UU. Entonces una matriz podría describir la superficie terrestre de cada estado.

Las tablas amplían esta idea al contener múltiples columnas almacenadas y representadas como matrices, cada una de las cuales describe un atributo diferente para cada elemento de una colección. De esta manera, las tablas nos permiten no sólo almacenar datos sobre muchas entidades sino también contener varios tipos de datos sobre cada entidad.

Por ejemplo, en la celda siguiente tenemos dos matrices. El primero, "cantidades_población", se definió anteriormente en la sección 2.2 y contiene la población mundial de cada año (estimada por la Oficina del Censo de EE. UU.). La segunda matriz, "años", contiene los años mismos. Estos elementos están en orden, por lo que el año y la población mundial de ese año tienen el mismo índice en sus matrices correspondientes.

```python
# Simplemente ejecuta esta celda

años = np.arange(1950, 2015+1)
print("Columna de población:", cantidades_población)
print("Columna de años:", años)
```

Supongamos que queremos responder a esta pregunta:

> ¿En qué año la población mundial superó los 6 mil millones?

Técnicamente, podrías responder a esta pregunta simplemente mirando las matrices, pero es un poco complicado, ya que tendrías que contar la posición donde la población cruzó por primera vez los 6 mil millones y luego encontrar el elemento correspondiente en la matriz de años. En casos como estos, podría ser más fácil colocar los datos en una *`Tabla`*, un tipo de conjunto de datos bidimensional.

La siguiente expresión:

- crea una tabla vacía usando la expresión `Table()`,
- agrega dos columnas llamando a `with_columns` con cuatro argumentos (correspondientes al nombre de la columna, los valores de esa columna, el nombre de la segunda columna y los valores de esa columna),
- asigna el resultado al nombre `población`, y finalmente
- evalúa la "población" para que podamos ver la tabla.

Las cadenas `"Año"` y `"Población"` son etiquetas de columnas que hemos elegido. Los nombres `población_cantidades` y `años` se asignaron arriba a dos matrices de la **misma longitud**. La función `with_columns` (puede encontrar la documentación [aquí](http://data8.org/fa21/python-reference.html)) acepta cadenas alternas (para representar etiquetas de columnas) y matrices (que representan los datos en esas columnas). Las cadenas y matrices están separadas por comas.

```python
población = Tabla().with_columns(
     "Población", cantidades_población,
     "Año", años
)
población
```

¡Ahora los datos se combinan en una sola tabla! Es mucho más fácil analizar estos datos. Si necesita saber cuál era la población en 1959, por ejemplo, podrá saberlo con un solo vistazo.

**Pregunta 3.1.** En la celda siguiente, hemos creado 2 matrices. Siguiendo los pasos anteriores, asigne `top_10_movies` a una tabla que tenga dos columnas llamadas "Nombre" y "Clasificación", que contienen `top_10_movie_names` y `top_10_movie_ratings` respectivamente.


```python
top_10_movie_names = make_array(
        'The Shawshank Redemption (1994)',
        'The Godfather (1972)',
        'The Godfather: Part II (1974)',
        'Pulp Fiction (1994)',
        "Schindler's List (1993)",
        'The Lord of the Rings: The Return of the King (2003)',
        '12 Angry Men (1957)',
        'The Dark Knight (2008)',
        'Il buono, il brutto, il cattivo (1966)',
        'The Lord of the Rings: The Fellowship of the Ring (2001)')
top_10_movie_ratings = make_array(9.2, 9.2, 9., 8.9, 8.9, 8.9, 8.9, 8.9, 8.9, 8.8)

top_10_movies = ...

# We've put this next line here 
# so your table will get printed out 
# when you run this cell.
top_10_movies
```

```python
grader.check("q31")
```

#### Cargando una tabla desde un archivo

En la mayoría de los casos, no nos tomaremos la molestia de escribir todos los datos manualmente. En lugar de ello, los cargamos desde una fuente externa, como un archivo de datos. Existen muchos formatos para archivos de datos, pero CSV ("valores separados por comas") es el más común.

`Table.read_table(...)` toma un argumento (una ruta a un archivo de datos en formato **cadena**) y devuelve una tabla.

**Pregunta 3.2.** `imdb.csv` contiene una tabla de información sobre las 250 películas mejor calificadas en IMDb. Cárguelo como una tabla llamada `imdb`.

(¡Quizás recuerde haber trabajado con esta tabla en el Laboratorio 2!)

```python
imdb = ...
imdb
```

```python
grader.check("q32")
```

¿De dónde vino "imdb.csv"? Eche un vistazo a [la carpeta de este laboratorio](./). Debería ver un archivo llamado `imdb.csv`.

Abra el archivo `imdb.csv` en esa carpeta y observe el formato. ¿Que notaste? La terminación del nombre de archivo `.csv` dice que este archivo está en el [formato CSV (valores separados por comas)] (http://edoceo.com/utilitas/csv-file-format).

## 4. ¡Más operaciones en la mesa!

Ahora que ha trabajado con matrices, agreguemos algunos métodos más a la lista de operaciones de tablas que vio en la Práctica 2.

### `columna`

`columna` toma el nombre de la columna de una tabla (en formato de cadena) como argumento y devuelve los valores de esa columna como una **matriz**.

```python
# Returns an array of movie names
top_10_movies.column('Name')
```

### `take`
El método de tabla `take` toma como argumento una matriz de números. Cada número debe ser el índice de una fila de la tabla. Devuelve una **nueva tabla** con solo esas filas.

Por lo general, querrás usar `take` junto con `np.arange` para tomar las primeras filas de una tabla.

```python
# Toma las primeras 5 películas de top_10_movies
top_10_movies.take(np.arange(0, 5, 1))
```

Las siguientes tres preguntas le brindarán práctica para combinar las operaciones que aprendió en esta práctica de laboratorio y la anterior para responder preguntas sobre las tablas "población" e "imdb". Primero, consulte la tabla de "población" de la sección 2.

```python
# Ejecute esta celda para mostrar la tabla de población.
población
```

**Pregunta 4.1.** Consulte la tabla de "población" de la sección 2 de esta práctica de laboratorio. Calcule el año en que la población mundial superó por primera vez los 6 mil millones. Asigne el año a `year_population_crossed_6_billion`.

_Sugerencia:_ ¿Cómo podemos filtrar la tabla solo para filas con valores de población superiores a 6 mil millones? ¿Cómo podemos encontrar el primer año de esta tabla filtrada?

```python
year_population_crossed_6_billion = ...
year_population_crossed_6_billion
```

```python
grader.check("q41")
```

**Pregunta 4.2.** Encuentre la calificación promedio de las películas estrenadas antes del año 2000 y la calificación promedio de las películas estrenadas en el año 2000 o después de las películas en `imdb`.

*Pista*: Piensa en los pasos que debes seguir (tomar el promedio, encontrar las calificaciones, buscar películas estrenadas en los siglos XX y XXI) e intenta ponerlas en un orden que tenga sentido.

```python
antes_2000 = ...
después_o_en_2000 = ...
antes_2000_avg = ...
after_or_in_2000_avg = ...
print("Calificación promedio antes de 2000:", antes_2000_avg)
print("Calificación promedio después o en 2000:", despues_or_in_2000_avg)
```

```python
grader.check("q42")
```

**Pregunta 4.3.** Aquí tienes un desafío: encuentra la cantidad de películas que se estrenaron en años *pares*.

*Pista:* El operador `%` calcula el resto al dividir por un número. Entonces "5% 2" es 1 y "6% 2" es 0. Un número es par si el resto es 0 cuando se divide por 2.

*Pista 2:* `%` se puede usar en matrices, operando por elementos como `+` o `*`. Entonces `make_array(5, 6, 7) % 2` es `array([1, 0, 1])`.

*Pista 3:* Agregue una nueva columna llamada "Año restante" que es el resto cuando el año de estreno de cada película se divide por 2. Para hacer esto, puede usar `tbl.with_column(col_name, col_values)`: un método de tabla que toma en el nombre de la nueva columna y una matriz de valores y devuelve una copia del `tbl` original con la nueva columna. Luego, use `where` para encontrar filas donde esa nueva columna sea igual a 0. Finalmente, use `num_rows` para contar el número de dichas filas.

*Nota:* Estos pasos pueden encadenarse en una sola declaración o dividirse en varias líneas con nombres intermedios asignados. ¡Siempre puedes resolver los problemas como quieras!

```python
núm_películas_año_pares=...
núm_películas_año_pares
```

```python
grader.check("q43")
```

## 5. Más práctica de matrices

¡Las siguientes preguntas serán una excelente práctica para aprender a generar matrices de diferentes maneras, así como para ejecutar diferentes métodos de aritmética de matrices!

###  5.1 `np.arange` (cont.)

##### Lecturas de temperatura
La NOAA (la Administración Nacional Oceánica y Atmosférica de EE. UU.) opera estaciones meteorológicas que miden las temperaturas de la superficie en diferentes sitios de los Estados Unidos. Las lecturas horarias están [disponibles públicamente] (http://www.ncdc.noaa.gov/qclcd/QCLCD?prior=N).

Supongamos que descargamos todos los datos horarios del sitio de Oakland, California, para el mes de diciembre de 2015. Para analizar los datos, queremos saber cuándo se tomó cada lectura, pero encontramos que los datos no incluyen las marcas de tiempo de las lecturas. (la hora a la que se tomó cada uno).

Sin embargo, sabemos que la primera lectura se tomó en el primer instante de diciembre de 2015 (medianoche del 1 de diciembre) y cada lectura posterior se tomó exactamente 1 hora después de la última.

**Pregunta 5.1.1.** Cree una matriz del *tiempo, en segundos, desde el inicio del mes* en el que se tomó **cada lectura horaria**. Nómbralo `collection_times`.

*Pista 1:* Hubo 31 días en diciembre, lo que equivale a ($31 \times 24$) horas o ($31 \times 24 \times 60 \times 60$) segundos. Entonces su matriz debe tener $31 \times 24$ elementos.

*Pista 2:* ¡La función `len` también funciona en matrices! Si su `collection_times` no pasa las pruebas, verifique su longitud y asegúrese de que tenga $31 \times 24$ elementos.

```python
tiempos_colección = ...
tiempos_de_colección
```

```python
grader.check("q511")
```

### 5.2 Hacer algo con cada elemento de una matriz (cont.)

Calcular una propina en varias cuentas de restaurante a la vez (en este caso solo 3):

```python
facturas_restaurante = make_array(20.12, 39.90, 31.01)
print("Facturas del restaurante:\t", facturas_restaurante)

# Multiplicación de matrices
propinas = .2 * facturas_restaurante
print("Consejos:\t\t\t", consejos)
```

<img src="array_multiplication.jpg">

**Pregunta 5.2.1** Supongamos que el cargo total en un restaurante es la factura original más la propina. Si la propina es del 20%, eso significa que podemos multiplicar la factura original por 1,2 para obtener el cargo total. Calcule el cargo total por cada factura en "facturas_restaurante" y asigne la matriz resultante a "cargos_totales".

```python
cargos_totales = ...
cargos_totales
```

```python
grader.check("q521")
```

**Pregunta 5.2.2.** ¡La matriz `more_restaurant_bills` contiene 100.000 billetes! Calcule el cargo total por cada uno en `more_restaurant_bills`. ¿En qué se diferencia tu código?

```python
more_restaurant_bills = Table.read_table("more_restaurant_bills.csv").column("Bill")
more_total_charges = ...
more_total_charges
```

```python
grader.check("q522")
```

La función `sum` toma una única matriz de números como argumento. Devuelve la suma de todos los números de esa matriz (por lo que devuelve un único número, no una matriz).

**Pregunta 5.2.3.** ¿Cuál fue la suma de todas las facturas en `more_restaurant_bills`, *incluidas las propinas*?


```python
sum_of_bills = ...
sum_of_bills
```

```python
grader.check("q523")
```

**Pregunta 5.2.4.** Las potencias de 2 ($2^0 = 1$, $2^1 = 2$, $2^2 = 4$, etc.) surgen con frecuencia en informática. (Por ejemplo, es posible que hayas notado que el almacenamiento en teléfonos inteligentes o dispositivos USB viene en potencias de 2, como 16 GB, 32 GB o 64 GB). Utilice `np.arange` y el operador de exponenciación `**` para calcular la primera 30 potencias de 2, comenzando desde `2^0`.

*Pista 1:* `np.arange(1, 2**30, 1)` crea una matriz con $2^{30}$ elementos y **bloqueará su kernel**.

*Pista 2:* Parte de tu solución involucrará `np.arange`, pero tu matriz no debe tener más de 30 elementos.


```python
potencias_de_2 = ...
poderes_de_2
```

```python
grader.check("q524")
```

---

<img src="java.png" alt="Picture of two brown cats sitting on a window" width="400"/>

¡¡**Java** y **Juno** están emocionados de que hayas terminado el Laboratorio 3!!

Para volver a verificar su trabajo, la celda a continuación volverá a ejecutar todas las pruebas del autocalificador.

---

¡Terminaste con el laboratorio!

**Información importante sobre el envío:**
- **Ejecuta todas las pruebas** y verifica que todas pasan
- **Guardar** desde el menú **Archivo**
- **Ejecute la celda final para generar el archivo zip**
- **Haga clic en el enlace para descargar el archivo zip**
- Luego, vaya a [Gradescope](https://www.gradescope.com/courses/576157) y envíe el archivo zip a la tarea correspondiente. El nombre de esta tarea es "Lab XX Autograder", donde XX es el número de laboratorio: 01, 02, 03, etc.

- Si terminas temprano en el laboratorio regular, **pídele a uno de los miembros del personal que te marque**.

**Es su responsabilidad asegurarse de que su trabajo esté guardado antes de ejecutar la última celda.**

## Envío

Asegúrese de haber ejecutado todas las celdas de su cuaderno en orden antes de ejecutar la siguiente celda, para que todas las imágenes/gráficos aparezcan en el resultado. La siguiente celda generará un archivo zip para que lo envíe. **¡Guarde antes de exportar!**

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q11 (points: None)
```
>>> woman_asking
'The woman asked:'
```
```
>>> gagarin_quote
'"As a matter of fact, I have!"'
```


Tests for q111 (points: None)
```
>>> new_word
'matchmaker'
```


Tests for q121 (points: None)
```
>>> abs(difference) == 30 or abs(difference) == 30.0
True
```


Tests for q131 (points: None)
```
>>> sentence_length
896
```


Tests for q211 (points: None)
```
>>> import numpy as np
>>> type(interesting_numbers) == np.ndarray
True
```
```
>>> len(interesting_numbers)
4
```
```
>>> import numpy as np
>>> all(interesting_numbers == np.array([0, 1, -1, math.pi]))
True
```


Tests for q212 (points: None)
```
>>> import numpy as np
>>> type(hello_world_components) == np.ndarray
True
```
```
>>> len(hello_world_components)
5
```
```
>>> import numpy as np
>>> all(hello_world_components == np.array(["Hello", ",", " ", "world", "!"]))
True
```


Tests for q213 (points: None)
```
>>> type(multiples_of_99) == np.ndarray
True
```
```
>>> len(multiples_of_99)
102
```
```
>>> all(multiples_of_99 == np.arange(0, 9999+99, 99))
True
```


Tests for q221 (points: None)
```
>>> population_1973
3942096442
```


Tests for q231 (points: None)
```
>>> # It looks like you're not making an array.  You shouldn't need to
>>> # use .item anywhere in your solution.
>>> import numpy as np
>>> type(population_rounded) == np.ndarray
True
```
```
>>> # You made an array, but it doesn't have the right numbers in it.
>>> import numpy as np
>>> sum(population_rounded) == 312868000000
True
```


Tests for q31 (points: None)
```
>>> type(top_10_movies) == tables.Table
True
```
```
>>> top_10_movies.select('Name', 'Rating')
Name                                                     | Rating
The Shawshank Redemption (1994)                          | 9.2
The Godfather (1972)                                     | 9.2
The Godfather: Part II (1974)                            | 9
Pulp Fiction (1994)                                      | 8.9
Schindler's List (1993)                                  | 8.9
The Lord of the Rings: The Return of the King (2003)     | 8.9
12 Angry Men (1957)                                      | 8.9
The Dark Knight (2008)                                   | 8.9
Il buono, il brutto, il cattivo (1966)                   | 8.9
The Lord of the Rings: The Fellowship of the Ring (2001) | 8.8
```


Tests for q32 (points: None)
```
>>> type(imdb) == tables.Table
True
```
```
>>> imdb.num_rows == 250
True
```
```
>>> imdb.select('Votes', 'Rating', 'Title', 'Year', 'Decade').sort(0).take(range(2,5))
Votes | Rating | Title                       | Year | Decade
31003 | 8.1    | Le salaire de la peur       | 1953 | 1950
32385 | 8      | La battaglia di Algeri      | 1966 | 1960
35983 | 8.1    | The Best Years of Our Lives | 1946 | 1940
```


Tests for q41 (points: None)
```
>>> # Oops, your name is assigned to the wrong data type!
>>> import numpy as np
>>> type(year_population_crossed_6_billion) == int or type(year_population_crossed_6_billion) == np.int64
True
```
```
>>> year_population_crossed_6_billion == 1999
True
```


Tests for q42 (points: None)
```
>>> abs(before_2000_avg - 8.2783625730994146) < 1e-5
True
```
```
>>> abs(after_or_in_2000_avg - 8.2379746835443033) < 1e-5
True
```


Tests for q43 (points: None)
```
>>> num_even_year_movies == 127
True
```


Tests for q511 (points: None)
```
>>> type(collection_times) == np.ndarray
True
```
```
>>> len(collection_times)
744
```
```
>>> collection_times.item(12) == 43200
True
```
```
>>> collection_times.item(37) == 133200
True
```


Tests for q521 (points: None)
```
>>> # It looks like you're not making an array.  You shouldn't need to
>>> # use .item anywhere in your solution.
>>> import numpy as np
>>> type(total_charges) == np.ndarray
True
```
```
>>> # You made an array, but it doesn't have the right numbers in it.
>>> import numpy as np
>>> sum(abs(total_charges - np.array([24.144, 47.88, 37.212]))) < 1e-6
True
```


Tests for q522 (points: None)
```
>>> # It looks like you're not making an array.  You shouldn't need to
>>> # use .item anywhere in your solution.
>>> import numpy as np
>>> type(more_total_charges) == np.ndarray
True
```
```
>>> # You made an array, but it doesn't have the right numbers in it.
>>> import numpy as np
>>> sum(abs(more_total_charges - 1.2 * more_restaurant_bills)) < 1e-6
True
```


Tests for q523 (points: None)
```
>>> round(sum_of_bills, 2) == 1795730.06
True
```


Tests for q524 (points: None)
```
>>> all(powers_of_2 == 2 ** np.arange(30))
True
```
