<a href="https://colab.research.google.com/github/darioblandon/materials-fa23/blob/main/lab/lab04/lab04_td.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab04.ipynb")
```

<p><img alt="Colaboratory logo" height="140px" src="https://upload.wikimedia.org/wikipedia/commons/archive/f/fb/20161010213812%21Escudo-UdeA.svg" align="left" hspace="10px" vspace="0px"></p>

<h1> Fundamentos en ciencias de datos</h1>

# Laboratorio 4: Funciones y Visualizaciones

¡Bienvenidos al Laboratorio 4! ¡Esta semana aprenderemos sobre funciones, métodos de tablas como "aplicar" y cómo generar visualizaciones!

Lectura recomendada:

* [Aplicar una función a una columna](https://inferentialthinking.com/chapters/08/1/Applying_a_Function_to_a_Column.html)
* [Visualizaciones](https://inferentialthinking.com/chapters/07/Visualization.html)
* [Referencia de Python](https://www.data8.org/fa23/reference/)

**Fecha límite de envío de laboratorio (solo laboratorio de autoservicio): viernes 27 de septiembre a las 11 p.m.**

**Obtener ayuda en el laboratorio**: Siempre que te sientas estancado o necesites más aclaraciones, busca un GSI o un tutor, ¡y estarán encantados de ayudarte!

Estas son las políticas para obtener crédito completo:

1. Para los estudiantes inscritos en <span style="color: #BC412B">**Regular Labs**</span> en persona, recibirán crédito de laboratorio al **asistir** a la sección de laboratorio y lograr un progreso significativo en el cuaderno de laboratorio. **<span style="text-decoration: underline">debes</span>** enviar esta tarea a Gradescope.

2. Los estudiantes inscritos en <span style="color: #04888B">**Laboratorio de autoservicio**</span> recibirán crédito de laboratorio al **completar** el cuaderno, aprobar los casos de prueba y **enviándolo** a Gradescope antes de las **11:59 p. m. de la fecha de entrega**.

**Envío**: una vez que haya terminado, ejecute todas las celdas además de la última, seleccione Archivo > Guardar cuaderno y luego ejecute la celda final. El resultado contendrá un archivo zip que podrá utilizar para enviarlo en Gradescope.

Comencemos configurando las pruebas y las importaciones ejecutando la siguiente celda.

First, set up the notebook by running the cell below.

```python
import numpy as np
from datascience import *

# Estas líneas configuran capacidades gráficas.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)
```

## 1. Definición de funciones

Comencemos escribiendo una función que convierta una proporción en porcentaje multiplicándola por 100. Por ejemplo, el valor de `to_percentage(.5)` debe ser el número 50 (sin signo de porcentaje).

Una definición de función tiene algunas partes.

##### `def`
Creamos una función con una declaración `def` (abreviatura de **def**ine):
<br>
<pre style="margin-left: 40px; font-family: 'Courier New';"><code><span style="font-weight: bold; color: green">def</span></code>
</pre>

##### Nombre
Luego viene el nombre de la función. Al igual que otros nombres que hemos definido, no puede comenzar con un número ni contener espacios. Llamemos a nuestra función `to_percentage`:
    
<br>
<pre style="margin-left: 40px; font-family: 'Courier New';"><code><span style="font-weight: bold; color: green">def</span> <span style="color: blue">to_percentage</span></code>
</pre>

##### Signature
Luego viene algo ___firma___ de la función. Esto le dice a Python el número de **argumentos** en la función y los nombres de esos argumentos. Un argumento es un valor que se _pasa_ a la función cuando se llama. Una función puede tener cualquier número de argumentos (¡incluido 0!).

`to_percentage` debe tomar un argumento, y a ese argumento lo llamaremos `proporción` ya que debe ser una proporción.
<br>
<pre style="margin-left: 40px; font-family: 'Courier New';"><code><span style="font-weight: bold; color: green">def</span> <span style="color: blue">to_percentage</span>(proportion)</code>
</pre>
    
Si queremos que nuestra función tome más de un argumento, agregamos una coma entre el nombre de cada argumento, es decir, to_percentage(proporción, decimales). Tenga en cuenta que si tuviéramos cero argumentos, aún colocaríamos los paréntesis () después de ese nombre.

Ponemos **dos puntos** después de la asignación para indicarle a Python que las siguientes líneas identadas son el cuerpo de la función. ¡Asegúrate de recordar los dos puntos!

<br>
<pre style="margin-left: 40px; font-family: 'Courier New';"><code><span style="font-weight: bold; color: green">def</span> <span style="color: blue">to_percentage</span>(proportion):</code>
</pre>

##### Documentación
Las funciones pueden hacer cosas complicadas, por lo que debes escribir una explicación de lo que hace tu función. Para funciones pequeñas, esto es menos importante, pero es un buen hábito aprender desde el principio (aunque la documentación no es estrictamente necesaria). Convencionalmente, las funciones de Python se documentan escribiendo una cadena **indentada** entre comillas triples:

<br>
<pre style="margin-left: 40px; font-family: 'Courier New';"><code><span style="font-weight: bold; color: green">def</span> <span style="color: blue">to_percentage</span>(proportion):<br><span style="padding-left: 40px"><span style="color: red">"""Converts a proportion to a percentage."""</span></span></code>
</pre>

##### Body
Ahora comenzamos a escribir código que se ejecuta cuando se llama a la función. Esto se llama *cuerpo* de la función y cada línea **debe tener una sangría con una tabulación**. Cualquier línea que *no* tenga sangría y esté alineada a la izquierda con la declaración def se considera fuera de la función.

Algunas notas sobre el cuerpo de la función:
- Podemos escribir código que escribiríamos en cualquier otro lugar.
- Usamos los argumentos definidos en la firma de la función. Podemos hacer esto porque se asignan valores a esos argumentos cuando llamamos a la función.
- Generalmente evitamos hacer referencia a variables definidas *fuera* de la función. Si desea hacer referencia a variables fuera de la función, ¡páselas como argumentos!


Ahora, vamos a darle un nombre al número por el que multiplicamos una proporción para obtener un porcentaje:

<br>
<pre style="margin-left: 40px; font-family: 'Courier New';"><code><span style="font-weight: bold; color: green">def</span> <span style="color: blue">to_percentage</span>(proportion):<br><span style="padding-left: 40px"><span style="color: red">"""Converts a proportion to a percentage."""</span></span>
<span style="padding-left: 40px">factor <span style="color: purple">=</span> <span style="color: green">100</span></span></code>
</pre>

##### `return`
La instrucción especial "return" es parte del cuerpo de la función y le dice a Python que haga que el valor de la llamada a la función sea igual a lo que viene justo después de "return". Queremos que el valor de `to_percentage(.5)` sea la proporción .5 veces el factor 100, así que escribimos:
<br>
<pre style="margin-left: 40px; font-family: 'Courier New';"><code><span style="font-weight: bold; color: green">def</span> <span style="color: blue">to_percentage</span>(proportion):<br><span style="padding-left: 40px"><span style="color: red">"""Converts a proportion to a percentage."""</span></span>
<span style="padding-left: 40px">factor <span style="color: purple">=</span> <span style="color: green">100</span></span>
<span style="padding-left: 40px"><span style="font-weight: bold; color: green">return</span> proportion <span style="color: purple">*</span> factor</span>
</code>
</pre>
        
`return` solo tiene sentido en el contexto de una función y **nunca se puede usar fuera de una función**. `return` es siempre la última línea de la función porque Python deja de ejecutar el cuerpo de una función una vez que llega a una declaración `return`. Asegúrese de incluir una declaración de devolución a menos que no espere que la función devuelva nada.

*Nota:* `return` dentro de una función le dice a Python a qué valor se evalúa la función. Sin embargo, hay otras funciones, como "imprimir", que no tienen valor de "retorno". Por ejemplo, "imprimir" simplemente imprime un valor determinado en la consola.

En resumen, "return" se usa cuando quieres decirle a la *computadora* cuál es el valor de alguna variable, mientras que "print" se usa para decirte a ti, un *humano*, su valor.

**Question 1.1.** Define `to_percentage` in the cell below.  Call your function to convert the proportion .2 to a percentage.  Name that percentage `twenty_percent`.


```python
def ...
    ''' (Replace this with your documentation) '''
    ... = ...
    return ...

twenty_percent = ...
twenty_percent
```

```python
grader.check("q11")
```

Aquí hay algo importante sobre las funciones: **los nombres asignados *dentro* del cuerpo de una función solo son accesibles dentro del cuerpo de la función**. Una vez que la función ha regresado, esos nombres desaparecen. Entonces, incluso si creó una variable llamada `factor` y definió `factor = 100` dentro del cuerpo de la función `to_percentage` y luego llamó `to_percentage`, `factor` no tendría un valor asignado fuera del cuerpo. de `a_porcentaje`:

**Nota**: A continuación, debería ver un mensaje de error "NameError" que indica que el "nombre 'factor' no está definido". Python arroja este error porque "factor" no se ha definido fuera del cuerpo de la función "to_percentage".

```python
# Deberías recibir un error al ejecutar esto. (Si no lo haces,
# es posible que haya definido el factor en algún lugar arriba).
factor
```

Como lo hizo con las funciones integradas en laboratorios anteriores (max, abs, etc.), puede pasar valores con nombre como argumentos a su función.

**Pregunta 1.2.** Utilice `to_percentage` nuevamente para convertir la proporción denominada `a_proportion` (definida a continuación) en un porcentaje llamado `a_percentage`.

*Nota:* ¡No es necesario volver a definir `to_percentage`! Al igual que otros valores con nombre, las funciones permanecen después de que las defines.

```python
a_proportion = 2**(0.5) / 2
a_percentage = ...
a_percentage
```

```python
grader.check("q12")
```

En la siguiente celda, definiremos una función llamada `disemvocal`. Toma una sola cadena como argumento. Devuelve una copia de esa cadena, pero con todos los caracteres que son vocales eliminados. (En inglés, las vocales son los caracteres "a", "e", "i", "o" y "u").

Para eliminar todas las "a" de una cadena, usamos `a_string.replace("a", "")`. El método `.replace` para cadenas devuelve una nueva cadena, por lo que podemos llamar a `replace` varias veces, una tras otra.

```python
def disemvowel(a_string):
    """Removes all vowels from a string."""
    return a_string.replace("a", "").replace("e", "").replace("i", "").replace("o", "").replace("u", "")

# Un ejemplo de llamada a la función. (A menudo es útil ejecutar
# una llamada de ejemplo de vez en cuando mientras escribimos una función,
# para ver cómo funciona actualmente).
disemvowels("¿Puedes leer esto sin vocales?")
```

##### Llamadas sobre llamadas sobre llamadas
Así como escribes una serie de líneas para construir un cálculo complejo, es útil definir una serie de pequeñas funciones que se complementan entre sí. Como puedes escribir cualquier código dentro del cuerpo de una función, puedes llamar a otras funciones que hayas escrito.

Si una función es como una receta, definir una función en términos de otras funciones es como tener una receta de pastel que te dice que sigas otra receta para hacer el glaseado y otra para hacer el relleno de mermelada. Esto hace que la receta del pastel sea más corta y clara, y evita tener un montón de recetas de glaseado duplicadas. Es la base de la programación productiva.

Por ejemplo, supongamos que desea contar el número de caracteres *que no son vocales* en un fragmento de texto. Una forma de hacerlo es eliminar todas las vocales y contar el tamaño de la cadena restante.

**Pregunta 1.3.** Escribe una función llamada `num_non_vocals`. Debería tomar una cadena como argumento y devolver un número. Ese número debe ser el número de caracteres en la cadena del argumento que no sean vocales. Debes usar la función `disemvocal` que proporcionamos arriba dentro de la función `num_non_vocals`.

*Pista:* La función `len` toma una cadena como argumento y devuelve el número de caracteres que contiene.


```python
def num_non_vowels(a_string):
    """The number of characters in a string, minus the vowels."""
    ...

# Try calling your function yourself to make sure the output is what
# you expect.
```

```python
grader.check("q13")
```

Functions can also encapsulate code that *displays output* instead of computing a value. For example, if you call `print` inside a function, and then call that function, something will get printed.

The `movies_by_year` dataset in the textbook has information about movie sales in recent years.  Suppose you'd like to display the year with the 5th-highest total gross movie sales, printed within a sentence.  You might do this:

```python
movies_by_year = Table.read_table("movies_by_year.csv")
rank = 5
fifth_from_top_movie_year = movies_by_year.sort("Total Gross", descending=True).column("Year").item(rank-1)
print("Year number", rank, "for total gross movie sales was:", fifth_from_top_movie_year)
```

Después de escribir esto, te das cuenta de que también querías imprimir el segundo y tercer año con mayor puntuación. En lugar de copiar tu código, decides ponerlo en una función. Dado que el rango varía, lo convierte en un argumento para su función.

**Pregunta 1.4.** Escribe una función llamada `print_kth_top_movie_year`. Debería tomar un único argumento, la clasificación del año (como 2, 3 o 5 en los ejemplos anteriores) y debería utilizar la tabla `movies_by_year`. Debería imprimir un mensaje como el de arriba.

*Nota:* Su función no debería tener una declaración de "retorno".


```python
def print_kth_top_movie_year(k):
    # COMENZAR LA SOLUCIÓN
    year = movies_by_year.sort("Total Gross", descending=True).column("Year").item(k-1)
    print("Year number", k, "for total gross movie sales was:", year)
    # SOLUCIÓN FINAL
    """ # BEGIN PROMPT
    ...
    print(...)
"""; # INDICACIÓN DE FINALIZACIÓN

# Ejemplos de llamadas a tu función:
print_kth_top_movie_year(2)
print_kth_top_movie_year(3)
```

```python
grader.check("q14")
```

### `print` No es la misma que `return`
La función `print_kth_top_movie_year(k)` imprime las ventas brutas totales de películas para el año proporcionado! Sin embargo, como no devolvimos ningún valor en esta función, no podemos usarla después de llamarla. Veamos un ejemplo de otra función que imprime un valor pero no lo devuelve.

```python
def print_number_five():
    print(5)
```

```python
print_number_five()
```

Sin embargo, si intentamos utilizar la salida de `print_number_five()`, vemos que se imprime el valor `5` pero obtenemos un TypeError cuando intentamos agregarle el número 2.

```python
print_number_five_output = print_number_five()
print_number_five_output + 2
```

Puede parecer que `print_number_five()` está devolviendo un valor, 5. En realidad, ¡solo te muestra el número 5 sin darte el valor real! Si su función imprime un valor **sin devolverlo** e intenta utilizar ese valor, se encontrará con errores, ¡así que tenga cuidado!

Piensa en cómo podrías agregar una línea de código a la función `print_number_five` (después de `print(5)`) para que el código `print_number_five_output + 5` dé como resultado el valor `10`, en lugar de un error.

## 2. Funciones e ingresos del CEO

En esta pregunta, veremos la remuneración de los directores ejecutivos de las 100 empresas más grandes de California en 2015. Los datos se compilaron a partir de un [análisis de Los Angeles Times] (http://spreadsheets.latimes.com/california-ceo-compensation/) y, en última instancia, provinieron de [presentaciones] (https://www.sec.gov/answers /proxyhtf.htm) exigido por la SEC para todas las empresas que cotizan en bolsa. Dos empresas tienen dos directores ejecutivos, por lo que hay 102 directores ejecutivos en el conjunto de datos.

Copiamos los datos sin procesar de la página del LA Times en un archivo llamado `raw_compensation.csv`. (La página señala que todos los montos en dólares están en **millones de dólares**).

```python
raw_compensation = Table.read_table('raw_compensation.csv')
raw_compensation
```

Queremos calcular el promedio del salario de los directores ejecutivos. Intente ejecutar la celda a continuación.

```python
np.average(raw_compensation.column("Total Pay"))
```

Deberías ver un "TypeError". Examinemos por qué ocurrió este error observando los valores en la columna "Pago total". Para hacerlo, podemos usar la función `tipo`. Esta función nos dice el tipo de datos del objeto que le pasamos. Ejecute las siguientes celdas para ver qué sucede cuando pasamos `23`, `3.5` y `"Hello"` a la función `type`. ¿Tienen sentido sus resultados?

```python
type(23)
```

```python
type(3.5)
```

```python
type("Hello")
```

**Pregunta 2.1.** Utilice la función `type` y establezca `total_pay_type` en el tipo del primer valor en la columna "Pago total".


```python
total_pay_type = ...
total_pay_type
```

```python
grader.check("q21")
```

**Pregunta 2.2.** Deberías haber descubierto que los valores en la columna "Pago total" son cadenas. No tiene sentido tomar el promedio de los valores de una cadena, por lo que debemos convertirlos a números. Extraiga el primer valor en "Pago total". Es el salario de Mark Hurd en 2015, en *millones* de dólares. Llámelo `mark_hurd_pay_string`.


```python
mark_hurd_pay_string = ...
mark_hurd_pay_string
```

```python
grader.check("q22")
```

**Pregunta 2.3.** Convierta `mark_hurd_pay_string` a una cantidad de *dólares*.

Algunas sugerencias, ya que esta pregunta requiere varios pasos:
- El método de cadena `strip` será útil para eliminar el signo de dólar; elimina un carácter específico del inicio o final de una cadena. Por ejemplo, el valor de `"100%".strip("%")` es la cadena `"100"`.
- También necesitarás la función `float`, que convierte una cadena que parece un número en un número real. No te preocupes por el espacio en blanco al final de la cadena; la función `float` ignorará esto.
- Por último, recuerda que la respuesta debe ser en dólares, no en millones de dólares.


```python
mark_hurd_pay = ...
mark_hurd_pay
```

```python
grader.check("q23")
```

Para calcular el salario promedio, debemos hacer esto para cada director ejecutivo. Pero parece que eso implicaría copiar este código 102 veces.

En su lugar, usaremos funciones para realizar este cálculo. Más adelante en esta práctica de laboratorio, veremos el beneficio: podemos llamar a esa función en cada cadena de pago en el conjunto de datos a la vez.

**Pregunta 2.4.** Copie la expresión que utilizó para calcular `mark_hurd_pay` y utilícela como expresión de retorno de la siguiente función. Pero asegúrese de reemplazar el `mark_hurd_pay_string` específico con el nombre genérico `pay_string` especificado en la primera línea de la declaración `def`.

*Sugerencia*: Cuando se trata de funciones, generalmente no se debe hacer referencia a ninguna variable fuera de la función. Por lo general, desea trabajar con los argumentos que se le pasan, como `pay_string` para esta función. Si estás usando `mark_hurd_pay_string` dentro de tu función, ¡estás haciendo referencia a una variable externa!

```python
def convert_pay_string_to_number(pay_string):
    """Converts a pay string like '$100' (in millions) to a number of dollars."""
    ...
```

```python
grader.check("q24")
```

Ejecutar esa celda no convierte ninguna cadena de pago en particular. En su lugar, crea una función llamada `convert_pay_string_to_number` que puede convertir *cualquier* cadena con el formato correcto en un número que represente millones de dólares.

Podemos llamar a nuestra función tal como llamamos a las funciones integradas que hemos visto. Toma un argumento, *una cadena*, y devuelve un flotante.

```python
convert_pay_string_to_number('$42')
```

```python
convert_pay_string_to_number(mark_hurd_pay_string)
```

```python
# También podemos calcular el salario de Safra Catz de la misma manera:
convert_pay_string_to_number(raw_compensation.where("Name", are.containing("Safra")).column("Total Pay").item(0))
```

Con esta función, no tenemos que copiar el código que convierte una cadena de pago en un número cada vez que queremos convertir una cadena de pago. Ahora simplemente llamamos a una función cuyo nombre dice exactamente lo que está haciendo.

## 3. `apply` funciones

Definir una función es muy parecido a darle un nombre a un valor con `=`. De hecho, una función es un valor como el número 1 o el texto "datos".

Por ejemplo, podemos crear un nuevo nombre para la función incorporada `max` si queremos:

```python
our_name_for_max = max
our_name_for_max(2, 6)
```

El antiguo nombre de "max" todavía existe:

```python
max(2, 6)
```

Intente simplemente escribir `max` o `our_name_for_max` (o el nombre de cualquier otra función) en una celda y ejecute esa celda. Python imprimirá una descripción (muy breve) de la función.

```python
max
```

Ahora intenta escribir `max?` o `our_name_for_max?` (o el nombre de cualquier otra función) en una celda y ejecuta esa celda. Debería aparecer un cuadro de información en la parte inferior de la pantalla con una descripción más larga de la función.

*Nota: ¡También puedes presionar Shift+Tab después de hacer clic en un nombre para ver información similar!*

```python
our_name_for_max?
```

Veamos qué sucede cuando configuramos "max" en un valor que no es función. Python ahora cree que estás intentando utilizar un número como una función, lo que provoca un error. Esté atento a las funciones a las que se les haya cambiado el nombre cuando encuentre este tipo de error.

```python
max = 6
max(2, 6)
```

```python
# Esta celda restablece max a la función incorporada. Simplemente ejecuta esta celda, no cambies su contenido.
import builtins
max = builtins.max
```

¿Por qué es esto útil? Como las funciones son sólo valores, es posible pasarlas como argumentos a otras funciones. Aquí hay un ejemplo simple pero no tan práctico: podemos crear una serie de funciones.

```python
make_array(max, np.average, are.equal_to)
```

**Pregunta 3.1.** Crea una matriz que contenga otras 3 funciones cualesquiera que hayas visto. Llámelo "some_functions".

```python
some_functions = ...
some_functions
```

```python
grader.check("q31")
```

Trabajar con funciones como valores puede generar un código de aspecto extraño. Por ejemplo, vea si puede descubrir por qué funciona el siguiente código.

```python
make_array(max, np.average, are.equal_to).item(0)(4, -2, 7)
```

Un ejemplo más útil de pasar funciones a otras funciones como argumentos es el método de tabla "aplicar".

`apply` llama a una función muchas veces, una vez en *cada* elemento en una columna de una tabla. Produce una *matriz* de los resultados. Aquí usamos "aplicar" para convertir el pago de cada CEO a un número, usando la función que definiste:

*Nota:* Verá una serie de números como `5.325e+07`. Esta es la forma en que Python representa la **notación científica**. Interpretamos "5.325e+07" como "5.325 * 10**7", o "53.250.000".

```python
raw_compensation.apply(convert_pay_string_to_number, "Total Pay")
```

Aquí hay una ilustración de lo que eso hizo:

<img src="https://github.com/darioblandon/materials-fa23/blob/main/lab/lab04/apply.png?raw=1"/>

Tenga en cuenta que no escribimos `raw_compensation.apply(convert_pay_string_to_number(), “Total Pay”)` o `raw_compensation.apply(convert_pay_string_to_number(“Total Pay”))`. Simplemente pasamos el nombre de la función, sin paréntesis, a "aplicar", porque todo lo que queremos hacer es que "aplicar" sepa el nombre de la función que nos gustaría usar y el nombre de la columna que Me gusta usarlo. ¡`apply` luego llamará a la función `convert_pay_string_to_number` en cada valor de la columna por nosotros! También tenga en cuenta que llamar a `tbl.apply` ___no___ altera la tabla original de ninguna manera.



**Pregunta 3.2.** Usando `apply`, cree una tabla que sea una copia de `raw_compensation` con una columna adicional llamada `Total Pay ($)`. Esa columna debe contener el resultado de aplicar `convert_pay_string_to_number` a la columna `Total Pay` (como hicimos anteriormente). Llame a la nueva tabla "compensación".


```python
compensation = raw_compensation.with_column(
    "Total Pay ($)",
    ...
    )
compensation
```

```python
grader.check("q32")
```

Ahora que tenemos todos los pagos como números, podemos aprender más sobre ellos mediante el cálculo.

**Pregunta 3.3.** Calcule el salario total promedio de los directores ejecutivos en el conjunto de datos.


```python
average_total_pay = ...
average_total_pay
```

```python
grader.check("q33")
```

**Pregunta 3.4** Las empresas pagan a los ejecutivos de diversas formas: en efectivo, otorgándoles acciones u otro capital de la empresa, o con beneficios auxiliares (como jets privados). Calcule la proporción del pago de cada director ejecutivo que fue en efectivo. (Su respuesta debería ser una serie de números, uno para cada director ejecutivo en el conjunto de datos).

*Pista:* ¿Qué función has definido que puede convertir una cadena en un número?


```python
cash_proportion = ...
cash_proportion
```

```python
grader.check("q34")
```

**Why is `apply` useful?**

For operations like arithmetic, or the functions in the NumPy library, you don't need to use `apply`, because they automatically work on each element of an array.  But there are many things that don't.  The string manipulation we did in today's lab is one example.  Since you can write any code you want in a function, `apply` gives you greater control over how you operate on data.

Consulte la columna "% de cambio" en "compensación". Muestra el aumento porcentual en la remuneración del director ejecutivo con respecto al año anterior. Para los directores ejecutivos que no tienen ningún año anterior registrado, en su lugar dice "(Sin año anterior)". Los valores en esta columna son *cadenas*, no números, por lo que, al igual que la columna "Pago total", no se puede utilizar sin un poco de trabajo adicional.

Dado su salario actual y el aumento porcentual del año anterior, puede calcular el salario del año anterior. Por ejemplo, si su salario es $\$120$ este año, y eso es un aumento del 50 % respecto al año anterior, entonces su salario del año anterior fue $\frac{\$120}{1 + \frac{50}{100} }$ o \$80.

**Pregunta 3.5** Cree una nueva tabla llamada `with_previous_compensation`. Debería ser una copia de "compensación", pero con los directores ejecutivos "(sin año anterior)" filtrados y con una columna adicional llamada "Pago total de 2014 ($)". Esa columna debería tener la remuneración de cada director ejecutivo en 2014.

*Pista 1:* Puedes imprimir tus resultados después de cada paso para asegurarte de que estás en el camino correcto.

*Pista 2:* Hemos proporcionado una estructura que puedes utilizar para llegar a la respuesta. Sin embargo, si le resulta confuso, no dude en eliminar la estructura actual y abordar el problema a su manera.


```python
# Definition to turn percent to number
def percent_string_to_num(percent_string):
    """Converts a percentage string to a number."""
    return ...

# Compensation table where there is a previous year
having_previous_year = ...

# Get the percent changes as numbers instead of strings
# We're still working off the table having_previous_year
percent_changes = ...

# Calculate the previous year's pay
# We're still working off the table having_previous_year
previous_pay = ...

# Put the previous pay column into the having_previous_year table
with_previous_compensation = ...

with_previous_compensation
```

```python
grader.check("q35")
```

**Pregunta 3.6** Determinar la remuneración media en 2014 de los CEOs que aparecen en la tabla `with_previous_compensationr`. Asigne este valor a la variable `average_pay_2014`.

```python
average_pay_2014 = ...
average_pay_2014
```

```python
grader.check("q36")
```

## 4. Histogramas
Anteriormente, calculamos el salario promedio entre los directores ejecutivos en nuestro conjunto de datos de 102 directores ejecutivos. Sin embargo, el promedio no nos dice todo sobre las cantidades que se les paga a los directores ejecutivos. Quizás sólo unos pocos directores ejecutivos ganen la mayor parte del dinero, incluso entre estos 102.

Podemos usar un método de *histograma* para mostrar la *distribución* de un conjunto de números. El método de tabla `hist` toma un único argumento, el nombre de una columna de números. Produce un histograma de los números en esa columna.

**Pregunta 4.1.** Haga un histograma del salario total de los CEO en `remuneración`. Consulte con un compañero o instructor para asegurarse de tener la trama correcta. *Si recibes una advertencia, ignórala.*

*Sugerencia:* Si no está seguro de cómo crear un histograma, consulte la [hoja de referencia de Python](https://www.data8.org/fa23/reference/).

```python
...
```

**Pregunta 4.2.** ¿Cuántos directores ejecutivos ganaron más de $30 millones en salario total? Encuentre el valor usando el código, luego verifique que el valor que encontró sea consistente con lo que ve en el histograma.

```python
num_ceos_more_than_30_million_2 = ...
num_ceos_more_than_30_million_2
```

```python
grader.check("q42")
```

<img src="https://github.com/darioblandon/materials-fa23/blob/main/lab/lab04/toby.png?raw=1" alt="Photo of small dog being pet" width="30%"/>

**Toby** está emocionado de que el sol vuelva a brillar. Ahora que has terminado el laboratorio 4, ¡puedes relajarte en el Claro con él!

## 5. Formulario de socio del proyecto 1

¡El Proyecto 1 se lanzará **este viernes** (27/9)! Tiene la opción de trabajar con un socio que esté inscrito en su laboratorio. Su GSI le enviará un formulario para conectarlo con un socio para este proyecto. También puedes indicar si trabajas solo o ya has encontrado pareja y no necesitas estar emparejado. Este formulario es **obligatorio**; complételo antes de enviar su laboratorio. Establezca "enviado" en "Verdadero" una vez que haya enviado el formulario.

Nota: Si está completando esta práctica de laboratorio antes de tiempo, es posible que el formulario aún no se haya enviado. Establezca "enviado" en "Verdadero" por ahora y esté atento a un correo electrónico de su GSI a finales de esta semana.

```python
submitted = ...
```

```python
grader.check("q5")
```


---

¡Terminaste con el laboratorio!

**Información importante sobre el envío:**
- **Ejecuta todas las pruebas** y verifica que todas pasan
- **Guardar** desde el menú **Archivo**
- **Ejecute la celda final para generar el archivo zip**
- **Haga clic en el enlace para descargar el archivo zip**
- Luego, ve a [Gradescope](https://www.gradescope.com/courses/576157) y envía el archivo zip a la tarea correspondiente. El nombre de esta tarea es "Lab XX Autograder", donde XX es el número de laboratorio: 01, 02, 03, etc.

- Si terminas temprano en el laboratorio regular, **pídele a uno de los miembros del personal que te marque**.

**Es su responsabilidad asegurarse de que su trabajo esté guardado antes de ejecutar la última celda.**

## Envío

Asegúrese de haber ejecutado todas las celdas de su cuaderno en orden antes de ejecutar la siguiente celda, para que todas las imágenes/gráficos aparezcan en el resultado. La siguiente celda generará un archivo zip para que lo envíe. **¡Guarde antes de exportar!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q11 (points: None)
```
>>> # Make sure your function has the proper syntax!
>>> to_percentage(.35) == 35.0
True
```
```
>>> # Make sure you are passing in the proper argument!
>>> twenty_percent == 20.0
True
```


Tests for q12 (points: None)
```
>>> # Make sure you're using a_proportion!
>>> np.isclose(a_percentage, 70.71067811865476)
True
```


Tests for q13 (points: None)
```
>>> num_non_vowels("Go bears!") == 6
True
```


Tests for q14 (points: None)
```
>>> print_kth_top_movie_year(4)
Year number 4 for total gross movie sales was: 2009
```


Tests for q21 (points: None)
```
>>> # Make sure you are examining the values in the column, not the column itself
>>> import numpy
>>> total_pay_type != numpy.ndarray
True
```
```
>>> 'str' in str(total_pay_type)
True
```
```
>>> isinstance(total_pay_type, type)
True
```
```
>>> # Make sure to call the type function on a value in the column
>>> total_pay_type != int
True
```


Tests for q22 (points: None)
```
>>> str(mark_hurd_pay_string)
'$53.25 '
```


Tests for q23 (points: None)
```
>>> # Your answer should be a number
>>> type(mark_hurd_pay) != str
True
```
```
>>> # Don't forget to give your answer in dollars, not millions of 
>>> # Dollars! 
>>> mark_hurd_pay != 5325
True
```
```
>>> # Don't forget to give your answer in dollars, not millions of 
>>> # Dollars! 
>>> mark_hurd_pay == 53250000
True
```


Tests for q24 (points: None)
```
>>> convert_pay_string_to_number("$100 ") == 100000000.0
True
```
```
>>> convert_pay_string_to_number("$23 ") == 23000000.0
True
```


Tests for q31 (points: None)
```
>>> len(some_functions)
3
```
```
>>> # The first thing in your array may not be a function
>>> callable(some_functions.item(0))
True
```
```
>>> # The second thing in your array may not be a function
>>> callable(some_functions.item(1))
True
```
```
>>> # The third thing in your array may not be a function
>>> callable(some_functions.item(2))
True
```


Tests for q32 (points: None)
```
>>> # You either didn't add the 'Total Pay ($)' column, 
>>> # or you mislabeled it
>>> 'Total Pay ($)' in compensation.labels
True
```
```
>>> # You have the column in your table, but the values may be wrong
>>> t = compensation.sort('Total Pay ($)', descending = True)
>>> np.isclose(t.column('Total Pay ($)').item(0), 53250000.0)
True
```


Tests for q33 (points: None)
```
>>> import math
>>> math.isclose(average_total_pay, 11558613.861386139, rel_tol = 0.1)
True
```


Tests for q34 (points: None)
```
>>> len(cash_proportion) == 101
True
```
```
>>> import math
>>> math.isclose(cash_proportion.item(0), 0.01784038, rel_tol = .001)
True
```


Tests for q35 (points: None)
```
>>> # Make sure to remove the "(No previous year)" CEOs 
>>> "(No previous year)" not in with_previous_compensation.column("% Change")
True
```
```
>>> import math
>>> # You have the column, but some of your values may be wrong
>>> t = with_previous_compensation.sort("2014 Total Pay ($)", descending = True)
>>> value = t.column("2014 Total Pay ($)").item(0)
>>> math.isclose(value, 67700000.0, rel_tol = 0.01)
True
```
```
>>> # You have the column, but your number of rows is off
>>> with_previous_compensation.num_rows == 80
True
```


Tests for q36 (points: None)
```
>>> import math
>>> math.isclose(average_pay_2014, 11794790.817048479, rel_tol = 0.01)
True
```


Tests for q42 (points: None)
```
>>> num_ceos_more_than_30_million_2 == 5
True
```


Tests for q5 (points: None)
```
>>> submitted
True
```
